unsat
(assume a0 (and (and (and (and (and (and (and (and (and (and (or (or (or (or (or (! (= (! (op e0 e0) :named @p_246) e0) :named @p_252) (! (= @p_246 e1) :named @p_251)) (! (= @p_246 e2) :named @p_250)) (! (= @p_246 e3) :named @p_249)) (! (= @p_246 e4) :named @p_248)) (! (= @p_246 e5) :named @p_247)) (or (or (or (or (or (! (= (! (op e0 e1) :named @p_239) e0) :named @p_245) (! (= @p_239 e1) :named @p_244)) (! (= @p_239 e2) :named @p_243)) (! (= @p_239 e3) :named @p_242)) (! (= @p_239 e4) :named @p_241)) (! (= @p_239 e5) :named @p_240))) (or (or (or (or (or (! (= (! (op e0 e2) :named @p_232) e0) :named @p_238) (! (= @p_232 e1) :named @p_237)) (! (= @p_232 e2) :named @p_236)) (! (= @p_232 e3) :named @p_235)) (! (= @p_232 e4) :named @p_234)) (! (= @p_232 e5) :named @p_233))) (or (or (or (or (or (! (= (! (op e0 e3) :named @p_225) e0) :named @p_231) (! (= @p_225 e1) :named @p_230)) (! (= @p_225 e2) :named @p_229)) (! (= @p_225 e3) :named @p_228)) (! (= @p_225 e4) :named @p_227)) (! (= @p_225 e5) :named @p_226))) (or (or (or (or (or (! (= (! (op e0 e4) :named @p_218) e0) :named @p_224) (! (= @p_218 e1) :named @p_223)) (! (= @p_218 e2) :named @p_222)) (! (= @p_218 e3) :named @p_221)) (! (= @p_218 e4) :named @p_220)) (! (= @p_218 e5) :named @p_219))) (or (or (or (or (or (! (= (! (op e0 e5) :named @p_211) e0) :named @p_217) (! (= @p_211 e1) :named @p_216)) (! (= @p_211 e2) :named @p_215)) (! (= @p_211 e3) :named @p_214)) (! (= @p_211 e4) :named @p_213)) (! (= @p_211 e5) :named @p_212))) (and (and (and (and (and (or (or (or (or (or (! (= (! (op e1 e0) :named @p_204) e0) :named @p_210) (! (= @p_204 e1) :named @p_209)) (! (= @p_204 e2) :named @p_208)) (! (= @p_204 e3) :named @p_207)) (! (= @p_204 e4) :named @p_206)) (! (= @p_204 e5) :named @p_205)) (or (or (or (or (or (! (= (! (op e1 e1) :named @p_197) e0) :named @p_203) (! (= @p_197 e1) :named @p_202)) (! (= @p_197 e2) :named @p_201)) (! (= @p_197 e3) :named @p_200)) (! (= @p_197 e4) :named @p_199)) (! (= @p_197 e5) :named @p_198))) (or (or (or (or (or (! (= (! (op e1 e2) :named @p_190) e0) :named @p_196) (! (= @p_190 e1) :named @p_195)) (! (= @p_190 e2) :named @p_194)) (! (= @p_190 e3) :named @p_193)) (! (= @p_190 e4) :named @p_192)) (! (= @p_190 e5) :named @p_191))) (or (or (or (or (or (! (= (! (op e1 e3) :named @p_183) e0) :named @p_189) (! (= @p_183 e1) :named @p_188)) (! (= @p_183 e2) :named @p_187)) (! (= @p_183 e3) :named @p_186)) (! (= @p_183 e4) :named @p_185)) (! (= @p_183 e5) :named @p_184))) (or (or (or (or (or (! (= (! (op e1 e4) :named @p_176) e0) :named @p_182) (! (= @p_176 e1) :named @p_181)) (! (= @p_176 e2) :named @p_180)) (! (= @p_176 e3) :named @p_179)) (! (= @p_176 e4) :named @p_178)) (! (= @p_176 e5) :named @p_177))) (or (or (or (or (or (! (= (! (op e1 e5) :named @p_169) e0) :named @p_175) (! (= @p_169 e1) :named @p_174)) (! (= @p_169 e2) :named @p_173)) (! (= @p_169 e3) :named @p_172)) (! (= @p_169 e4) :named @p_171)) (! (= @p_169 e5) :named @p_170)))) (and (and (and (and (and (or (or (or (or (or (! (= (! (op e2 e0) :named @p_162) e0) :named @p_168) (! (= @p_162 e1) :named @p_167)) (! (= @p_162 e2) :named @p_166)) (! (= @p_162 e3) :named @p_165)) (! (= @p_162 e4) :named @p_164)) (! (= @p_162 e5) :named @p_163)) (or (or (or (or (or (! (= (! (op e2 e1) :named @p_155) e0) :named @p_161) (! (= @p_155 e1) :named @p_160)) (! (= @p_155 e2) :named @p_159)) (! (= @p_155 e3) :named @p_158)) (! (= @p_155 e4) :named @p_157)) (! (= @p_155 e5) :named @p_156))) (or (or (or (or (or (! (= (! (op e2 e2) :named @p_148) e0) :named @p_154) (! (= @p_148 e1) :named @p_153)) (! (= @p_148 e2) :named @p_152)) (! (= @p_148 e3) :named @p_151)) (! (= @p_148 e4) :named @p_150)) (! (= @p_148 e5) :named @p_149))) (or (or (or (or (or (! (= (! (op e2 e3) :named @p_141) e0) :named @p_147) (! (= @p_141 e1) :named @p_146)) (! (= @p_141 e2) :named @p_145)) (! (= @p_141 e3) :named @p_144)) (! (= @p_141 e4) :named @p_143)) (! (= @p_141 e5) :named @p_142))) (or (or (or (or (or (! (= (! (op e2 e4) :named @p_134) e0) :named @p_140) (! (= @p_134 e1) :named @p_139)) (! (= @p_134 e2) :named @p_138)) (! (= @p_134 e3) :named @p_137)) (! (= @p_134 e4) :named @p_136)) (! (= @p_134 e5) :named @p_135))) (or (or (or (or (or (! (= (! (op e2 e5) :named @p_127) e0) :named @p_133) (! (= @p_127 e1) :named @p_132)) (! (= @p_127 e2) :named @p_131)) (! (= @p_127 e3) :named @p_130)) (! (= @p_127 e4) :named @p_129)) (! (= @p_127 e5) :named @p_128)))) (and (and (and (and (and (or (or (or (or (or (! (= (! (op e3 e0) :named @p_120) e0) :named @p_126) (! (= @p_120 e1) :named @p_125)) (! (= @p_120 e2) :named @p_124)) (! (= @p_120 e3) :named @p_123)) (! (= @p_120 e4) :named @p_122)) (! (= @p_120 e5) :named @p_121)) (or (or (or (or (or (! (= (! (op e3 e1) :named @p_113) e0) :named @p_119) (! (= @p_113 e1) :named @p_118)) (! (= @p_113 e2) :named @p_117)) (! (= @p_113 e3) :named @p_116)) (! (= @p_113 e4) :named @p_115)) (! (= @p_113 e5) :named @p_114))) (or (or (or (or (or (! (= (! (op e3 e2) :named @p_106) e0) :named @p_112) (! (= @p_106 e1) :named @p_111)) (! (= @p_106 e2) :named @p_110)) (! (= @p_106 e3) :named @p_109)) (! (= @p_106 e4) :named @p_108)) (! (= @p_106 e5) :named @p_107))) (or (or (or (or (or (! (= (! (op e3 e3) :named @p_99) e0) :named @p_105) (! (= @p_99 e1) :named @p_104)) (! (= @p_99 e2) :named @p_103)) (! (= @p_99 e3) :named @p_102)) (! (= @p_99 e4) :named @p_101)) (! (= @p_99 e5) :named @p_100))) (or (or (or (or (or (! (= (! (op e3 e4) :named @p_92) e0) :named @p_98) (! (= @p_92 e1) :named @p_97)) (! (= @p_92 e2) :named @p_96)) (! (= @p_92 e3) :named @p_95)) (! (= @p_92 e4) :named @p_94)) (! (= @p_92 e5) :named @p_93))) (or (or (or (or (or (! (= (! (op e3 e5) :named @p_85) e0) :named @p_91) (! (= @p_85 e1) :named @p_90)) (! (= @p_85 e2) :named @p_89)) (! (= @p_85 e3) :named @p_88)) (! (= @p_85 e4) :named @p_87)) (! (= @p_85 e5) :named @p_86)))) (and (and (and (and (and (or (or (or (or (or (! (= (! (op e4 e0) :named @p_78) e0) :named @p_84) (! (= @p_78 e1) :named @p_83)) (! (= @p_78 e2) :named @p_82)) (! (= @p_78 e3) :named @p_81)) (! (= @p_78 e4) :named @p_80)) (! (= @p_78 e5) :named @p_79)) (or (or (or (or (or (! (= (! (op e4 e1) :named @p_71) e0) :named @p_77) (! (= @p_71 e1) :named @p_76)) (! (= @p_71 e2) :named @p_75)) (! (= @p_71 e3) :named @p_74)) (! (= @p_71 e4) :named @p_73)) (! (= @p_71 e5) :named @p_72))) (or (or (or (or (or (! (= (! (op e4 e2) :named @p_64) e0) :named @p_70) (! (= @p_64 e1) :named @p_69)) (! (= @p_64 e2) :named @p_68)) (! (= @p_64 e3) :named @p_67)) (! (= @p_64 e4) :named @p_66)) (! (= @p_64 e5) :named @p_65))) (or (or (or (or (or (! (= (! (op e4 e3) :named @p_57) e0) :named @p_63) (! (= @p_57 e1) :named @p_62)) (! (= @p_57 e2) :named @p_61)) (! (= @p_57 e3) :named @p_60)) (! (= @p_57 e4) :named @p_59)) (! (= @p_57 e5) :named @p_58))) (or (or (or (or (or (! (= (! (op e4 e4) :named @p_50) e0) :named @p_56) (! (= @p_50 e1) :named @p_55)) (! (= @p_50 e2) :named @p_54)) (! (= @p_50 e3) :named @p_53)) (! (= @p_50 e4) :named @p_52)) (! (= @p_50 e5) :named @p_51))) (or (or (or (or (or (! (= (! (op e4 e5) :named @p_43) e0) :named @p_49) (! (= @p_43 e1) :named @p_48)) (! (= @p_43 e2) :named @p_47)) (! (= @p_43 e3) :named @p_46)) (! (= @p_43 e4) :named @p_45)) (! (= @p_43 e5) :named @p_44)))) (and (and (and (and (and (or (or (or (or (or (! (= (! (op e5 e0) :named @p_36) e0) :named @p_42) (! (= @p_36 e1) :named @p_41)) (! (= @p_36 e2) :named @p_40)) (! (= @p_36 e3) :named @p_39)) (! (= @p_36 e4) :named @p_38)) (! (= @p_36 e5) :named @p_37)) (or (or (or (or (or (! (= (! (op e5 e1) :named @p_29) e0) :named @p_35) (! (= @p_29 e1) :named @p_34)) (! (= @p_29 e2) :named @p_33)) (! (= @p_29 e3) :named @p_32)) (! (= @p_29 e4) :named @p_31)) (! (= @p_29 e5) :named @p_30))) (or (or (or (or (or (! (= (! (op e5 e2) :named @p_22) e0) :named @p_28) (! (= @p_22 e1) :named @p_27)) (! (= @p_22 e2) :named @p_26)) (! (= @p_22 e3) :named @p_25)) (! (= @p_22 e4) :named @p_24)) (! (= @p_22 e5) :named @p_23))) (or (or (or (or (or (! (= (! (op e5 e3) :named @p_15) e0) :named @p_21) (! (= @p_15 e1) :named @p_20)) (! (= @p_15 e2) :named @p_19)) (! (= @p_15 e3) :named @p_18)) (! (= @p_15 e4) :named @p_17)) (! (= @p_15 e5) :named @p_16))) (or (or (or (or (or (! (= (! (op e5 e4) :named @p_8) e0) :named @p_14) (! (= @p_8 e1) :named @p_13)) (! (= @p_8 e2) :named @p_12)) (! (= @p_8 e3) :named @p_11)) (! (= @p_8 e4) :named @p_10)) (! (= @p_8 e5) :named @p_9))) (or (or (or (or (or (! (= (! (op e5 e5) :named @p_1) e0) :named @p_7) (! (= @p_1 e1) :named @p_6)) (! (= @p_1 e2) :named @p_5)) (! (= @p_1 e3) :named @p_4)) (! (= @p_1 e4) :named @p_3)) (! (= @p_1 e5) :named @p_2)))))
(assume a1 (and (and (and (and (and (and (and (and (and (and (and (or (or (or (or (or @p_252 @p_245) @p_238) @p_231) @p_224) @p_217) (or (or (or (or (or @p_252 @p_210) @p_168) @p_126) @p_84) @p_42)) (and (or (or (or (or (or @p_251 @p_244) @p_237) @p_230) @p_223) @p_216) (or (or (or (or (or @p_251 @p_209) @p_167) @p_125) @p_83) @p_41))) (and (or (or (or (or (or @p_250 @p_243) @p_236) @p_229) @p_222) @p_215) (or (or (or (or (or @p_250 @p_208) @p_166) @p_124) @p_82) @p_40))) (and (or (or (or (or (or @p_249 @p_242) @p_235) @p_228) @p_221) @p_214) (or (or (or (or (or @p_249 @p_207) @p_165) @p_123) @p_81) @p_39))) (and (or (or (or (or (or @p_248 @p_241) @p_234) @p_227) @p_220) @p_213) (or (or (or (or (or @p_248 @p_206) @p_164) @p_122) @p_80) @p_38))) (and (or (or (or (or (or @p_247 @p_240) @p_233) @p_226) @p_219) @p_212) (or (or (or (or (or @p_247 @p_205) @p_163) @p_121) @p_79) @p_37))) (and (and (and (and (and (and (or (or (or (or (or @p_210 @p_203) @p_196) @p_189) @p_182) @p_175) (or (or (or (or (or @p_245 @p_203) @p_161) @p_119) @p_77) @p_35)) (and (or (or (or (or (or @p_209 @p_202) @p_195) @p_188) @p_181) @p_174) (or (or (or (or (or @p_244 @p_202) @p_160) @p_118) @p_76) @p_34))) (and (or (or (or (or (or @p_208 @p_201) @p_194) @p_187) @p_180) @p_173) (or (or (or (or (or @p_243 @p_201) @p_159) @p_117) @p_75) @p_33))) (and (or (or (or (or (or @p_207 @p_200) @p_193) @p_186) @p_179) @p_172) (or (or (or (or (or @p_242 @p_200) @p_158) @p_116) @p_74) @p_32))) (and (or (or (or (or (or @p_206 @p_199) @p_192) @p_185) @p_178) @p_171) (or (or (or (or (or @p_241 @p_199) @p_157) @p_115) @p_73) @p_31))) (and (or (or (or (or (or @p_205 @p_198) @p_191) @p_184) @p_177) @p_170) (or (or (or (or (or @p_240 @p_198) @p_156) @p_114) @p_72) @p_30)))) (and (and (and (and (and (and (or (or (or (or (or @p_168 @p_161) @p_154) @p_147) @p_140) @p_133) (or (or (or (or (or @p_238 @p_196) @p_154) @p_112) @p_70) @p_28)) (and (or (or (or (or (or @p_167 @p_160) @p_153) @p_146) @p_139) @p_132) (or (or (or (or (or @p_237 @p_195) @p_153) @p_111) @p_69) @p_27))) (and (or (or (or (or (or @p_166 @p_159) @p_152) @p_145) @p_138) @p_131) (or (or (or (or (or @p_236 @p_194) @p_152) @p_110) @p_68) @p_26))) (and (or (or (or (or (or @p_165 @p_158) @p_151) @p_144) @p_137) @p_130) (or (or (or (or (or @p_235 @p_193) @p_151) @p_109) @p_67) @p_25))) (and (or (or (or (or (or @p_164 @p_157) @p_150) @p_143) @p_136) @p_129) (or (or (or (or (or @p_234 @p_192) @p_150) @p_108) @p_66) @p_24))) (and (or (or (or (or (or @p_163 @p_156) @p_149) @p_142) @p_135) @p_128) (or (or (or (or (or @p_233 @p_191) @p_149) @p_107) @p_65) @p_23)))) (and (and (and (and (and (and (or (or (or (or (or @p_126 @p_119) @p_112) @p_105) @p_98) @p_91) (or (or (or (or (or @p_231 @p_189) @p_147) @p_105) @p_63) @p_21)) (and (or (or (or (or (or @p_125 @p_118) @p_111) @p_104) @p_97) @p_90) (or (or (or (or (or @p_230 @p_188) @p_146) @p_104) @p_62) @p_20))) (and (or (or (or (or (or @p_124 @p_117) @p_110) @p_103) @p_96) @p_89) (or (or (or (or (or @p_229 @p_187) @p_145) @p_103) @p_61) @p_19))) (and (or (or (or (or (or @p_123 @p_116) @p_109) @p_102) @p_95) @p_88) (or (or (or (or (or @p_228 @p_186) @p_144) @p_102) @p_60) @p_18))) (and (or (or (or (or (or @p_122 @p_115) @p_108) @p_101) @p_94) @p_87) (or (or (or (or (or @p_227 @p_185) @p_143) @p_101) @p_59) @p_17))) (and (or (or (or (or (or @p_121 @p_114) @p_107) @p_100) @p_93) @p_86) (or (or (or (or (or @p_226 @p_184) @p_142) @p_100) @p_58) @p_16)))) (and (and (and (and (and (and (or (or (or (or (or @p_84 @p_77) @p_70) @p_63) @p_56) @p_49) (or (or (or (or (or @p_224 @p_182) @p_140) @p_98) @p_56) @p_14)) (and (or (or (or (or (or @p_83 @p_76) @p_69) @p_62) @p_55) @p_48) (or (or (or (or (or @p_223 @p_181) @p_139) @p_97) @p_55) @p_13))) (and (or (or (or (or (or @p_82 @p_75) @p_68) @p_61) @p_54) @p_47) (or (or (or (or (or @p_222 @p_180) @p_138) @p_96) @p_54) @p_12))) (and (or (or (or (or (or @p_81 @p_74) @p_67) @p_60) @p_53) @p_46) (or (or (or (or (or @p_221 @p_179) @p_137) @p_95) @p_53) @p_11))) (and (or (or (or (or (or @p_80 @p_73) @p_66) @p_59) @p_52) @p_45) (or (or (or (or (or @p_220 @p_178) @p_136) @p_94) @p_52) @p_10))) (and (or (or (or (or (or @p_79 @p_72) @p_65) @p_58) @p_51) @p_44) (or (or (or (or (or @p_219 @p_177) @p_135) @p_93) @p_51) @p_9)))) (and (and (and (and (and (and (or (or (or (or (or @p_42 @p_35) @p_28) @p_21) @p_14) @p_7) (or (or (or (or (or @p_217 @p_175) @p_133) @p_91) @p_49) @p_7)) (and (or (or (or (or (or @p_41 @p_34) @p_27) @p_20) @p_13) @p_6) (or (or (or (or (or @p_216 @p_174) @p_132) @p_90) @p_48) @p_6))) (and (or (or (or (or (or @p_40 @p_33) @p_26) @p_19) @p_12) @p_5) (or (or (or (or (or @p_215 @p_173) @p_131) @p_89) @p_47) @p_5))) (and (or (or (or (or (or @p_39 @p_32) @p_25) @p_18) @p_11) @p_4) (or (or (or (or (or @p_214 @p_172) @p_130) @p_88) @p_46) @p_4))) (and (or (or (or (or (or @p_38 @p_31) @p_24) @p_17) @p_10) @p_3) (or (or (or (or (or @p_213 @p_171) @p_129) @p_87) @p_45) @p_3))) (and (or (or (or (or (or @p_37 @p_30) @p_23) @p_16) @p_9) @p_2) (or (or (or (or (or @p_212 @p_170) @p_128) @p_86) @p_44) @p_2)))))
(assume a2 (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (= (! (op unit e0) :named @p_274) e0) :named @p_275) (= (op e0 unit) e0)) :named @p_276) (! (and (= (op unit e1) e1) (! (= (! (op e1 unit) :named @p_271) e1) :named @p_272)) :named @p_273)) :named @p_277) (! (and (= (op unit e2) e2) (! (= (! (op e2 unit) :named @p_268) e2) :named @p_269)) :named @p_270)) :named @p_278) (! (and (= (op unit e3) e3) (! (= (! (op e3 unit) :named @p_265) e3) :named @p_266)) :named @p_267)) :named @p_279) (! (and (= (op unit e4) e4) (! (= (! (op e4 unit) :named @p_262) e4) :named @p_263)) :named @p_264)) :named @p_280) (! (and (= (op unit e5) e5) (! (= (! (op e5 unit) :named @p_259) e5) :named @p_260)) :named @p_261)) :named @p_281) (or (or (or (or (or (! (= unit e0) :named @p_258) (! (= unit e1) :named @p_257)) (! (= unit e2) :named @p_256)) (! (= unit e3) :named @p_255)) (! (= unit e4) :named @p_254)) (! (= unit e5) :named @p_253))) :named @p_282))
(assume a3 (or (or (or (or (or (or (or (or (or (or (not (! (= @p_246 @p_246) :named @p_283)) (not (= @p_204 @p_239))) (not (= @p_162 @p_232))) (not (= @p_120 @p_225))) (not (= @p_78 @p_218))) (not (= @p_36 @p_211))) (or (or (or (or (or (not (= @p_239 @p_204)) (not (= @p_197 @p_197))) (not (= @p_155 @p_190))) (not (= @p_113 @p_183))) (not (= @p_71 @p_176))) (not (= @p_29 @p_169)))) (or (or (or (or (or (not (= @p_232 @p_162)) (not (= @p_190 @p_155))) (not (= @p_148 @p_148))) (not (= @p_106 @p_141))) (not (= @p_64 @p_134))) (not (= @p_22 @p_127)))) (or (or (or (or (or (not (= @p_225 @p_120)) (not (= @p_183 @p_113))) (not (= @p_141 @p_106))) (not (= @p_99 @p_99))) (not (= @p_57 @p_92))) (not (= @p_15 @p_85)))) (or (or (or (or (or (not (= @p_218 @p_78)) (not (= @p_176 @p_71))) (not (= @p_134 @p_64))) (not (= @p_92 @p_57))) (not (= @p_50 @p_50))) (not (= @p_8 @p_43)))) (or (or (or (or (or (not (= @p_211 @p_36)) (not (= @p_169 @p_29))) (not (= @p_127 @p_22))) (not (= @p_85 @p_15))) (not (= @p_43 @p_8))) (not (= @p_1 @p_1)))))
(assume a4 (or (or (or (or (or (and (and (and (and (and (not (= (op @p_246 e0) e0)) (not (= (op @p_239 e0) e1))) (not (= (op @p_232 e0) e2))) (not (= (op @p_225 e0) e3))) (not (= (op @p_218 e0) e4))) (not (= (op @p_211 e0) e5))) (and (and (and (and (and (not (= (op @p_204 e1) e0)) (not (= (op @p_197 e1) e1))) (not (= (op @p_190 e1) e2))) (not (= (op @p_183 e1) e3))) (not (= (op @p_176 e1) e4))) (not (= (op @p_169 e1) e5)))) (and (and (and (and (and (not (= (op @p_162 e2) e0)) (not (= (op @p_155 e2) e1))) (not (= (op @p_148 e2) e2))) (not (= (op @p_141 e2) e3))) (not (= (op @p_134 e2) e4))) (not (= (op @p_127 e2) e5)))) (and (and (and (and (and (not (= (op @p_120 e3) e0)) (not (= (op @p_113 e3) e1))) (not (= (op @p_106 e3) e2))) (not (= (op @p_99 e3) e3))) (not (= (op @p_92 e3) e4))) (not (= (op @p_85 e3) e5)))) (and (and (and (and (and (not (= (op @p_78 e4) e0)) (not (= (op @p_71 e4) e1))) (not (= (op @p_64 e4) e2))) (not (= (op @p_57 e4) e3))) (not (= (op @p_50 e4) e4))) (not (= (op @p_43 e4) e5)))) (and (and (and (and (and (not (= (op @p_36 e5) e0)) (not (= (op @p_29 e5) e1))) (not (= (op @p_22 e5) e2))) (not (= (op @p_15 e5) e3))) (not (= (op @p_8 e5) e4))) (not (= (op @p_1 e5) e5)))))
(assume a5 (and (and (and (and (and (or (or (or (or (or (! (= (op @p_246 @p_246) e0) :named @p_319) (= (! (op @p_239 @p_239) :named @p_318) e1)) (= (! (op @p_232 @p_232) :named @p_317) e2)) (= (! (op @p_225 @p_225) :named @p_316) e3)) (= (! (op @p_218 @p_218) :named @p_315) e4)) (= (! (op @p_211 @p_211) :named @p_314) e5)) (or (or (or (or (or (= (! (op @p_204 @p_204) :named @p_313) e0) (! (= (op @p_197 @p_197) e1) :named @p_312)) (= (! (op @p_190 @p_190) :named @p_311) e2)) (= (! (op @p_183 @p_183) :named @p_310) e3)) (= (! (op @p_176 @p_176) :named @p_309) e4)) (= (! (op @p_169 @p_169) :named @p_308) e5))) (or (or (or (or (or (= (! (op @p_162 @p_162) :named @p_307) e0) (= (! (op @p_155 @p_155) :named @p_306) e1)) (! (= (op @p_148 @p_148) e2) :named @p_305)) (= (! (op @p_141 @p_141) :named @p_304) e3)) (= (! (op @p_134 @p_134) :named @p_303) e4)) (= (! (op @p_127 @p_127) :named @p_302) e5))) (or (or (or (or (or (= (! (op @p_120 @p_120) :named @p_301) e0) (= (! (op @p_113 @p_113) :named @p_300) e1)) (= (! (op @p_106 @p_106) :named @p_299) e2)) (! (= (op @p_99 @p_99) e3) :named @p_298)) (= (! (op @p_92 @p_92) :named @p_297) e4)) (= (! (op @p_85 @p_85) :named @p_296) e5))) (or (or (or (or (or (= (! (op @p_78 @p_78) :named @p_295) e0) (= (! (op @p_71 @p_71) :named @p_294) e1)) (= (! (op @p_64 @p_64) :named @p_293) e2)) (= (! (op @p_57 @p_57) :named @p_292) e3)) (! (= (op @p_50 @p_50) e4) :named @p_291)) (= (! (op @p_43 @p_43) :named @p_290) e5))) (or (or (or (or (or (= (! (op @p_36 @p_36) :named @p_289) e0) (= (! (op @p_29 @p_29) :named @p_288) e1)) (= (! (op @p_22 @p_22) :named @p_287) e2)) (= (! (op @p_15 @p_15) :named @p_286) e3)) (= (! (op @p_8 @p_8) :named @p_285) e4)) (! (= (op @p_1 @p_1) e5) :named @p_284))))
(assume a6 (or (or (or (or (or (and (and (and (and (and (not @p_319) (not (= @p_313 e1))) (not (= @p_307 e2))) (not (= @p_301 e3))) (not (= @p_295 e4))) (not (= @p_289 e5))) (and (and (and (and (and (not (= @p_318 e0)) (not @p_312)) (not (= @p_306 e2))) (not (= @p_300 e3))) (not (= @p_294 e4))) (not (= @p_288 e5)))) (and (and (and (and (and (not (= @p_317 e0)) (not (= @p_311 e1))) (not @p_305)) (not (= @p_299 e3))) (not (= @p_293 e4))) (not (= @p_287 e5)))) (and (and (and (and (and (not (= @p_316 e0)) (not (= @p_310 e1))) (not (= @p_304 e2))) (not @p_298)) (not (= @p_292 e4))) (not (= @p_286 e5)))) (and (and (and (and (and (not (= @p_315 e0)) (not (= @p_309 e1))) (not (= @p_303 e2))) (not (= @p_297 e3))) (not @p_291)) (not (= @p_285 e5)))) (and (and (and (and (and (not (= @p_314 e0)) (not (= @p_308 e1))) (not (= @p_302 e2))) (not (= @p_296 e3))) (not (= @p_290 e4))) (not @p_284))))
(assume a7 (! (and (! (and (! (and (! (and (! (and (! (and (or (not (! (= (! (op e0 @p_246) :named @p_357) e0) :named @p_358)) (= e0 unit)) (! (or (! (not (! (= (! (op e1 @p_197) :named @p_328) e1) :named @p_329)) :named @p_379) (! (= e1 unit) :named @p_378)) :named @p_380)) :named @p_381) (! (or (! (not (! (= (! (op e2 @p_148) :named @p_326) e2) :named @p_327)) :named @p_376) (! (= e2 unit) :named @p_375)) :named @p_377)) :named @p_382) (! (or (! (not (! (= (! (op e3 @p_99) :named @p_324) e3) :named @p_325)) :named @p_373) (! (= e3 unit) :named @p_372)) :named @p_374)) :named @p_383) (! (or (! (not (! (= (! (op e4 @p_50) :named @p_322) e4) :named @p_323)) :named @p_370) (! (= e4 unit) :named @p_369)) :named @p_371)) :named @p_384) (! (or (! (not (! (= (! (op e5 @p_1) :named @p_320) e5) :named @p_321)) :named @p_367) (! (= e5 unit) :named @p_366)) :named @p_368)) :named @p_385) (! (and (! (and (! (and (! (and (! (and (! (or (or (or (or (or (! (and (! (= (! (op e0 @p_246) :named @p_357) e0) :named @p_358) (! (not (! (= @p_246 unit) :named @p_355)) :named @p_356)) :named @p_359) (! (and (! (= (! (op e1 @p_204) :named @p_352) e0) :named @p_353) (! (not (! (= @p_204 unit) :named @p_350)) :named @p_351)) :named @p_354)) (! (and (! (= (! (op e2 @p_162) :named @p_347) e0) :named @p_348) (! (not (! (= @p_162 unit) :named @p_345)) :named @p_346)) :named @p_349)) (! (and (! (= (! (op e3 @p_120) :named @p_342) e0) :named @p_343) (! (not (! (= @p_120 unit) :named @p_340)) :named @p_341)) :named @p_344)) (! (and (! (= (! (op e4 @p_78) :named @p_337) e0) :named @p_338) (! (not (! (= @p_78 unit) :named @p_335)) :named @p_336)) :named @p_339)) (! (and (! (= (! (op e5 @p_36) :named @p_332) e0) :named @p_333) (! (not (! (= @p_36 unit) :named @p_330)) :named @p_331)) :named @p_334)) :named @p_360) (or (or (or (or (or (and (= (op e0 @p_239) e1) (not (= @p_239 unit))) (and @p_329 (not (= @p_197 unit)))) (and (= (op e2 @p_155) e1) (not (= @p_155 unit)))) (and (= (op e3 @p_113) e1) (not (= @p_113 unit)))) (and (= (op e4 @p_71) e1) (not (= @p_71 unit)))) (and (= (op e5 @p_29) e1) (not (= @p_29 unit))))) :named @p_361) (or (or (or (or (or (and (= (op e0 @p_232) e2) (not (= @p_232 unit))) (and (= (op e1 @p_190) e2) (not (= @p_190 unit)))) (and @p_327 (not (= @p_148 unit)))) (and (= (op e3 @p_106) e2) (not (= @p_106 unit)))) (and (= (op e4 @p_64) e2) (not (= @p_64 unit)))) (and (= (op e5 @p_22) e2) (not (= @p_22 unit))))) :named @p_362) (or (or (or (or (or (and (= (op e0 @p_225) e3) (not (= @p_225 unit))) (and (= (op e1 @p_183) e3) (not (= @p_183 unit)))) (and (= (op e2 @p_141) e3) (not (= @p_141 unit)))) (and @p_325 (not (= @p_99 unit)))) (and (= (op e4 @p_57) e3) (not (= @p_57 unit)))) (and (= (op e5 @p_15) e3) (not (= @p_15 unit))))) :named @p_363) (or (or (or (or (or (and (= (op e0 @p_218) e4) (not (= @p_218 unit))) (and (= (op e1 @p_176) e4) (not (= @p_176 unit)))) (and (= (op e2 @p_134) e4) (not (= @p_134 unit)))) (and (= (op e3 @p_92) e4) (not (= @p_92 unit)))) (and @p_323 (not (= @p_50 unit)))) (and (= (op e5 @p_8) e4) (not (= @p_8 unit))))) :named @p_364) (or (or (or (or (or (and (= (op e0 @p_211) e5) (not (= @p_211 unit))) (and (= (op e1 @p_169) e5) (not (= @p_169 unit)))) (and (= (op e2 @p_127) e5) (not (= @p_127 unit)))) (and (= (op e3 @p_85) e5) (not (= @p_85 unit)))) (and (= (op e4 @p_43) e5) (not (= @p_43 unit)))) (and @p_321 (not (= @p_1 unit))))) :named @p_365)) :named @p_386))
(assume a8 @p_258)
(assume a9 (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= @p_246 @p_204)) (not (= @p_246 @p_162))) (not (= @p_246 @p_120))) (not (= @p_246 @p_78))) (not (= @p_246 @p_36))) (not (= @p_204 @p_162))) (not (= @p_204 @p_120))) (not (= @p_204 @p_78))) (not (= @p_204 @p_36))) (not (= @p_162 @p_120))) (not (= @p_162 @p_78))) (not (= @p_162 @p_36))) (not (= @p_120 @p_78))) (not (= @p_120 @p_36))) (not (= @p_78 @p_36))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= @p_239 @p_197)) (not (= @p_239 @p_155))) (not (= @p_239 @p_113))) (not (= @p_239 @p_71))) (not (= @p_239 @p_29))) (not (= @p_197 @p_155))) (not (= @p_197 @p_113))) (not (= @p_197 @p_71))) (not (= @p_197 @p_29))) (not (= @p_155 @p_113))) (not (= @p_155 @p_71))) (not (= @p_155 @p_29))) (not (= @p_113 @p_71))) (not (= @p_113 @p_29))) (not (= @p_71 @p_29)))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= @p_232 @p_190)) (not (= @p_232 @p_148))) (not (= @p_232 @p_106))) (not (= @p_232 @p_64))) (not (= @p_232 @p_22))) (not (= @p_190 @p_148))) (not (= @p_190 @p_106))) (not (= @p_190 @p_64))) (not (= @p_190 @p_22))) (not (= @p_148 @p_106))) (not (= @p_148 @p_64))) (not (= @p_148 @p_22))) (not (= @p_106 @p_64))) (not (= @p_106 @p_22))) (not (= @p_64 @p_22)))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= @p_225 @p_183)) (not (= @p_225 @p_141))) (not (= @p_225 @p_99))) (not (= @p_225 @p_57))) (not (= @p_225 @p_15))) (not (= @p_183 @p_141))) (not (= @p_183 @p_99))) (not (= @p_183 @p_57))) (not (= @p_183 @p_15))) (not (= @p_141 @p_99))) (not (= @p_141 @p_57))) (not (= @p_141 @p_15))) (not (= @p_99 @p_57))) (not (= @p_99 @p_15))) (not (= @p_57 @p_15)))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= @p_218 @p_176)) (not (= @p_218 @p_134))) (not (= @p_218 @p_92))) (not (= @p_218 @p_50))) (not (= @p_218 @p_8))) (not (= @p_176 @p_134))) (not (= @p_176 @p_92))) (not (= @p_176 @p_50))) (not (= @p_176 @p_8))) (not (= @p_134 @p_92))) (not (= @p_134 @p_50))) (not (= @p_134 @p_8))) (not (= @p_92 @p_50))) (not (= @p_92 @p_8))) (not (= @p_50 @p_8)))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= @p_211 @p_169)) (not (= @p_211 @p_127))) (not (= @p_211 @p_85))) (not (= @p_211 @p_43))) (not (= @p_211 @p_1))) (not (= @p_169 @p_127))) (not (= @p_169 @p_85))) (not (= @p_169 @p_43))) (not (= @p_169 @p_1))) (not (= @p_127 @p_85))) (not (= @p_127 @p_43))) (not (= @p_127 @p_1))) (not (= @p_85 @p_43))) (not (= @p_85 @p_1))) (not (= @p_43 @p_1)))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= @p_246 @p_239)) (not (= @p_246 @p_232))) (not (= @p_246 @p_225))) (not (= @p_246 @p_218))) (not (= @p_246 @p_211))) (not (= @p_239 @p_232))) (not (= @p_239 @p_225))) (not (= @p_239 @p_218))) (not (= @p_239 @p_211))) (not (= @p_232 @p_225))) (not (= @p_232 @p_218))) (not (= @p_232 @p_211))) (not (= @p_225 @p_218))) (not (= @p_225 @p_211))) (not (= @p_218 @p_211))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= @p_204 @p_197)) (not (= @p_204 @p_190))) (not (= @p_204 @p_183))) (not (= @p_204 @p_176))) (not (= @p_204 @p_169))) (not (= @p_197 @p_190))) (not (= @p_197 @p_183))) (not (= @p_197 @p_176))) (not (= @p_197 @p_169))) (not (= @p_190 @p_183))) (not (= @p_190 @p_176))) (not (= @p_190 @p_169))) (not (= @p_183 @p_176))) (not (= @p_183 @p_169))) (not (= @p_176 @p_169)))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= @p_162 @p_155)) (not (= @p_162 @p_148))) (not (= @p_162 @p_141))) (not (= @p_162 @p_134))) (not (= @p_162 @p_127))) (not (= @p_155 @p_148))) (not (= @p_155 @p_141))) (not (= @p_155 @p_134))) (not (= @p_155 @p_127))) (not (= @p_148 @p_141))) (not (= @p_148 @p_134))) (not (= @p_148 @p_127))) (not (= @p_141 @p_134))) (not (= @p_141 @p_127))) (not (= @p_134 @p_127)))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= @p_120 @p_113)) (not (= @p_120 @p_106))) (not (= @p_120 @p_99))) (not (= @p_120 @p_92))) (not (= @p_120 @p_85))) (not (= @p_113 @p_106))) (not (= @p_113 @p_99))) (not (= @p_113 @p_92))) (not (= @p_113 @p_85))) (not (= @p_106 @p_99))) (not (= @p_106 @p_92))) (not (= @p_106 @p_85))) (not (= @p_99 @p_92))) (not (= @p_99 @p_85))) (not (= @p_92 @p_85)))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= @p_78 @p_71)) (not (= @p_78 @p_64))) (not (= @p_78 @p_57))) (not (= @p_78 @p_50))) (not (= @p_78 @p_43))) (not (= @p_71 @p_64))) (not (= @p_71 @p_57))) (not (= @p_71 @p_50))) (not (= @p_71 @p_43))) (not (= @p_64 @p_57))) (not (= @p_64 @p_50))) (not (= @p_64 @p_43))) (not (= @p_57 @p_50))) (not (= @p_57 @p_43))) (not (= @p_50 @p_43)))) (and (and (and (and (and (and (and (and (and (and (and (and (and (and (not (= @p_36 @p_29)) (not (= @p_36 @p_22))) (not (= @p_36 @p_15))) (not (= @p_36 @p_8))) (not (= @p_36 @p_1))) (not (= @p_29 @p_22))) (not (= @p_29 @p_15))) (not (= @p_29 @p_8))) (not (= @p_29 @p_1))) (not (= @p_22 @p_15))) (not (= @p_22 @p_8))) (not (= @p_22 @p_1))) (not (= @p_15 @p_8))) (not (= @p_15 @p_1))) (not (= @p_8 @p_1))))))
(assume a10 (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (not (! (= e0 e1) :named @p_395)) :named @p_396) (! (not (! (= e0 e2) :named @p_393)) :named @p_394)) :named @p_397) (! (not (! (= e0 e3) :named @p_391)) :named @p_392)) :named @p_398) (! (not (! (= e0 e4) :named @p_389)) :named @p_390)) :named @p_399) (! (not (! (= e0 e5) :named @p_387)) :named @p_388)) :named @p_400) (not (= e1 e2))) :named @p_401) (not (= e1 e3))) :named @p_402) (not (= e1 e4))) :named @p_403) (not (= e1 e5))) :named @p_404) (not (= e2 e3))) :named @p_405) (not (= e2 e4))) :named @p_406) (not (= e2 e5))) :named @p_407) (not (= e3 e4))) :named @p_408) (not (= e3 e5))) :named @p_409) (not (= e4 e5))) :named @p_410))
(step t0 (cl (! (and (! (= e2 @p_162) :named @p_411) (! (= e0 @p_347) :named @p_421)) :named @p_537) (! (not @p_411) :named @p_535) (! (not @p_421) :named @p_536)) :rule and_neg)
(step t1 (cl (! (=> @p_537 (! (= e2 @p_326) :named @p_414)) :named @p_539) @p_537) :rule implies_neg1)
(anchor :step t2)
(assume t2.a0 @p_411)
(assume t2.a1 @p_421)
(step t2.t0 (cl (not (! (= @p_269 @p_411) :named @p_412)) (not @p_269) @p_411) :rule equiv_pos2)
(step t2.t1 (cl (= @p_269 (! (= e2 @p_268) :named @p_413))) :rule all_simplify)
(step t2.t2 (cl (= e2 e2)) :rule refl)
(step t2.t3 (cl (= @p_268 @p_162)) :rule cong :premises (t2.t2 a8))
(step t2.t4 (cl (= @p_413 @p_411)) :rule cong :premises (t2.t2 t2.t3))
(step t2.t5 (cl @p_412) :rule trans :premises (t2.t1 t2.t4))
(step t2.t6 (cl @p_281) :rule and :premises (a2))
(step t2.t7 (cl @p_280) :rule and :premises (t2.t6))
(step t2.t8 (cl @p_279) :rule and :premises (t2.t7))
(step t2.t9 (cl @p_278) :rule and :premises (t2.t8))
(step t2.t10 (cl @p_270) :rule and :premises (t2.t9))
(step t2.t11 (cl @p_269) :rule and :premises (t2.t10))
(step t2.t12 (cl @p_411) :rule resolution :premises (t2.t0 t2.t5 t2.t11))
(step t2.t13 (cl @p_348) :rule symm :premises (t2.a1))
(step t2.t14 (cl @p_421) :rule symm :premises (t2.t13))
(step t2.t15 (cl @p_166) :rule symm :premises (t2.t12))
(step t2.t16 (cl (= @p_347 @p_148)) :rule cong :premises (t2.t2 t2.t15))
(step t2.t17 (cl (= e0 @p_148)) :rule trans :premises (t2.t14 t2.t16))
(step t2.t18 (cl (= @p_162 @p_326)) :rule cong :premises (t2.t2 t2.t17))
(step t2.t19 (cl @p_414) :rule trans :premises (t2.t12 t2.t18))
(step t2 (cl @p_535 @p_536 @p_414) :rule subproof :discharge (t2.a0 t2.a1))
(step t3 (cl (! (not @p_537) :named @p_538) @p_411) :rule and_pos)
(step t4 (cl @p_538 @p_421) :rule and_pos)
(step t5 (cl @p_414 @p_538 @p_538) :rule resolution :premises (t2 t3 t4))
(step t6 (cl @p_538 @p_538 @p_414) :rule reordering :premises (t5))
(step t7 (cl @p_538 @p_414) :rule contraction :premises (t6))
(step t8 (cl @p_539 @p_414) :rule resolution :premises (t1 t7))
(step t9 (cl @p_539 (! (not @p_414) :named @p_415)) :rule implies_neg2)
(step t10 (cl @p_539 @p_539) :rule resolution :premises (t8 t9))
(step t11 (cl @p_539) :rule contraction :premises (t10))
(step t12 (cl @p_538 @p_414) :rule implies :premises (t11))
(step t13 (cl @p_535 @p_536 @p_414) :rule resolution :premises (t0 t12))
(step t14 (cl @p_414 @p_535 @p_536) :rule reordering :premises (t13))
(step t15 (cl (! (not (! (and @p_421 (! (not (! (= e0 @p_162) :named @p_422)) :named @p_423)) :named @p_424)) :named @p_534) @p_421) :rule and_pos)
(step t16 (cl @p_421 @p_534) :rule reordering :premises (t15))
(step t17 (cl (not (! (= @p_360 (! (or (! (and (! (= e0 @p_332) :named @p_427) (! (not (! (= e0 @p_36) :named @p_425)) :named @p_426)) :named @p_428) (! (and (! (= e0 @p_337) :named @p_447) (! (not (! (= e0 @p_78) :named @p_445)) :named @p_446)) :named @p_448) (! (and (! (= e0 @p_342) :named @p_467) (! (not (! (= e0 @p_120) :named @p_465)) :named @p_466)) :named @p_468) @p_424 (! (and (! (= e0 @p_357) :named @p_487) (! (not (! (= e0 @p_246) :named @p_485)) :named @p_486)) :named @p_488) (! (and (! (= e0 @p_352) :named @p_494) (! (not (! (= e0 @p_204) :named @p_492)) :named @p_493)) :named @p_495)) :named @p_512)) :named @p_513)) (not @p_360) @p_512) :rule equiv_pos2)
(step t18 (cl (= @p_360 (! (or @p_334 @p_339 @p_344 @p_349 @p_359 @p_354) :named @p_533))) :rule all_simplify)
(step t19 (cl (= @p_333 @p_427)) :rule all_simplify)
(step t20 (cl (= @p_330 (! (= unit @p_36) :named @p_529))) :rule all_simplify)
(step t21 (cl (= @p_331 (! (not @p_529) :named @p_530))) :rule cong :premises (t20))
(step t22 (cl (= @p_334 (! (and @p_427 @p_530) :named @p_531))) :rule cong :premises (t19 t21))
(step t23 (cl (= @p_338 @p_447)) :rule all_simplify)
(step t24 (cl (= @p_335 (! (= unit @p_78) :named @p_526))) :rule all_simplify)
(step t25 (cl (= @p_336 (! (not @p_526) :named @p_527))) :rule cong :premises (t24))
(step t26 (cl (= @p_339 (! (and @p_447 @p_527) :named @p_528))) :rule cong :premises (t23 t25))
(step t27 (cl (= @p_343 @p_467)) :rule all_simplify)
(step t28 (cl (= @p_340 (! (= unit @p_120) :named @p_523))) :rule all_simplify)
(step t29 (cl (= @p_341 (! (not @p_523) :named @p_524))) :rule cong :premises (t28))
(step t30 (cl (= @p_344 (! (and @p_467 @p_524) :named @p_525))) :rule cong :premises (t27 t29))
(step t31 (cl (= @p_348 @p_421)) :rule all_simplify)
(step t32 (cl (= @p_345 (! (= unit @p_162) :named @p_520))) :rule all_simplify)
(step t33 (cl (= @p_346 (! (not @p_520) :named @p_521))) :rule cong :premises (t32))
(step t34 (cl (= @p_349 (! (and @p_421 @p_521) :named @p_522))) :rule cong :premises (t31 t33))
(step t35 (cl (= @p_358 @p_487)) :rule all_simplify)
(step t36 (cl (= @p_355 (! (= unit @p_246) :named @p_517))) :rule all_simplify)
(step t37 (cl (= @p_356 (! (not @p_517) :named @p_518))) :rule cong :premises (t36))
(step t38 (cl (= @p_359 (! (and @p_487 @p_518) :named @p_519))) :rule cong :premises (t35 t37))
(step t39 (cl (= @p_353 @p_494)) :rule all_simplify)
(step t40 (cl (= @p_350 (! (= unit @p_204) :named @p_514))) :rule all_simplify)
(step t41 (cl (= @p_351 (! (not @p_514) :named @p_515))) :rule cong :premises (t40))
(step t42 (cl (= @p_354 (! (and @p_494 @p_515) :named @p_516))) :rule cong :premises (t39 t41))
(step t43 (cl (= @p_533 (! (or @p_531 @p_528 @p_525 @p_522 @p_519 @p_516) :named @p_532))) :rule cong :premises (t22 t26 t30 t34 t38 t42))
(step t44 (cl (= @p_360 @p_532)) :rule trans :premises (t18 t43))
(step t45 (cl (= @p_427 @p_427)) :rule refl)
(step t46 (cl (= @p_36 @p_36)) :rule refl)
(step t47 (cl (= @p_529 @p_425)) :rule cong :premises (a8 t46))
(step t48 (cl (= @p_530 @p_426)) :rule cong :premises (t47))
(step t49 (cl (= @p_531 @p_428)) :rule cong :premises (t45 t48))
(step t50 (cl (= @p_447 @p_447)) :rule refl)
(step t51 (cl (= @p_78 @p_78)) :rule refl)
(step t52 (cl (= @p_526 @p_445)) :rule cong :premises (a8 t51))
(step t53 (cl (= @p_527 @p_446)) :rule cong :premises (t52))
(step t54 (cl (= @p_528 @p_448)) :rule cong :premises (t50 t53))
(step t55 (cl (= @p_467 @p_467)) :rule refl)
(step t56 (cl (= @p_120 @p_120)) :rule refl)
(step t57 (cl (= @p_523 @p_465)) :rule cong :premises (a8 t56))
(step t58 (cl (= @p_524 @p_466)) :rule cong :premises (t57))
(step t59 (cl (= @p_525 @p_468)) :rule cong :premises (t55 t58))
(step t60 (cl (= @p_421 @p_421)) :rule refl)
(step t61 (cl (= @p_162 @p_162)) :rule refl)
(step t62 (cl (= @p_520 @p_422)) :rule cong :premises (a8 t61))
(step t63 (cl (= @p_521 @p_423)) :rule cong :premises (t62))
(step t64 (cl (= @p_522 @p_424)) :rule cong :premises (t60 t63))
(step t65 (cl (= @p_487 @p_487)) :rule refl)
(step t66 (cl @p_283) :rule refl)
(step t67 (cl (= @p_517 @p_485)) :rule cong :premises (a8 t66))
(step t68 (cl (= @p_518 @p_486)) :rule cong :premises (t67))
(step t69 (cl (= @p_519 @p_488)) :rule cong :premises (t65 t68))
(step t70 (cl (= @p_494 @p_494)) :rule refl)
(step t71 (cl (= @p_204 @p_204)) :rule refl)
(step t72 (cl (= @p_514 @p_492)) :rule cong :premises (a8 t71))
(step t73 (cl (= @p_515 @p_493)) :rule cong :premises (t72))
(step t74 (cl (= @p_516 @p_495)) :rule cong :premises (t70 t73))
(step t75 (cl (= @p_532 @p_512)) :rule cong :premises (t49 t54 t59 t64 t69 t74))
(step t76 (cl @p_513) :rule trans :premises (t44 t75))
(step t77 (cl @p_365) :rule and :premises (a7))
(step t78 (cl @p_364) :rule and :premises (t77))
(step t79 (cl @p_363) :rule and :premises (t78))
(step t80 (cl @p_362) :rule and :premises (t79))
(step t81 (cl @p_361) :rule and :premises (t80))
(step t82 (cl @p_360) :rule and :premises (t81))
(step t83 (cl @p_512) :rule resolution :premises (t17 t76 t82))
(step t84 (cl @p_428 @p_448 @p_468 @p_424 @p_488 @p_495) :rule or :premises (t83))
(step t85 (cl @p_495 @p_488 @p_424 @p_468 @p_448 @p_428) :rule reordering :premises (t84))
(step t86 (cl (! (not @p_495) :named @p_496) @p_494) :rule and_pos)
(step t87 (cl @p_494 @p_496) :rule reordering :premises (t86))
(step t88 (cl (! (and (! (= e1 @p_204) :named @p_498) @p_494) :named @p_509) (! (not @p_498) :named @p_508) (! (not @p_494) :named @p_497)) :rule and_neg)
(step t89 (cl (! (=> @p_509 (! (= e1 @p_328) :named @p_501)) :named @p_511) @p_509) :rule implies_neg1)
(anchor :step t90)
(assume t90.a0 @p_498)
(assume t90.a1 @p_494)
(step t90.t0 (cl (not (! (= @p_272 @p_498) :named @p_499)) (not @p_272) @p_498) :rule equiv_pos2)
(step t90.t1 (cl (= @p_272 (! (= e1 @p_271) :named @p_500))) :rule all_simplify)
(step t90.t2 (cl (= e1 e1)) :rule refl)
(step t90.t3 (cl (= @p_271 @p_204)) :rule cong :premises (t90.t2 a8))
(step t90.t4 (cl (= @p_500 @p_498)) :rule cong :premises (t90.t2 t90.t3))
(step t90.t5 (cl @p_499) :rule trans :premises (t90.t1 t90.t4))
(step t90.t6 (cl @p_281) :rule and :premises (a2))
(step t90.t7 (cl @p_280) :rule and :premises (t90.t6))
(step t90.t8 (cl @p_279) :rule and :premises (t90.t7))
(step t90.t9 (cl @p_278) :rule and :premises (t90.t8))
(step t90.t10 (cl @p_277) :rule and :premises (t90.t9))
(step t90.t11 (cl @p_273) :rule and :premises (t90.t10))
(step t90.t12 (cl @p_272) :rule and :premises (t90.t11))
(step t90.t13 (cl @p_498) :rule resolution :premises (t90.t0 t90.t5 t90.t12))
(step t90.t14 (cl @p_353) :rule symm :premises (t90.a1))
(step t90.t15 (cl @p_494) :rule symm :premises (t90.t14))
(step t90.t16 (cl @p_209) :rule symm :premises (t90.t13))
(step t90.t17 (cl (= @p_352 @p_197)) :rule cong :premises (t90.t2 t90.t16))
(step t90.t18 (cl (= e0 @p_197)) :rule trans :premises (t90.t15 t90.t17))
(step t90.t19 (cl (= @p_204 @p_328)) :rule cong :premises (t90.t2 t90.t18))
(step t90.t20 (cl @p_501) :rule trans :premises (t90.t13 t90.t19))
(step t90 (cl @p_508 @p_497 @p_501) :rule subproof :discharge (t90.a0 t90.a1))
(step t91 (cl (! (not @p_509) :named @p_510) @p_498) :rule and_pos)
(step t92 (cl @p_510 @p_494) :rule and_pos)
(step t93 (cl @p_501 @p_510 @p_510) :rule resolution :premises (t90 t91 t92))
(step t94 (cl @p_510 @p_510 @p_501) :rule reordering :premises (t93))
(step t95 (cl @p_510 @p_501) :rule contraction :premises (t94))
(step t96 (cl @p_511 @p_501) :rule resolution :premises (t89 t95))
(step t97 (cl @p_511 (! (not @p_501) :named @p_502)) :rule implies_neg2)
(step t98 (cl @p_511 @p_511) :rule resolution :premises (t96 t97))
(step t99 (cl @p_511) :rule contraction :premises (t98))
(step t100 (cl @p_510 @p_501) :rule implies :premises (t99))
(step t101 (cl @p_508 @p_497 @p_501) :rule resolution :premises (t88 t100))
(step t102 (cl @p_501 @p_508 @p_497) :rule reordering :premises (t101))
(step t103 (cl (not (! (= @p_380 (! (or @p_502 (! (= e1 e0) :named @p_503)) :named @p_504)) :named @p_505)) (not @p_380) @p_504) :rule equiv_pos2)
(step t104 (cl (= @p_329 @p_501)) :rule all_simplify)
(step t105 (cl (= @p_379 @p_502)) :rule cong :premises (t104))
(step t106 (cl (= @p_378 @p_257)) :rule all_simplify)
(step t107 (cl (= @p_380 (! (or @p_502 @p_257) :named @p_506))) :rule cong :premises (t105 t106))
(step t108 (cl (= @p_502 @p_502)) :rule refl)
(step t109 (cl (= e1 e1)) :rule refl)
(step t110 (cl (= @p_257 @p_395)) :rule cong :premises (a8 t109))
(step t111 (cl (= @p_506 (! (or @p_502 @p_395) :named @p_507))) :rule cong :premises (t108 t110))
(step t112 (cl (= @p_395 @p_503)) :rule all_simplify)
(step t113 (cl (= @p_507 @p_504)) :rule cong :premises (t108 t112))
(step t114 (cl (= @p_506 @p_504)) :rule trans :premises (t111 t113))
(step t115 (cl @p_505) :rule trans :premises (t107 t114))
(step t116 (cl @p_385) :rule and :premises (a7))
(step t117 (cl @p_384) :rule and :premises (t116))
(step t118 (cl @p_383) :rule and :premises (t117))
(step t119 (cl @p_382) :rule and :premises (t118))
(step t120 (cl @p_381) :rule and :premises (t119))
(step t121 (cl @p_380) :rule and :premises (t120))
(step t122 (cl @p_504) :rule resolution :premises (t103 t115 t121))
(step t123 (cl @p_502 @p_503) :rule or :premises (t122))
(step t124 (cl @p_503 @p_502) :rule reordering :premises (t123))
(step t125 (cl @p_409) :rule and :premises (a10))
(step t126 (cl @p_408) :rule and :premises (t125))
(step t127 (cl @p_407) :rule and :premises (t126))
(step t128 (cl @p_406) :rule and :premises (t127))
(step t129 (cl @p_405) :rule and :premises (t128))
(step t130 (cl @p_404) :rule and :premises (t129))
(step t131 (cl @p_403) :rule and :premises (t130))
(step t132 (cl @p_402) :rule and :premises (t131))
(step t133 (cl @p_401) :rule and :premises (t132))
(step t134 (cl @p_400) :rule and :premises (t133))
(step t135 (cl @p_399) :rule and :premises (t134))
(step t136 (cl @p_398) :rule and :premises (t135))
(step t137 (cl @p_397) :rule and :premises (t136))
(step t138 (cl @p_396) :rule and :premises (t137))
(step t139 (cl (not @p_503)) :rule not_symm :premises (t138))
(step t140 (cl @p_502) :rule resolution :premises (t124 t139))
(step t141 (cl (not @p_499) (not @p_272) @p_498) :rule equiv_pos2)
(step t142 (cl (= @p_272 @p_500)) :rule all_simplify)
(step t143 (cl (= @p_271 @p_204)) :rule cong :premises (t109 a8))
(step t144 (cl (= @p_500 @p_498)) :rule cong :premises (t109 t143))
(step t145 (cl @p_499) :rule trans :premises (t142 t144))
(step t146 (cl @p_281) :rule and :premises (a2))
(step t147 (cl @p_280) :rule and :premises (t146))
(step t148 (cl @p_279) :rule and :premises (t147))
(step t149 (cl @p_278) :rule and :premises (t148))
(step t150 (cl @p_277) :rule and :premises (t149))
(step t151 (cl @p_273) :rule and :premises (t150))
(step t152 (cl @p_272) :rule and :premises (t151))
(step t153 (cl @p_498) :rule resolution :premises (t141 t145 t152))
(step t154 (cl @p_497) :rule resolution :premises (t102 t140 t153))
(step t155 (cl @p_496) :rule resolution :premises (t87 t154))
(step t156 (cl (! (not @p_488) :named @p_489) @p_486) :rule and_pos)
(step t157 (cl @p_486 @p_489) :rule reordering :premises (t156))
(step t158 (cl (not (! (= @p_275 @p_485) :named @p_490)) (not @p_275) @p_485) :rule equiv_pos2)
(step t159 (cl (= @p_275 (! (= e0 @p_274) :named @p_491))) :rule all_simplify)
(step t160 (cl (= e0 e0)) :rule refl)
(step t161 (cl (= @p_274 @p_246)) :rule cong :premises (a8 t160))
(step t162 (cl (= @p_491 @p_485)) :rule cong :premises (t160 t161))
(step t163 (cl @p_490) :rule trans :premises (t159 t162))
(step t164 (cl @p_276) :rule and :premises (t150))
(step t165 (cl @p_275) :rule and :premises (t164))
(step t166 (cl @p_485) :rule resolution :premises (t158 t163 t165))
(step t167 (cl @p_489) :rule resolution :premises (t157 t166))
(step t168 (cl (! (not @p_468) :named @p_469) @p_467) :rule and_pos)
(step t169 (cl @p_467 @p_469) :rule reordering :premises (t168))
(step t170 (cl (! (and (! (= e3 @p_120) :named @p_471) @p_467) :named @p_482) (! (not @p_471) :named @p_481) (! (not @p_467) :named @p_470)) :rule and_neg)
(step t171 (cl (! (=> @p_482 (! (= e3 @p_324) :named @p_474)) :named @p_484) @p_482) :rule implies_neg1)
(anchor :step t172)
(assume t172.a0 @p_471)
(assume t172.a1 @p_467)
(step t172.t0 (cl (not (! (= @p_266 @p_471) :named @p_472)) (not @p_266) @p_471) :rule equiv_pos2)
(step t172.t1 (cl (= @p_266 (! (= e3 @p_265) :named @p_473))) :rule all_simplify)
(step t172.t2 (cl (= e3 e3)) :rule refl)
(step t172.t3 (cl (= @p_265 @p_120)) :rule cong :premises (t172.t2 a8))
(step t172.t4 (cl (= @p_473 @p_471)) :rule cong :premises (t172.t2 t172.t3))
(step t172.t5 (cl @p_472) :rule trans :premises (t172.t1 t172.t4))
(step t172.t6 (cl @p_267) :rule and :premises (t148))
(step t172.t7 (cl @p_266) :rule and :premises (t172.t6))
(step t172.t8 (cl @p_471) :rule resolution :premises (t172.t0 t172.t5 t172.t7))
(step t172.t9 (cl @p_343) :rule symm :premises (t172.a1))
(step t172.t10 (cl @p_467) :rule symm :premises (t172.t9))
(step t172.t11 (cl @p_123) :rule symm :premises (t172.t8))
(step t172.t12 (cl (= @p_342 @p_99)) :rule cong :premises (t172.t2 t172.t11))
(step t172.t13 (cl (= e0 @p_99)) :rule trans :premises (t172.t10 t172.t12))
(step t172.t14 (cl (= @p_120 @p_324)) :rule cong :premises (t172.t2 t172.t13))
(step t172.t15 (cl @p_474) :rule trans :premises (t172.t8 t172.t14))
(step t172 (cl @p_481 @p_470 @p_474) :rule subproof :discharge (t172.a0 t172.a1))
(step t173 (cl (! (not @p_482) :named @p_483) @p_471) :rule and_pos)
(step t174 (cl @p_483 @p_467) :rule and_pos)
(step t175 (cl @p_474 @p_483 @p_483) :rule resolution :premises (t172 t173 t174))
(step t176 (cl @p_483 @p_483 @p_474) :rule reordering :premises (t175))
(step t177 (cl @p_483 @p_474) :rule contraction :premises (t176))
(step t178 (cl @p_484 @p_474) :rule resolution :premises (t171 t177))
(step t179 (cl @p_484 (! (not @p_474) :named @p_475)) :rule implies_neg2)
(step t180 (cl @p_484 @p_484) :rule resolution :premises (t178 t179))
(step t181 (cl @p_484) :rule contraction :premises (t180))
(step t182 (cl @p_483 @p_474) :rule implies :premises (t181))
(step t183 (cl @p_481 @p_470 @p_474) :rule resolution :premises (t170 t182))
(step t184 (cl @p_474 @p_481 @p_470) :rule reordering :premises (t183))
(step t185 (cl (not (! (= @p_374 (! (or @p_475 (! (= e3 e0) :named @p_476)) :named @p_477)) :named @p_478)) (not @p_374) @p_477) :rule equiv_pos2)
(step t186 (cl (= @p_325 @p_474)) :rule all_simplify)
(step t187 (cl (= @p_373 @p_475)) :rule cong :premises (t186))
(step t188 (cl (= @p_372 @p_255)) :rule all_simplify)
(step t189 (cl (= @p_374 (! (or @p_475 @p_255) :named @p_479))) :rule cong :premises (t187 t188))
(step t190 (cl (= @p_475 @p_475)) :rule refl)
(step t191 (cl (= e3 e3)) :rule refl)
(step t192 (cl (= @p_255 @p_391)) :rule cong :premises (a8 t191))
(step t193 (cl (= @p_479 (! (or @p_475 @p_391) :named @p_480))) :rule cong :premises (t190 t192))
(step t194 (cl (= @p_391 @p_476)) :rule all_simplify)
(step t195 (cl (= @p_480 @p_477)) :rule cong :premises (t190 t194))
(step t196 (cl (= @p_479 @p_477)) :rule trans :premises (t193 t195))
(step t197 (cl @p_478) :rule trans :premises (t189 t196))
(step t198 (cl @p_374) :rule and :premises (t118))
(step t199 (cl @p_477) :rule resolution :premises (t185 t197 t198))
(step t200 (cl @p_475 @p_476) :rule or :premises (t199))
(step t201 (cl @p_476 @p_475) :rule reordering :premises (t200))
(step t202 (cl @p_392) :rule and :premises (t136))
(step t203 (cl (not @p_476)) :rule not_symm :premises (t202))
(step t204 (cl @p_475) :rule resolution :premises (t201 t203))
(step t205 (cl (not @p_472) (not @p_266) @p_471) :rule equiv_pos2)
(step t206 (cl (= @p_266 @p_473)) :rule all_simplify)
(step t207 (cl (= @p_265 @p_120)) :rule cong :premises (t191 a8))
(step t208 (cl (= @p_473 @p_471)) :rule cong :premises (t191 t207))
(step t209 (cl @p_472) :rule trans :premises (t206 t208))
(step t210 (cl @p_267) :rule and :premises (t148))
(step t211 (cl @p_266) :rule and :premises (t210))
(step t212 (cl @p_471) :rule resolution :premises (t205 t209 t211))
(step t213 (cl @p_470) :rule resolution :premises (t184 t204 t212))
(step t214 (cl @p_469) :rule resolution :premises (t169 t213))
(step t215 (cl (! (not @p_448) :named @p_449) @p_447) :rule and_pos)
(step t216 (cl @p_447 @p_449) :rule reordering :premises (t215))
(step t217 (cl (! (and (! (= e4 @p_78) :named @p_451) @p_447) :named @p_462) (! (not @p_451) :named @p_461) (! (not @p_447) :named @p_450)) :rule and_neg)
(step t218 (cl (! (=> @p_462 (! (= e4 @p_322) :named @p_454)) :named @p_464) @p_462) :rule implies_neg1)
(anchor :step t219)
(assume t219.a0 @p_451)
(assume t219.a1 @p_447)
(step t219.t0 (cl (not (! (= @p_263 @p_451) :named @p_452)) (not @p_263) @p_451) :rule equiv_pos2)
(step t219.t1 (cl (= @p_263 (! (= e4 @p_262) :named @p_453))) :rule all_simplify)
(step t219.t2 (cl (= e4 e4)) :rule refl)
(step t219.t3 (cl (= @p_262 @p_78)) :rule cong :premises (t219.t2 a8))
(step t219.t4 (cl (= @p_453 @p_451)) :rule cong :premises (t219.t2 t219.t3))
(step t219.t5 (cl @p_452) :rule trans :premises (t219.t1 t219.t4))
(step t219.t6 (cl @p_264) :rule and :premises (t147))
(step t219.t7 (cl @p_263) :rule and :premises (t219.t6))
(step t219.t8 (cl @p_451) :rule resolution :premises (t219.t0 t219.t5 t219.t7))
(step t219.t9 (cl @p_338) :rule symm :premises (t219.a1))
(step t219.t10 (cl @p_447) :rule symm :premises (t219.t9))
(step t219.t11 (cl @p_80) :rule symm :premises (t219.t8))
(step t219.t12 (cl (= @p_337 @p_50)) :rule cong :premises (t219.t2 t219.t11))
(step t219.t13 (cl (= e0 @p_50)) :rule trans :premises (t219.t10 t219.t12))
(step t219.t14 (cl (= @p_78 @p_322)) :rule cong :premises (t219.t2 t219.t13))
(step t219.t15 (cl @p_454) :rule trans :premises (t219.t8 t219.t14))
(step t219 (cl @p_461 @p_450 @p_454) :rule subproof :discharge (t219.a0 t219.a1))
(step t220 (cl (! (not @p_462) :named @p_463) @p_451) :rule and_pos)
(step t221 (cl @p_463 @p_447) :rule and_pos)
(step t222 (cl @p_454 @p_463 @p_463) :rule resolution :premises (t219 t220 t221))
(step t223 (cl @p_463 @p_463 @p_454) :rule reordering :premises (t222))
(step t224 (cl @p_463 @p_454) :rule contraction :premises (t223))
(step t225 (cl @p_464 @p_454) :rule resolution :premises (t218 t224))
(step t226 (cl @p_464 (! (not @p_454) :named @p_455)) :rule implies_neg2)
(step t227 (cl @p_464 @p_464) :rule resolution :premises (t225 t226))
(step t228 (cl @p_464) :rule contraction :premises (t227))
(step t229 (cl @p_463 @p_454) :rule implies :premises (t228))
(step t230 (cl @p_461 @p_450 @p_454) :rule resolution :premises (t217 t229))
(step t231 (cl @p_454 @p_461 @p_450) :rule reordering :premises (t230))
(step t232 (cl (not (! (= @p_371 (! (or @p_455 (! (= e4 e0) :named @p_456)) :named @p_457)) :named @p_458)) (not @p_371) @p_457) :rule equiv_pos2)
(step t233 (cl (= @p_323 @p_454)) :rule all_simplify)
(step t234 (cl (= @p_370 @p_455)) :rule cong :premises (t233))
(step t235 (cl (= @p_369 @p_254)) :rule all_simplify)
(step t236 (cl (= @p_371 (! (or @p_455 @p_254) :named @p_459))) :rule cong :premises (t234 t235))
(step t237 (cl (= @p_455 @p_455)) :rule refl)
(step t238 (cl (= e4 e4)) :rule refl)
(step t239 (cl (= @p_254 @p_389)) :rule cong :premises (a8 t238))
(step t240 (cl (= @p_459 (! (or @p_455 @p_389) :named @p_460))) :rule cong :premises (t237 t239))
(step t241 (cl (= @p_389 @p_456)) :rule all_simplify)
(step t242 (cl (= @p_460 @p_457)) :rule cong :premises (t237 t241))
(step t243 (cl (= @p_459 @p_457)) :rule trans :premises (t240 t242))
(step t244 (cl @p_458) :rule trans :premises (t236 t243))
(step t245 (cl @p_371) :rule and :premises (t117))
(step t246 (cl @p_457) :rule resolution :premises (t232 t244 t245))
(step t247 (cl @p_455 @p_456) :rule or :premises (t246))
(step t248 (cl @p_456 @p_455) :rule reordering :premises (t247))
(step t249 (cl @p_390) :rule and :premises (t135))
(step t250 (cl (not @p_456)) :rule not_symm :premises (t249))
(step t251 (cl @p_455) :rule resolution :premises (t248 t250))
(step t252 (cl (not @p_452) (not @p_263) @p_451) :rule equiv_pos2)
(step t253 (cl (= @p_263 @p_453)) :rule all_simplify)
(step t254 (cl (= @p_262 @p_78)) :rule cong :premises (t238 a8))
(step t255 (cl (= @p_453 @p_451)) :rule cong :premises (t238 t254))
(step t256 (cl @p_452) :rule trans :premises (t253 t255))
(step t257 (cl @p_264) :rule and :premises (t147))
(step t258 (cl @p_263) :rule and :premises (t257))
(step t259 (cl @p_451) :rule resolution :premises (t252 t256 t258))
(step t260 (cl @p_450) :rule resolution :premises (t231 t251 t259))
(step t261 (cl @p_449) :rule resolution :premises (t216 t260))
(step t262 (cl (! (not @p_428) :named @p_429) @p_427) :rule and_pos)
(step t263 (cl @p_427 @p_429) :rule reordering :premises (t262))
(step t264 (cl (! (and (! (= e5 @p_36) :named @p_431) @p_427) :named @p_442) (! (not @p_431) :named @p_441) (! (not @p_427) :named @p_430)) :rule and_neg)
(step t265 (cl (! (=> @p_442 (! (= e5 @p_320) :named @p_434)) :named @p_444) @p_442) :rule implies_neg1)
(anchor :step t266)
(assume t266.a0 @p_431)
(assume t266.a1 @p_427)
(step t266.t0 (cl (not (! (= @p_260 @p_431) :named @p_432)) (not @p_260) @p_431) :rule equiv_pos2)
(step t266.t1 (cl (= @p_260 (! (= e5 @p_259) :named @p_433))) :rule all_simplify)
(step t266.t2 (cl (= e5 e5)) :rule refl)
(step t266.t3 (cl (= @p_259 @p_36)) :rule cong :premises (t266.t2 a8))
(step t266.t4 (cl (= @p_433 @p_431)) :rule cong :premises (t266.t2 t266.t3))
(step t266.t5 (cl @p_432) :rule trans :premises (t266.t1 t266.t4))
(step t266.t6 (cl @p_261) :rule and :premises (t146))
(step t266.t7 (cl @p_260) :rule and :premises (t266.t6))
(step t266.t8 (cl @p_431) :rule resolution :premises (t266.t0 t266.t5 t266.t7))
(step t266.t9 (cl @p_333) :rule symm :premises (t266.a1))
(step t266.t10 (cl @p_427) :rule symm :premises (t266.t9))
(step t266.t11 (cl @p_37) :rule symm :premises (t266.t8))
(step t266.t12 (cl (= @p_332 @p_1)) :rule cong :premises (t266.t2 t266.t11))
(step t266.t13 (cl (= e0 @p_1)) :rule trans :premises (t266.t10 t266.t12))
(step t266.t14 (cl (= @p_36 @p_320)) :rule cong :premises (t266.t2 t266.t13))
(step t266.t15 (cl @p_434) :rule trans :premises (t266.t8 t266.t14))
(step t266 (cl @p_441 @p_430 @p_434) :rule subproof :discharge (t266.a0 t266.a1))
(step t267 (cl (! (not @p_442) :named @p_443) @p_431) :rule and_pos)
(step t268 (cl @p_443 @p_427) :rule and_pos)
(step t269 (cl @p_434 @p_443 @p_443) :rule resolution :premises (t266 t267 t268))
(step t270 (cl @p_443 @p_443 @p_434) :rule reordering :premises (t269))
(step t271 (cl @p_443 @p_434) :rule contraction :premises (t270))
(step t272 (cl @p_444 @p_434) :rule resolution :premises (t265 t271))
(step t273 (cl @p_444 (! (not @p_434) :named @p_435)) :rule implies_neg2)
(step t274 (cl @p_444 @p_444) :rule resolution :premises (t272 t273))
(step t275 (cl @p_444) :rule contraction :premises (t274))
(step t276 (cl @p_443 @p_434) :rule implies :premises (t275))
(step t277 (cl @p_441 @p_430 @p_434) :rule resolution :premises (t264 t276))
(step t278 (cl @p_434 @p_441 @p_430) :rule reordering :premises (t277))
(step t279 (cl (not (! (= @p_368 (! (or @p_435 (! (= e5 e0) :named @p_436)) :named @p_437)) :named @p_438)) (not @p_368) @p_437) :rule equiv_pos2)
(step t280 (cl (= @p_321 @p_434)) :rule all_simplify)
(step t281 (cl (= @p_367 @p_435)) :rule cong :premises (t280))
(step t282 (cl (= @p_366 @p_253)) :rule all_simplify)
(step t283 (cl (= @p_368 (! (or @p_435 @p_253) :named @p_439))) :rule cong :premises (t281 t282))
(step t284 (cl (= @p_435 @p_435)) :rule refl)
(step t285 (cl (= e5 e5)) :rule refl)
(step t286 (cl (= @p_253 @p_387)) :rule cong :premises (a8 t285))
(step t287 (cl (= @p_439 (! (or @p_435 @p_387) :named @p_440))) :rule cong :premises (t284 t286))
(step t288 (cl (= @p_387 @p_436)) :rule all_simplify)
(step t289 (cl (= @p_440 @p_437)) :rule cong :premises (t284 t288))
(step t290 (cl (= @p_439 @p_437)) :rule trans :premises (t287 t289))
(step t291 (cl @p_438) :rule trans :premises (t283 t290))
(step t292 (cl @p_368) :rule and :premises (t116))
(step t293 (cl @p_437) :rule resolution :premises (t279 t291 t292))
(step t294 (cl @p_435 @p_436) :rule or :premises (t293))
(step t295 (cl @p_436 @p_435) :rule reordering :premises (t294))
(step t296 (cl @p_388) :rule and :premises (t134))
(step t297 (cl (not @p_436)) :rule not_symm :premises (t296))
(step t298 (cl @p_435) :rule resolution :premises (t295 t297))
(step t299 (cl (not @p_432) (not @p_260) @p_431) :rule equiv_pos2)
(step t300 (cl (= @p_260 @p_433)) :rule all_simplify)
(step t301 (cl (= @p_259 @p_36)) :rule cong :premises (t285 a8))
(step t302 (cl (= @p_433 @p_431)) :rule cong :premises (t285 t301))
(step t303 (cl @p_432) :rule trans :premises (t300 t302))
(step t304 (cl @p_261) :rule and :premises (t146))
(step t305 (cl @p_260) :rule and :premises (t304))
(step t306 (cl @p_431) :rule resolution :premises (t299 t303 t305))
(step t307 (cl @p_430) :rule resolution :premises (t278 t298 t306))
(step t308 (cl @p_429) :rule resolution :premises (t263 t307))
(step t309 (cl @p_424) :rule resolution :premises (t85 t155 t167 t214 t261 t308))
(step t310 (cl @p_421) :rule resolution :premises (t16 t309))
(step t311 (cl (not (! (= @p_377 (! (or @p_415 (! (= e2 e0) :named @p_416)) :named @p_417)) :named @p_418)) (not @p_377) @p_417) :rule equiv_pos2)
(step t312 (cl (= @p_327 @p_414)) :rule all_simplify)
(step t313 (cl (= @p_376 @p_415)) :rule cong :premises (t312))
(step t314 (cl (= @p_375 @p_256)) :rule all_simplify)
(step t315 (cl (= @p_377 (! (or @p_415 @p_256) :named @p_419))) :rule cong :premises (t313 t314))
(step t316 (cl (= @p_415 @p_415)) :rule refl)
(step t317 (cl (= e2 e2)) :rule refl)
(step t318 (cl (= @p_256 @p_393)) :rule cong :premises (a8 t317))
(step t319 (cl (= @p_419 (! (or @p_415 @p_393) :named @p_420))) :rule cong :premises (t316 t318))
(step t320 (cl (= @p_393 @p_416)) :rule all_simplify)
(step t321 (cl (= @p_420 @p_417)) :rule cong :premises (t316 t320))
(step t322 (cl (= @p_419 @p_417)) :rule trans :premises (t319 t321))
(step t323 (cl @p_418) :rule trans :premises (t315 t322))
(step t324 (cl @p_377) :rule and :premises (t119))
(step t325 (cl @p_417) :rule resolution :premises (t311 t323 t324))
(step t326 (cl @p_415 @p_416) :rule or :premises (t325))
(step t327 (cl @p_416 @p_415) :rule reordering :premises (t326))
(step t328 (cl @p_394) :rule and :premises (t137))
(step t329 (cl (not @p_416)) :rule not_symm :premises (t328))
(step t330 (cl @p_415) :rule resolution :premises (t327 t329))
(step t331 (cl (not @p_412) (not @p_269) @p_411) :rule equiv_pos2)
(step t332 (cl (= @p_269 @p_413)) :rule all_simplify)
(step t333 (cl (= @p_268 @p_162)) :rule cong :premises (t317 a8))
(step t334 (cl (= @p_413 @p_411)) :rule cong :premises (t317 t333))
(step t335 (cl @p_412) :rule trans :premises (t332 t334))
(step t336 (cl @p_270) :rule and :premises (t149))
(step t337 (cl @p_269) :rule and :premises (t336))
(step t338 (cl @p_411) :rule resolution :premises (t331 t335 t337))
(step t339 (cl) :rule resolution :premises (t14 t310 t330 t338))
