unsat
(assume a0 (! (and (! (and (and (and (and (and (and (and (or (or (or (or (! (= (! (op e0 e0) :named @p_150) e0) :named @p_155) (! (= @p_150 e1) :named @p_154)) (! (= @p_150 e2) :named @p_153)) (! (= @p_150 e3) :named @p_152)) (! (= @p_150 e4) :named @p_151)) (or (or (or (or (! (= (! (op e0 e1) :named @p_144) e0) :named @p_149) (! (= @p_144 e1) :named @p_148)) (! (= @p_144 e2) :named @p_147)) (! (= @p_144 e3) :named @p_146)) (! (= @p_144 e4) :named @p_145))) (or (or (or (or (! (= (! (op e0 e2) :named @p_138) e0) :named @p_143) (! (= @p_138 e1) :named @p_142)) (! (= @p_138 e2) :named @p_141)) (! (= @p_138 e3) :named @p_140)) (! (= @p_138 e4) :named @p_139))) (or (or (or (or (! (= (! (op e0 e3) :named @p_132) e0) :named @p_137) (! (= @p_132 e1) :named @p_136)) (! (= @p_132 e2) :named @p_135)) (! (= @p_132 e3) :named @p_134)) (! (= @p_132 e4) :named @p_133))) (or (or (or (or (! (= (! (op e0 e4) :named @p_126) e0) :named @p_131) (! (= @p_126 e1) :named @p_130)) (! (= @p_126 e2) :named @p_129)) (! (= @p_126 e3) :named @p_128)) (! (= @p_126 e4) :named @p_127))) (and (and (and (and (or (or (or (or (! (= (! (op e1 e0) :named @p_120) e0) :named @p_125) (! (= @p_120 e1) :named @p_124)) (! (= @p_120 e2) :named @p_123)) (! (= @p_120 e3) :named @p_122)) (! (= @p_120 e4) :named @p_121)) (or (or (or (or (! (= (! (op e1 e1) :named @p_114) e0) :named @p_119) (! (= @p_114 e1) :named @p_118)) (! (= @p_114 e2) :named @p_117)) (! (= @p_114 e3) :named @p_116)) (! (= @p_114 e4) :named @p_115))) (or (or (or (or (! (= (! (op e1 e2) :named @p_108) e0) :named @p_113) (! (= @p_108 e1) :named @p_112)) (! (= @p_108 e2) :named @p_111)) (! (= @p_108 e3) :named @p_110)) (! (= @p_108 e4) :named @p_109))) (or (or (or (or (! (= (! (op e1 e3) :named @p_102) e0) :named @p_107) (! (= @p_102 e1) :named @p_106)) (! (= @p_102 e2) :named @p_105)) (! (= @p_102 e3) :named @p_104)) (! (= @p_102 e4) :named @p_103))) (or (or (or (or (! (= (! (op e1 e4) :named @p_96) e0) :named @p_101) (! (= @p_96 e1) :named @p_100)) (! (= @p_96 e2) :named @p_99)) (! (= @p_96 e3) :named @p_98)) (! (= @p_96 e4) :named @p_97)))) (and (and (and (and (or (or (or (or (! (= (! (op e2 e0) :named @p_90) e0) :named @p_95) (! (= @p_90 e1) :named @p_94)) (! (= @p_90 e2) :named @p_93)) (! (= @p_90 e3) :named @p_92)) (! (= @p_90 e4) :named @p_91)) (or (or (or (or (! (= (! (op e2 e1) :named @p_84) e0) :named @p_89) (! (= @p_84 e1) :named @p_88)) (! (= @p_84 e2) :named @p_87)) (! (= @p_84 e3) :named @p_86)) (! (= @p_84 e4) :named @p_85))) (or (or (or (or (! (= (! (op e2 e2) :named @p_78) e0) :named @p_83) (! (= @p_78 e1) :named @p_82)) (! (= @p_78 e2) :named @p_81)) (! (= @p_78 e3) :named @p_80)) (! (= @p_78 e4) :named @p_79))) (or (or (or (or (! (= (! (op e2 e3) :named @p_72) e0) :named @p_77) (! (= @p_72 e1) :named @p_76)) (! (= @p_72 e2) :named @p_75)) (! (= @p_72 e3) :named @p_74)) (! (= @p_72 e4) :named @p_73))) (or (or (or (or (! (= (! (op e2 e4) :named @p_66) e0) :named @p_71) (! (= @p_66 e1) :named @p_70)) (! (= @p_66 e2) :named @p_69)) (! (= @p_66 e3) :named @p_68)) (! (= @p_66 e4) :named @p_67)))) (! (and (and (and (and (or (or (or (or (! (= (! (op e3 e0) :named @p_59) e0) :named @p_64) (! (= @p_59 e1) :named @p_63)) (! (= @p_59 e2) :named @p_62)) (! (= @p_59 e3) :named @p_61)) (! (= @p_59 e4) :named @p_60)) (or (or (or (or (! (= (! (op e3 e1) :named @p_53) e0) :named @p_58) (! (= @p_53 e1) :named @p_57)) (! (= @p_53 e2) :named @p_56)) (! (= @p_53 e3) :named @p_55)) (! (= @p_53 e4) :named @p_54))) (or (or (or (or (! (= (! (op e3 e2) :named @p_47) e0) :named @p_52) (! (= @p_47 e1) :named @p_51)) (! (= @p_47 e2) :named @p_50)) (! (= @p_47 e3) :named @p_49)) (! (= @p_47 e4) :named @p_48))) (or (or (or (or (! (= (! (op e3 e3) :named @p_41) e0) :named @p_46) (! (= @p_41 e1) :named @p_45)) (! (= @p_41 e2) :named @p_44)) (! (= @p_41 e3) :named @p_43)) (! (= @p_41 e4) :named @p_42))) (! (or (or (or (or (! (= (! (op e3 e4) :named @p_34) e0) :named @p_39) (! (= @p_34 e1) :named @p_38)) (! (= @p_34 e2) :named @p_37)) (! (= @p_34 e3) :named @p_36)) (! (= @p_34 e4) :named @p_35)) :named @p_40)) :named @p_65)) :named @p_156) (! (and (! (and (and (and (or (or (or (or (! (= (! (op e4 e0) :named @p_26) e0) :named @p_31) (! (= @p_26 e1) :named @p_30)) (! (= @p_26 e2) :named @p_29)) (! (= @p_26 e3) :named @p_28)) (! (= @p_26 e4) :named @p_27)) (or (or (or (or (! (= (! (op e4 e1) :named @p_20) e0) :named @p_25) (! (= @p_20 e1) :named @p_24)) (! (= @p_20 e2) :named @p_23)) (! (= @p_20 e3) :named @p_22)) (! (= @p_20 e4) :named @p_21))) (or (or (or (or (! (= (! (op e4 e2) :named @p_14) e0) :named @p_19) (! (= @p_14 e1) :named @p_18)) (! (= @p_14 e2) :named @p_17)) (! (= @p_14 e3) :named @p_16)) (! (= @p_14 e4) :named @p_15))) (! (or (or (or (or (! (= (! (op e4 e3) :named @p_7) e0) :named @p_12) (! (= @p_7 e1) :named @p_11)) (! (= @p_7 e2) :named @p_10)) (! (= @p_7 e3) :named @p_9)) (! (= @p_7 e4) :named @p_8)) :named @p_13)) :named @p_32) (or (or (or (or (! (= (! (op e4 e4) :named @p_1) e0) :named @p_6) (! (= @p_1 e1) :named @p_5)) (! (= @p_1 e2) :named @p_4)) (! (= @p_1 e3) :named @p_3)) (! (= @p_1 e4) :named @p_2))) :named @p_33)) :named @p_157))
(assume a1 (! (and (and (and (and (and (and (and (and (and (or (or (or (or @p_155 @p_149) @p_143) @p_137) @p_131) (or (or (or (or @p_155 @p_125) @p_95) @p_64) @p_31)) (and (or (or (or (or @p_154 @p_148) @p_142) @p_136) @p_130) (or (or (or (or @p_154 @p_124) @p_94) @p_63) @p_30))) (and (or (or (or (or @p_153 @p_147) @p_141) @p_135) @p_129) (or (or (or (or @p_153 @p_123) @p_93) @p_62) @p_29))) (and (or (or (or (or @p_152 @p_146) @p_140) @p_134) @p_128) (or (or (or (or @p_152 @p_122) @p_92) @p_61) @p_28))) (and (or (or (or (or @p_151 @p_145) @p_139) @p_133) @p_127) (or (or (or (or @p_151 @p_121) @p_91) @p_60) @p_27))) (and (and (and (and (and (or (or (or (or @p_125 @p_119) @p_113) @p_107) @p_101) (or (or (or (or @p_149 @p_119) @p_89) @p_58) @p_25)) (and (or (or (or (or @p_124 @p_118) @p_112) @p_106) @p_100) (or (or (or (or @p_148 @p_118) @p_88) @p_57) @p_24))) (and (or (or (or (or @p_123 @p_117) @p_111) @p_105) @p_99) (or (or (or (or @p_147 @p_117) @p_87) @p_56) @p_23))) (and (or (or (or (or @p_122 @p_116) @p_110) @p_104) @p_98) (or (or (or (or @p_146 @p_116) @p_86) @p_55) @p_22))) (and (or (or (or (or @p_121 @p_115) @p_109) @p_103) @p_97) (or (or (or (or @p_145 @p_115) @p_85) @p_54) @p_21)))) (and (and (and (and (and (or (or (or (or @p_95 @p_89) @p_83) @p_77) @p_71) (or (or (or (or @p_143 @p_113) @p_83) @p_52) @p_19)) (and (or (or (or (or @p_94 @p_88) @p_82) @p_76) @p_70) (or (or (or (or @p_142 @p_112) @p_82) @p_51) @p_18))) (and (or (or (or (or @p_93 @p_87) @p_81) @p_75) @p_69) (or (or (or (or @p_141 @p_111) @p_81) @p_50) @p_17))) (and (or (or (or (or @p_92 @p_86) @p_80) @p_74) @p_68) (or (or (or (or @p_140 @p_110) @p_80) @p_49) @p_16))) (and (or (or (or (or @p_91 @p_85) @p_79) @p_73) @p_67) (or (or (or (or @p_139 @p_109) @p_79) @p_48) @p_15)))) (and (and (and (and (and (or (or (or (or @p_64 @p_58) @p_52) @p_46) @p_39) (or (or (or (or @p_137 @p_107) @p_77) @p_46) @p_12)) (and (or (or (or (or @p_63 @p_57) @p_51) @p_45) @p_38) (or (or (or (or @p_136 @p_106) @p_76) @p_45) @p_11))) (and (or (or (or (or @p_62 @p_56) @p_50) @p_44) @p_37) (or (or (or (or @p_135 @p_105) @p_75) @p_44) @p_10))) (and (or (or (or (or @p_61 @p_55) @p_49) @p_43) @p_36) (or (or (or (or @p_134 @p_104) @p_74) @p_43) @p_9))) (and (or (or (or (or @p_60 @p_54) @p_48) @p_42) @p_35) (or (or (or (or @p_133 @p_103) @p_73) @p_42) @p_8)))) (! (and (! (and (and (and (and (or (or (or (or @p_31 @p_25) @p_19) @p_12) @p_6) (or (or (or (or @p_131 @p_101) @p_71) @p_39) @p_6)) (and (or (or (or (or @p_30 @p_24) @p_18) @p_11) @p_5) (or (or (or (or @p_130 @p_100) @p_70) @p_38) @p_5))) (and (or (or (or (or @p_29 @p_23) @p_17) @p_10) @p_4) (or (or (or (or @p_129 @p_99) @p_69) @p_37) @p_4))) (! (and (! (or (or (or (or @p_28 @p_22) @p_16) @p_9) @p_3) :named @p_159) (! (or (or (or (or @p_128 @p_98) @p_68) @p_36) @p_3) :named @p_158)) :named @p_160)) :named @p_161) (and (or (or (or (or @p_27 @p_21) @p_15) @p_8) @p_2) (or (or (or (or @p_127 @p_97) @p_67) @p_35) @p_2))) :named @p_162)) :named @p_163))
(assume a2 (! (and (! (and (! (and (! (and (or (or (or (or @p_155 @p_119) @p_83) @p_46) @p_6) (! (or (or (or (or @p_154 @p_118) @p_82) @p_45) @p_5) :named @p_165)) :named @p_166) (or (or (or (or @p_153 @p_117) @p_81) @p_44) @p_4)) :named @p_167) (! (or (or (or (or @p_152 @p_116) @p_80) @p_43) @p_3) :named @p_164)) :named @p_168) (or (or (or (or @p_151 @p_115) @p_79) @p_42) @p_2)) :named @p_169))
(assume a3 (and (and (and (and (or (or (or (or @p_155 @p_124) @p_93) @p_61) @p_27) (or (or (or (or @p_149 @p_118) @p_87) @p_55) @p_21)) (or (or (or (or @p_143 @p_112) @p_81) @p_49) @p_15)) (or (or (or (or @p_137 @p_106) @p_75) @p_43) @p_8)) (or (or (or (or @p_131 @p_100) @p_69) @p_36) @p_2)))
(assume a4 (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (= @p_150 @p_150) (! (= @p_120 @p_144) :named @p_196)) :named @p_197) (! (= @p_90 @p_138) :named @p_195)) :named @p_198) (! (= @p_59 @p_132) :named @p_194)) :named @p_199) (! (= @p_26 @p_126) :named @p_193)) :named @p_200) (! (and (! (and (! (and (and (! (= @p_144 @p_120) :named @p_189) (! (= @p_114 @p_114) :named @p_188)) (! (= @p_84 @p_108) :named @p_187)) :named @p_190) (! (= @p_53 @p_102) :named @p_186)) :named @p_191) (! (= @p_20 @p_96) :named @p_185)) :named @p_192)) :named @p_201) (! (and (! (and (and (and (! (= @p_138 @p_90) :named @p_182) (! (= @p_108 @p_84) :named @p_181)) (= @p_78 @p_78)) (! (= @p_47 @p_72) :named @p_180)) :named @p_183) (! (= @p_14 @p_66) :named @p_179)) :named @p_184)) :named @p_202) (! (and (and (and (and (! (= @p_132 @p_59) :named @p_177) (! (= @p_102 @p_53) :named @p_176)) (! (= @p_72 @p_47) :named @p_175)) (= @p_41 @p_41)) (! (= @p_7 @p_34) :named @p_174)) :named @p_178)) :named @p_203) (and (and (and (and (! (= @p_126 @p_26) :named @p_173) (! (= @p_96 @p_20) :named @p_172)) (! (= @p_66 @p_14) :named @p_171)) (! (= @p_34 @p_7) :named @p_170)) (= @p_1 @p_1))) :named @p_204))
(assume a5 (or (or (or (or @p_155 @p_118) @p_81) @p_43) @p_2))
(assume a6 (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (= (! (op e0 @p_150) :named @p_226) e0) :named @p_227) (= (op e0 @p_144) e1)) :named @p_228) (! (= (! (op e0 @p_138) :named @p_224) e2) :named @p_225)) :named @p_229) (= (op e0 @p_132) e3)) :named @p_230) (= (op e0 @p_126) e4)) :named @p_231) (! (and (! (and (! (and (! (and (= (op e1 @p_120) e0) (! (= (! (op e1 @p_114) :named @p_218) e1) :named @p_219)) :named @p_220) (= (op e1 @p_108) e2)) :named @p_221) (= (op e1 @p_102) e3)) :named @p_222) (! (= (! (op e1 @p_96) :named @p_216) e4) :named @p_217)) :named @p_223)) :named @p_232) (! (and (! (and (! (and (! (and (! (= (! (op e2 @p_90) :named @p_210) e0) :named @p_211) (= (op e2 @p_84) e1)) :named @p_212) (! (= (! (op e2 @p_78) :named @p_208) e2) :named @p_209)) :named @p_213) (= (op e2 @p_72) e3)) :named @p_214) (= (op e2 @p_66) e4)) :named @p_215)) :named @p_233) (and (and (and (and (= (op e3 @p_59) e0) (= (op e3 @p_53) e1)) (= (op e3 @p_47) e2)) (= (op e3 @p_41) e3)) (= (op e3 @p_34) e4))) :named @p_234) (! (and (and (and (and (= (op e4 @p_26) e0) (= (op e4 @p_20) e1)) (= (op e4 @p_14) e2)) (= (op e4 @p_7) e3)) (! (= (! (op e4 @p_1) :named @p_205) e4) :named @p_206)) :named @p_207)) :named @p_235))
(assume a7 (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (not (= @p_150 @p_120)) (! (not (! (= @p_150 @p_90) :named @p_293)) :named @p_294)) :named @p_295) (not (= @p_150 @p_59))) :named @p_296) (not (= @p_150 @p_26))) :named @p_297) (not (= @p_120 @p_90))) :named @p_298) (not (= @p_120 @p_59))) :named @p_299) (not (= @p_120 @p_26))) :named @p_300) (! (not (! (= @p_90 @p_59) :named @p_291)) :named @p_292)) :named @p_301) (! (not (! (= @p_90 @p_26) :named @p_289)) :named @p_290)) :named @p_302) (not (= @p_59 @p_26))) :named @p_303) (and (and (and (and (and (and (and (and (and (not (= @p_144 @p_114)) (not (= @p_144 @p_84))) (not (= @p_144 @p_53))) (not (= @p_144 @p_20))) (not (= @p_114 @p_84))) (not (= @p_114 @p_53))) (not (= @p_114 @p_20))) (not (= @p_84 @p_53))) (not (= @p_84 @p_20))) (not (= @p_53 @p_20)))) :named @p_304) (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (not (! (= @p_138 @p_108) :named @p_278)) :named @p_279) (not (= @p_138 @p_78))) :named @p_280) (not (= @p_138 @p_47))) :named @p_281) (not (= @p_138 @p_14))) :named @p_282) (not (= @p_108 @p_78))) :named @p_283) (not (= @p_108 @p_47))) :named @p_284) (not (= @p_108 @p_14))) :named @p_285) (not (= @p_78 @p_47))) :named @p_286) (not (= @p_78 @p_14))) :named @p_287) (not (= @p_47 @p_14))) :named @p_288)) :named @p_305) (! (and (! (and (! (and (! (and (! (and (! (and (! (and (and (and (not (= @p_132 @p_102)) (not (= @p_132 @p_72))) (not (= @p_132 @p_41))) (! (not (! (= @p_132 @p_7) :named @p_269)) :named @p_270)) :named @p_271) (not (= @p_102 @p_72))) :named @p_272) (not (= @p_102 @p_41))) :named @p_273) (not (= @p_102 @p_7))) :named @p_274) (not (= @p_72 @p_41))) :named @p_275) (not (= @p_72 @p_7))) :named @p_276) (! (not (! (= @p_41 @p_7) :named @p_267)) :named @p_268)) :named @p_277)) :named @p_306) (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (and (not (= @p_126 @p_96)) (not (= @p_126 @p_66))) (! (not (! (= @p_126 @p_34) :named @p_257)) :named @p_258)) :named @p_259) (not (= @p_126 @p_1))) :named @p_260) (not (= @p_96 @p_66))) :named @p_261) (not (= @p_96 @p_34))) :named @p_262) (not (= @p_96 @p_1))) :named @p_263) (not (= @p_66 @p_34))) :named @p_264) (not (= @p_66 @p_1))) :named @p_265) (! (not (! (= @p_34 @p_1) :named @p_255)) :named @p_256)) :named @p_266)) :named @p_307) (! (and (! (and (! (and (and (and (and (and (and (and (and (and (and (and (not (= @p_150 @p_144)) (not (= @p_150 @p_138))) (not (= @p_150 @p_132))) (not (= @p_150 @p_126))) (not (= @p_144 @p_138))) (not (= @p_144 @p_132))) (not (= @p_144 @p_126))) (not (= @p_138 @p_132))) (not (= @p_138 @p_126))) (not (= @p_132 @p_126))) (and (and (and (and (and (and (and (and (and (not (= @p_120 @p_114)) (not (= @p_120 @p_108))) (not (= @p_120 @p_102))) (not (= @p_120 @p_96))) (not (= @p_114 @p_108))) (not (= @p_114 @p_102))) (not (= @p_114 @p_96))) (not (= @p_108 @p_102))) (not (= @p_108 @p_96))) (not (= @p_102 @p_96)))) (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (not (= @p_90 @p_84)) (! (not (! (= @p_90 @p_78) :named @p_241)) :named @p_242)) :named @p_243) (not (= @p_90 @p_72))) :named @p_244) (! (not (! (= @p_90 @p_66) :named @p_239)) :named @p_240)) :named @p_245) (not (= @p_84 @p_78))) :named @p_246) (not (= @p_84 @p_72))) :named @p_247) (not (= @p_84 @p_66))) :named @p_248) (not (= @p_78 @p_72))) :named @p_249) (not (= @p_78 @p_66))) :named @p_250) (not (= @p_72 @p_66))) :named @p_251)) :named @p_252) (and (and (and (and (and (and (and (and (and (not (= @p_59 @p_53)) (not (= @p_59 @p_47))) (not (= @p_59 @p_41))) (not (= @p_59 @p_34))) (not (= @p_53 @p_47))) (not (= @p_53 @p_41))) (not (= @p_53 @p_34))) (not (= @p_47 @p_41))) (not (= @p_47 @p_34))) (not (= @p_41 @p_34)))) :named @p_253) (! (and (and (and (and (and (and (and (and (and (not (= @p_26 @p_20)) (not (= @p_26 @p_14))) (not (= @p_26 @p_7))) (not (= @p_26 @p_1))) (not (= @p_20 @p_14))) (not (= @p_20 @p_7))) (not (= @p_20 @p_1))) (not (= @p_14 @p_7))) (not (= @p_14 @p_1))) (! (not (! (= @p_7 @p_1) :named @p_236)) :named @p_237)) :named @p_238)) :named @p_254)) :named @p_308))
(assume a8 (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (not (! (= e0 e1) :named @p_317)) :named @p_318) (not (= e0 e2))) :named @p_319) (! (not (! (= e0 e3) :named @p_315)) :named @p_316)) :named @p_320) (not (= e0 e4))) :named @p_321) (! (not (! (= e1 e2) :named @p_313)) :named @p_314)) :named @p_322) (not (= e1 e3))) :named @p_323) (! (not (! (= e1 e4) :named @p_311)) :named @p_312)) :named @p_324) (! (not (! (= e2 e3) :named @p_309)) :named @p_310)) :named @p_325) (not (= e2 e4))) :named @p_326) (not (= e3 e4))) :named @p_327))
(assume a9 (! (and (! (and (! (and (! (= e0 @p_1) :named @p_331) (! (= e2 @p_114) :named @p_330)) :named @p_332) (! (= e3 @p_90) :named @p_329)) :named @p_333) (! (= e4 @p_78) :named @p_328)) :named @p_334))
(assume a10 (! (not (! (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and @p_154 @p_149) @p_140) @p_135) @p_127) @p_125) @p_117) @p_112) @p_103) @p_98) @p_92) @p_88) @p_79) @p_77) @p_69) @p_62) @p_54) @p_52) @p_43) @p_38) @p_27) @p_22) @p_17) @p_11) @p_6) :named @p_335)) :named @p_336))
(step t0 (cl (not (! (= @p_336 (! (not (! (and (! (= e1 (! (op (! (op @p_114 @p_114) :named @p_337) (! (op @p_114 (! (op @p_337 @p_337) :named @p_338)) :named @p_340)) :named @p_398)) :named @p_399) (! (= @p_114 (! (op @p_337 @p_114) :named @p_454)) :named @p_455) (! (= @p_340 (! (op @p_337 e1) :named @p_459)) :named @p_460) (! (= (! (op @p_114 @p_114) :named @p_337) (! (op @p_337 @p_338) :named @p_355)) :named @p_356) (! (= e1 (! (op @p_340 @p_337) :named @p_403)) :named @p_485) (! (= @p_340 (! (op @p_340 @p_340) :named @p_359)) :named @p_360) (! (= @p_338 (! (op @p_340 @p_114) :named @p_509)) :named @p_510) (! (= (! (op @p_114 @p_114) :named @p_337) (! (op @p_340 e1) :named @p_514)) :named @p_515) (! (= @p_114 (! (op @p_340 @p_338) :named @p_464)) :named @p_528) (! (= @p_114 (! (op @p_114 @p_337) :named @p_351)) :named @p_352) (! (= @p_338 (! (op @p_114 @p_340) :named @p_347)) :named @p_348) (! (= e1 (! (op @p_114 e1) :named @p_529)) :named @p_530) (! (= @p_340 (! (op e1 @p_337) :named @p_520)) :named @p_533) (! (= (! (op @p_114 @p_114) :named @p_337) (! (op e1 @p_340) :named @p_516)) :named @p_550) (! (= e1 @p_218) :named @p_346) (! (= @p_338 (! (op e1 @p_338) :named @p_588)) :named @p_589) (! (= (! (op @p_114 @p_114) :named @p_337) (! (op @p_338 @p_337) :named @p_404)) :named @p_413) (! (= @p_114 (! (op @p_338 @p_340) :named @p_428)) :named @p_434) (! (= @p_340 (! (op @p_338 @p_114) :named @p_339)) :named @p_341) (! (= e1 (! (op @p_338 @p_338) :named @p_383)) :named @p_551) (! (= @p_338 (! (op @p_338 e1) :named @p_590)) :named @p_598)) :named @p_620)) :named @p_621)) :named @p_622)) (not @p_336) @p_621) :rule equiv_pos2)
(step t1 (cl (= @p_335 (! (and @p_6 @p_11 @p_17 @p_22 @p_27 @p_38 @p_43 @p_52 @p_54 @p_62 @p_69 @p_77 @p_79 @p_88 @p_92 @p_98 @p_103 @p_112 @p_117 @p_125 @p_127 @p_135 @p_140 @p_154 @p_149) :named @p_644))) :rule all_simplify)
(step t2 (cl (= @p_6 @p_331)) :rule all_simplify)
(step t3 (cl (= @p_11 (! (= e1 @p_7) :named @p_447))) :rule all_simplify)
(step t4 (cl (= @p_17 (! (= e2 @p_14) :named @p_637))) :rule all_simplify)
(step t5 (cl (= @p_22 (! (= e3 @p_20) :named @p_479))) :rule all_simplify)
(step t6 (cl (= @p_27 (! (= e4 @p_26) :named @p_636))) :rule all_simplify)
(step t7 (cl (= @p_38 (! (= e1 @p_34) :named @p_502))) :rule all_simplify)
(step t8 (cl (= @p_43 (! (= e3 @p_41) :named @p_392))) :rule all_simplify)
(step t9 (cl (= @p_52 (! (= e0 @p_47) :named @p_635))) :rule all_simplify)
(step t10 (cl (= @p_54 (! (= e4 @p_53) :named @p_634))) :rule all_simplify)
(step t11 (cl (= @p_62 (! (= e2 @p_59) :named @p_633))) :rule all_simplify)
(step t12 (cl (= @p_69 (! (= e2 @p_66) :named @p_632))) :rule all_simplify)
(step t13 (cl (= @p_77 (! (= e0 @p_72) :named @p_631))) :rule all_simplify)
(step t14 (cl (= @p_79 @p_328)) :rule all_simplify)
(step t15 (cl (= @p_88 (! (= e1 @p_84) :named @p_630))) :rule all_simplify)
(step t16 (cl (= @p_92 @p_329)) :rule all_simplify)
(step t17 (cl (= @p_98 (! (= e3 @p_96) :named @p_544))) :rule all_simplify)
(step t18 (cl (= @p_103 (! (= e4 @p_102) :named @p_629))) :rule all_simplify)
(step t19 (cl (= @p_112 (! (= e1 @p_108) :named @p_628))) :rule all_simplify)
(step t20 (cl (= @p_117 @p_330)) :rule all_simplify)
(step t21 (cl (= @p_125 (! (= e0 @p_120) :named @p_627))) :rule all_simplify)
(step t22 (cl (= @p_127 (! (= e4 @p_126) :named @p_626))) :rule all_simplify)
(step t23 (cl (= @p_135 (! (= e2 @p_132) :named @p_625))) :rule all_simplify)
(step t24 (cl (= @p_140 (! (= e3 @p_138) :named @p_624))) :rule all_simplify)
(step t25 (cl (= @p_154 (! (= e1 @p_150) :named @p_582))) :rule all_simplify)
(step t26 (cl (= @p_149 (! (= e0 @p_144) :named @p_623))) :rule all_simplify)
(step t27 (cl (= @p_644 (! (and @p_331 @p_447 @p_637 @p_479 @p_636 @p_502 @p_392 @p_635 @p_634 @p_633 @p_632 @p_631 @p_328 @p_630 @p_329 @p_544 @p_629 @p_628 @p_330 @p_627 @p_626 @p_625 @p_624 @p_582 @p_623) :named @p_638))) :rule cong :premises (t2 t3 t4 t5 t6 t7 t8 t9 t10 t11 t12 t13 t14 t15 t16 t17 t18 t19 t20 t21 t22 t23 t24 t25 t26))
(step t28 (cl (= @p_335 @p_638)) :rule trans :premises (t1 t27))
(step t29 (cl (= @p_336 (! (not @p_638) :named @p_639))) :rule cong :premises (t28))
(step t30 (cl @p_333) :rule and :premises (a9))
(step t31 (cl @p_332) :rule and :premises (t30))
(step t32 (cl @p_331) :rule and :premises (t31))
(step t33 (cl (not (! (= @p_328 (! (= e4 @p_337) :named @p_344)) :named @p_345)) (not @p_328) @p_344) :rule equiv_pos2)
(step t34 (cl (= e4 e4)) :rule refl)
(step t35 (cl @p_330) :rule and :premises (t31))
(step t36 (cl (= @p_78 @p_337)) :rule cong :premises (t35 t35))
(step t37 (cl @p_345) :rule cong :premises (t34 t36))
(step t38 (cl @p_328) :rule and :premises (a9))
(step t39 (cl @p_344) :rule resolution :premises (t33 t37 t38))
(step t40 (cl (= @p_1 @p_338)) :rule cong :premises (t39 t39))
(step t41 (cl (= e0 @p_338)) :rule trans :premises (t32 t40))
(step t42 (cl (= @p_331 (! (= @p_338 @p_338) :named @p_436))) :rule cong :premises (t41 t40))
(step t43 (cl (= e1 e1)) :rule refl)
(step t44 (cl (not (! (= @p_329 (! (= e3 (! (op @p_114 @p_1) :named @p_368)) :named @p_369)) :named @p_370)) (not @p_329) @p_369) :rule equiv_pos2)
(step t45 (cl (= e3 e3)) :rule refl)
(step t46 (cl (= @p_90 @p_368)) :rule cong :premises (t35 t32))
(step t47 (cl @p_370) :rule cong :premises (t45 t46))
(step t48 (cl @p_329) :rule and :premises (t30))
(step t49 (cl @p_369) :rule resolution :premises (t44 t47 t48))
(step t50 (cl @p_188) :rule refl)
(step t51 (cl (= @p_368 @p_340)) :rule cong :premises (t50 t40))
(step t52 (cl (= e3 @p_340)) :rule trans :premises (t49 t51))
(step t53 (cl (= @p_7 @p_398)) :rule cong :premises (t39 t52))
(step t54 (cl (= @p_447 @p_399)) :rule cong :premises (t43 t53))
(step t55 (cl (= @p_14 @p_454)) :rule cong :premises (t39 t35))
(step t56 (cl (= @p_637 @p_455)) :rule cong :premises (t35 t55))
(step t57 (cl (= @p_20 @p_459)) :rule cong :premises (t39 t43))
(step t58 (cl (= @p_479 @p_460)) :rule cong :premises (t52 t57))
(step t59 (cl (= @p_26 @p_355)) :rule cong :premises (t39 t41))
(step t60 (cl (= @p_636 @p_356)) :rule cong :premises (t39 t59))
(step t61 (cl (= @p_34 @p_403)) :rule cong :premises (t52 t39))
(step t62 (cl (= @p_502 @p_485)) :rule cong :premises (t43 t61))
(step t63 (cl (= @p_41 @p_359)) :rule cong :premises (t52 t52))
(step t64 (cl (= @p_392 @p_360)) :rule cong :premises (t52 t63))
(step t65 (cl (= @p_47 @p_509)) :rule cong :premises (t52 t35))
(step t66 (cl (= @p_635 @p_510)) :rule cong :premises (t41 t65))
(step t67 (cl (= @p_53 @p_514)) :rule cong :premises (t52 t43))
(step t68 (cl (= @p_634 @p_515)) :rule cong :premises (t39 t67))
(step t69 (cl (= @p_59 @p_464)) :rule cong :premises (t52 t41))
(step t70 (cl (= @p_633 @p_528)) :rule cong :premises (t35 t69))
(step t71 (cl (= @p_66 @p_351)) :rule cong :premises (t35 t39))
(step t72 (cl (= @p_632 @p_352)) :rule cong :premises (t35 t71))
(step t73 (cl (= @p_72 @p_347)) :rule cong :premises (t35 t52))
(step t74 (cl (= @p_631 @p_348)) :rule cong :premises (t41 t73))
(step t75 (cl (= @p_328 (! (= @p_337 @p_337) :named @p_640))) :rule cong :premises (t39 t36))
(step t76 (cl (= @p_84 @p_529)) :rule cong :premises (t35 t43))
(step t77 (cl (= @p_630 @p_530)) :rule cong :premises (t43 t76))
(step t78 (cl (= @p_90 @p_340)) :rule cong :premises (t35 t41))
(step t79 (cl (= @p_329 (! (= @p_340 @p_340) :named @p_473))) :rule cong :premises (t52 t78))
(step t80 (cl (= @p_96 @p_520)) :rule cong :premises (t43 t39))
(step t81 (cl (= @p_544 @p_533)) :rule cong :premises (t52 t80))
(step t82 (cl (= @p_102 @p_516)) :rule cong :premises (t43 t52))
(step t83 (cl (= @p_629 @p_550)) :rule cong :premises (t39 t82))
(step t84 (cl (= @p_108 @p_218)) :rule cong :premises (t43 t35))
(step t85 (cl (= @p_628 @p_346)) :rule cong :premises (t43 t84))
(step t86 (cl (= @p_330 @p_188)) :rule cong :premises (t35 t50))
(step t87 (cl (= @p_120 @p_588)) :rule cong :premises (t43 t41))
(step t88 (cl (= @p_627 @p_589)) :rule cong :premises (t41 t87))
(step t89 (cl (= @p_126 @p_404)) :rule cong :premises (t41 t39))
(step t90 (cl (= @p_626 @p_413)) :rule cong :premises (t39 t89))
(step t91 (cl (= @p_132 @p_428)) :rule cong :premises (t41 t52))
(step t92 (cl (= @p_625 @p_434)) :rule cong :premises (t35 t91))
(step t93 (cl (= @p_138 @p_339)) :rule cong :premises (t41 t35))
(step t94 (cl (= @p_624 @p_341)) :rule cong :premises (t52 t93))
(step t95 (cl (= @p_150 @p_383)) :rule cong :premises (t41 t41))
(step t96 (cl (= @p_582 @p_551)) :rule cong :premises (t43 t95))
(step t97 (cl (= @p_144 @p_590)) :rule cong :premises (t41 t43))
(step t98 (cl (= @p_623 @p_598)) :rule cong :premises (t41 t97))
(step t99 (cl (= @p_638 (! (and @p_436 @p_399 @p_455 @p_460 @p_356 @p_485 @p_360 @p_510 @p_515 @p_528 @p_352 @p_348 @p_640 @p_530 @p_473 @p_533 @p_550 @p_346 @p_188 @p_589 @p_413 @p_434 @p_341 @p_551 @p_598) :named @p_641))) :rule cong :premises (t42 t54 t56 t58 t60 t62 t64 t66 t68 t70 t72 t74 t75 t77 t79 t81 t83 t85 t86 t88 t90 t92 t94 t96 t98))
(step t100 (cl (= @p_639 (! (not @p_641) :named @p_642))) :rule cong :premises (t99))
(step t101 (cl (= @p_436 true)) :rule all_simplify)
(step t102 (cl (= @p_399 @p_399)) :rule refl)
(step t103 (cl (= @p_455 @p_455)) :rule refl)
(step t104 (cl (= @p_460 @p_460)) :rule refl)
(step t105 (cl (= @p_356 @p_356)) :rule refl)
(step t106 (cl (= @p_485 @p_485)) :rule refl)
(step t107 (cl (= @p_360 @p_360)) :rule refl)
(step t108 (cl (= @p_510 @p_510)) :rule refl)
(step t109 (cl (= @p_515 @p_515)) :rule refl)
(step t110 (cl (= @p_528 @p_528)) :rule refl)
(step t111 (cl (= @p_352 @p_352)) :rule refl)
(step t112 (cl (= @p_348 @p_348)) :rule refl)
(step t113 (cl (= @p_640 true)) :rule all_simplify)
(step t114 (cl (= @p_530 @p_530)) :rule refl)
(step t115 (cl (= @p_473 true)) :rule all_simplify)
(step t116 (cl (= @p_533 @p_533)) :rule refl)
(step t117 (cl (= @p_550 @p_550)) :rule refl)
(step t118 (cl (= @p_346 @p_346)) :rule refl)
(step t119 (cl (= @p_188 true)) :rule all_simplify)
(step t120 (cl (= @p_589 @p_589)) :rule refl)
(step t121 (cl (= @p_413 @p_413)) :rule refl)
(step t122 (cl (= @p_434 @p_434)) :rule refl)
(step t123 (cl (= @p_341 @p_341)) :rule refl)
(step t124 (cl (= @p_551 @p_551)) :rule refl)
(step t125 (cl (= @p_598 @p_598)) :rule refl)
(step t126 (cl (= @p_641 (! (and true @p_399 @p_455 @p_460 @p_356 @p_485 @p_360 @p_510 @p_515 @p_528 @p_352 @p_348 true @p_530 true @p_533 @p_550 @p_346 true @p_589 @p_413 @p_434 @p_341 @p_551 @p_598) :named @p_643))) :rule cong :premises (t101 t102 t103 t104 t105 t106 t107 t108 t109 t110 t111 t112 t113 t114 t115 t116 t117 t118 t119 t120 t121 t122 t123 t124 t125))
(step t127 (cl (= @p_643 @p_620)) :rule all_simplify)
(step t128 (cl (= @p_641 @p_620)) :rule trans :premises (t126 t127))
(step t129 (cl (= @p_642 @p_621)) :rule cong :premises (t128))
(step t130 (cl (= @p_639 @p_621)) :rule trans :premises (t100 t129))
(step t131 (cl @p_622) :rule trans :premises (t29 t130))
(step t132 (cl @p_621) :rule resolution :premises (t0 t131 a10))
(step t133 (cl (! (not @p_399) :named @p_619) (! (not @p_455) :named @p_618) (! (not @p_460) :named @p_617) (! (not @p_356) :named @p_616) (! (not @p_485) :named @p_615) (! (not @p_360) :named @p_614) (! (not @p_510) :named @p_613) (! (not @p_515) :named @p_612) (! (not @p_528) :named @p_611) (! (not @p_352) :named @p_610) (! (not @p_348) :named @p_609) (! (not @p_530) :named @p_608) (! (not @p_533) :named @p_607) (! (not @p_550) :named @p_606) (! (not @p_346) :named @p_605) (! (not @p_589) :named @p_604) (! (not @p_413) :named @p_603) (! (not @p_434) :named @p_602) (! (not @p_341) :named @p_601) (! (not @p_551) :named @p_599) (! (not @p_598) :named @p_600)) :rule not_and :premises (t132))
(step t134 (cl @p_599 @p_600 @p_601 @p_602 @p_603 @p_604 @p_605 @p_606 @p_607 @p_608 @p_609 @p_610 @p_611 @p_612 @p_613 @p_614 @p_615 @p_616 @p_617 @p_618 @p_619) :rule reordering :premises (t133))
(step t135 (cl (not (! (= @p_227 (! (= @p_338 (! (op @p_338 @p_383) :named @p_594)) :named @p_595)) :named @p_596)) (not @p_227) @p_595) :rule equiv_pos2)
(step t136 (cl (= @p_227 (! (= e0 @p_226) :named @p_597))) :rule all_simplify)
(step t137 (cl (= @p_226 @p_594)) :rule cong :premises (t41 t95))
(step t138 (cl (= @p_597 @p_595)) :rule cong :premises (t41 t137))
(step t139 (cl @p_596) :rule trans :premises (t136 t138))
(step t140 (cl @p_234) :rule and :premises (a6))
(step t141 (cl @p_233) :rule and :premises (t140))
(step t142 (cl @p_232) :rule and :premises (t141))
(step t143 (cl @p_231) :rule and :premises (t142))
(step t144 (cl @p_230) :rule and :premises (t143))
(step t145 (cl @p_229) :rule and :premises (t144))
(step t146 (cl @p_228) :rule and :premises (t145))
(step t147 (cl @p_227) :rule and :premises (t146))
(step t148 (cl @p_595) :rule resolution :premises (t135 t139 t147))
(step t149 (cl @p_436) :rule refl)
(step t150 (cl (not (! (= @p_165 (! (or (! (= e1 @p_338) :named @p_569) (! (= e1 @p_359) :named @p_552) (! (= e1 @p_337) :named @p_562) @p_551 (! (= e1 @p_114) :named @p_574)) :named @p_580)) :named @p_581)) (not @p_165) @p_580) :rule equiv_pos2)
(step t151 (cl (= @p_165 (! (or @p_5 @p_45 @p_82 @p_154 @p_118) :named @p_587))) :rule all_simplify)
(step t152 (cl (= @p_5 (! (= e1 @p_1) :named @p_585))) :rule all_simplify)
(step t153 (cl (= @p_45 (! (= e1 @p_41) :named @p_584))) :rule all_simplify)
(step t154 (cl (= @p_82 (! (= e1 @p_78) :named @p_583))) :rule all_simplify)
(step t155 (cl (= @p_118 @p_574)) :rule all_simplify)
(step t156 (cl (= @p_587 (! (or @p_585 @p_584 @p_583 @p_582 @p_574) :named @p_586))) :rule cong :premises (t152 t153 t154 t25 t155))
(step t157 (cl (= @p_165 @p_586)) :rule trans :premises (t151 t156))
(step t158 (cl (= @p_585 @p_569)) :rule cong :premises (t43 t40))
(step t159 (cl (= @p_584 @p_552)) :rule cong :premises (t43 t63))
(step t160 (cl (= @p_583 @p_562)) :rule cong :premises (t43 t36))
(step t161 (cl (= @p_574 @p_574)) :rule refl)
(step t162 (cl (= @p_586 @p_580)) :rule cong :premises (t158 t159 t160 t96 t161))
(step t163 (cl @p_581) :rule trans :premises (t157 t162))
(step t164 (cl @p_168) :rule and :premises (a2))
(step t165 (cl @p_167) :rule and :premises (t164))
(step t166 (cl @p_166) :rule and :premises (t165))
(step t167 (cl @p_165) :rule and :premises (t166))
(step t168 (cl @p_580) :rule resolution :premises (t150 t163 t167))
(step t169 (cl @p_569 @p_552 @p_562 @p_551 @p_574) :rule or :premises (t168))
(step t170 (cl @p_574 @p_569 @p_562 @p_551 @p_552) :rule reordering :premises (t169))
(step t171 (cl (not (! (= @p_314 (! (not @p_574) :named @p_575)) :named @p_576)) (not @p_314) @p_575) :rule equiv_pos2)
(step t172 (cl (= @p_313 (! (= e2 e1) :named @p_577))) :rule all_simplify)
(step t173 (cl (= @p_314 (! (not @p_577) :named @p_578))) :rule cong :premises (t172))
(step t174 (cl (= @p_577 @p_118)) :rule cong :premises (t35 t43))
(step t175 (cl (= @p_578 (! (not @p_118) :named @p_579))) :rule cong :premises (t174))
(step t176 (cl (= @p_579 @p_575)) :rule cong :premises (t155))
(step t177 (cl (= @p_578 @p_575)) :rule trans :premises (t175 t176))
(step t178 (cl @p_576) :rule trans :premises (t173 t177))
(step t179 (cl @p_326) :rule and :premises (a8))
(step t180 (cl @p_325) :rule and :premises (t179))
(step t181 (cl @p_324) :rule and :premises (t180))
(step t182 (cl @p_323) :rule and :premises (t181))
(step t183 (cl @p_322) :rule and :premises (t182))
(step t184 (cl @p_314) :rule and :premises (t183))
(step t185 (cl @p_575) :rule resolution :premises (t171 t178 t184))
(step t186 (cl (not (! (= @p_318 (! (not @p_569) :named @p_570)) :named @p_571)) (not @p_318) @p_570) :rule equiv_pos2)
(step t187 (cl (= @p_317 (! (= e1 e0) :named @p_572))) :rule all_simplify)
(step t188 (cl (= @p_318 (! (not @p_572) :named @p_573))) :rule cong :premises (t187))
(step t189 (cl (= @p_572 @p_569)) :rule cong :premises (t43 t41))
(step t190 (cl (= @p_573 @p_570)) :rule cong :premises (t189))
(step t191 (cl @p_571) :rule trans :premises (t188 t190))
(step t192 (cl @p_321) :rule and :premises (t183))
(step t193 (cl @p_320) :rule and :premises (t192))
(step t194 (cl @p_319) :rule and :premises (t193))
(step t195 (cl @p_318) :rule and :premises (t194))
(step t196 (cl @p_570) :rule resolution :premises (t186 t191 t195))
(step t197 (cl (not (! (= @p_312 (! (not @p_562) :named @p_563)) :named @p_564)) (not @p_312) @p_563) :rule equiv_pos2)
(step t198 (cl (= @p_311 (! (= e4 e1) :named @p_565))) :rule all_simplify)
(step t199 (cl (= @p_312 (! (not @p_565) :named @p_566))) :rule cong :premises (t198))
(step t200 (cl (= @p_565 (! (= @p_337 e1) :named @p_567))) :rule cong :premises (t39 t43))
(step t201 (cl (= @p_566 (! (not @p_567) :named @p_568))) :rule cong :premises (t200))
(step t202 (cl (= @p_567 @p_562)) :rule all_simplify)
(step t203 (cl (= @p_568 @p_563)) :rule cong :premises (t202))
(step t204 (cl (= @p_566 @p_563)) :rule trans :premises (t201 t203))
(step t205 (cl @p_564) :rule trans :premises (t199 t204))
(step t206 (cl @p_312) :rule and :premises (t181))
(step t207 (cl @p_563) :rule resolution :premises (t197 t205 t206))
(step t208 (cl (= (! (= @p_552 false) :named @p_554) (! (not @p_552) :named @p_553))) :rule equiv_simplify)
(step t209 (cl (not @p_554) @p_553) :rule equiv1 :premises (t208))
(step t210 (cl @p_223) :rule and :premises (t142))
(step t211 (cl @p_222) :rule and :premises (t210))
(step t212 (cl @p_221) :rule and :premises (t211))
(step t213 (cl @p_220) :rule and :premises (t212))
(step t214 (cl @p_219) :rule and :premises (t213))
(step t215 (cl @p_346) :rule symm :premises (t214))
(step t216 (cl (not (! (= @p_182 @p_341) :named @p_342)) (not @p_182) @p_341) :rule equiv_pos2)
(step t217 (cl (= @p_182 (! (= @p_339 @p_340) :named @p_343))) :rule cong :premises (t93 t78))
(step t218 (cl (= @p_343 @p_341)) :rule all_simplify)
(step t219 (cl @p_342) :rule trans :premises (t217 t218))
(step t220 (cl @p_203) :rule and :premises (a4))
(step t221 (cl @p_202) :rule and :premises (t220))
(step t222 (cl @p_201) :rule and :premises (t221))
(step t223 (cl @p_200) :rule and :premises (t222))
(step t224 (cl @p_199) :rule and :premises (t223))
(step t225 (cl @p_198) :rule and :premises (t224))
(step t226 (cl @p_195) :rule and :premises (t225))
(step t227 (cl @p_182) :rule symm :premises (t226))
(step t228 (cl @p_341) :rule resolution :premises (t216 t219 t227))
(step t229 (cl @p_343) :rule symm :premises (t228))
(step t230 (cl (not (! (= @p_164 (! (or (! (= @p_338 @p_340) :named @p_371) @p_360 (! (= @p_337 @p_340) :named @p_378) (! (= @p_383 @p_340) :named @p_384) (! (= @p_114 @p_340) :named @p_361)) :named @p_387)) :named @p_388)) (not @p_164) @p_387) :rule equiv_pos2)
(step t231 (cl (= @p_164 (! (or @p_3 @p_43 @p_80 @p_152 @p_116) :named @p_397))) :rule all_simplify)
(step t232 (cl (= @p_3 (! (= e3 @p_1) :named @p_393))) :rule all_simplify)
(step t233 (cl (= @p_80 (! (= e3 @p_78) :named @p_391))) :rule all_simplify)
(step t234 (cl (= @p_152 (! (= e3 @p_150) :named @p_390))) :rule all_simplify)
(step t235 (cl (= @p_116 (! (= e3 @p_114) :named @p_389))) :rule all_simplify)
(step t236 (cl (= @p_397 (! (or @p_393 @p_392 @p_391 @p_390 @p_389) :named @p_394))) :rule cong :premises (t232 t8 t233 t234 t235))
(step t237 (cl (= @p_164 @p_394)) :rule trans :premises (t231 t236))
(step t238 (cl (= @p_393 (! (= @p_340 @p_338) :named @p_376))) :rule cong :premises (t52 t40))
(step t239 (cl (= @p_391 (! (= @p_340 @p_337) :named @p_381))) :rule cong :premises (t52 t36))
(step t240 (cl (= @p_390 (! (= @p_340 @p_383) :named @p_395))) :rule cong :premises (t52 t95))
(step t241 (cl (= @p_389 (! (= @p_340 @p_114) :named @p_366))) :rule cong :premises (t52 t50))
(step t242 (cl (= @p_394 (! (or @p_376 @p_360 @p_381 @p_395 @p_366) :named @p_396))) :rule cong :premises (t238 t64 t239 t240 t241))
(step t243 (cl (= @p_376 @p_371)) :rule all_simplify)
(step t244 (cl (= @p_381 @p_378)) :rule all_simplify)
(step t245 (cl (= @p_395 @p_384)) :rule all_simplify)
(step t246 (cl (= @p_366 @p_361)) :rule all_simplify)
(step t247 (cl (= @p_396 @p_387)) :rule cong :premises (t243 t107 t244 t245 t246))
(step t248 (cl (= @p_394 @p_387)) :rule trans :premises (t242 t247))
(step t249 (cl @p_388) :rule trans :premises (t237 t248))
(step t250 (cl @p_164) :rule and :premises (t164))
(step t251 (cl @p_387) :rule resolution :premises (t230 t249 t250))
(step t252 (cl @p_371 @p_360 @p_378 @p_384 @p_361) :rule or :premises (t251))
(step t253 (cl @p_384 @p_378 @p_371 @p_361 @p_360) :rule reordering :premises (t252))
(step t254 (cl (not (! (= @p_294 (! (not @p_384) :named @p_385)) :named @p_386)) (not @p_294) @p_385) :rule equiv_pos2)
(step t255 (cl (= @p_293 @p_384)) :rule cong :premises (t95 t78))
(step t256 (cl @p_386) :rule cong :premises (t255))
(step t257 (cl @p_307) :rule and :premises (a7))
(step t258 (cl @p_306) :rule and :premises (t257))
(step t259 (cl @p_305) :rule and :premises (t258))
(step t260 (cl @p_304) :rule and :premises (t259))
(step t261 (cl @p_303) :rule and :premises (t260))
(step t262 (cl @p_302) :rule and :premises (t261))
(step t263 (cl @p_301) :rule and :premises (t262))
(step t264 (cl @p_300) :rule and :premises (t263))
(step t265 (cl @p_299) :rule and :premises (t264))
(step t266 (cl @p_298) :rule and :premises (t265))
(step t267 (cl @p_297) :rule and :premises (t266))
(step t268 (cl @p_296) :rule and :premises (t267))
(step t269 (cl @p_295) :rule and :premises (t268))
(step t270 (cl @p_294) :rule and :premises (t269))
(step t271 (cl @p_385) :rule resolution :premises (t254 t256 t270))
(step t272 (cl (not (! (= @p_242 (! (not @p_378) :named @p_379)) :named @p_380)) (not @p_242) @p_379) :rule equiv_pos2)
(step t273 (cl (= @p_241 @p_381)) :rule cong :premises (t78 t36))
(step t274 (cl (= @p_242 (! (not @p_381) :named @p_382))) :rule cong :premises (t273))
(step t275 (cl (= @p_382 @p_379)) :rule cong :premises (t244))
(step t276 (cl @p_380) :rule trans :premises (t274 t275))
(step t277 (cl @p_254) :rule and :premises (a7))
(step t278 (cl @p_253) :rule and :premises (t277))
(step t279 (cl @p_252) :rule and :premises (t278))
(step t280 (cl @p_251) :rule and :premises (t279))
(step t281 (cl @p_250) :rule and :premises (t280))
(step t282 (cl @p_249) :rule and :premises (t281))
(step t283 (cl @p_248) :rule and :premises (t282))
(step t284 (cl @p_247) :rule and :premises (t283))
(step t285 (cl @p_246) :rule and :premises (t284))
(step t286 (cl @p_245) :rule and :premises (t285))
(step t287 (cl @p_244) :rule and :premises (t286))
(step t288 (cl @p_243) :rule and :premises (t287))
(step t289 (cl @p_242) :rule and :premises (t288))
(step t290 (cl @p_379) :rule resolution :premises (t272 t276 t289))
(step t291 (cl (not (! (= @p_316 (! (not @p_371) :named @p_372)) :named @p_373)) (not @p_316) @p_372) :rule equiv_pos2)
(step t292 (cl (= @p_315 (! (= e3 e0) :named @p_374))) :rule all_simplify)
(step t293 (cl (= @p_316 (! (not @p_374) :named @p_375))) :rule cong :premises (t292))
(step t294 (cl (= @p_374 @p_376)) :rule cong :premises (t52 t41))
(step t295 (cl (= @p_375 (! (not @p_376) :named @p_377))) :rule cong :premises (t294))
(step t296 (cl (= @p_377 @p_372)) :rule cong :premises (t243))
(step t297 (cl (= @p_375 @p_372)) :rule trans :premises (t295 t296))
(step t298 (cl @p_373) :rule trans :premises (t293 t297))
(step t299 (cl @p_316) :rule and :premises (t193))
(step t300 (cl @p_372) :rule resolution :premises (t291 t298 t299))
(step t301 (cl (not (! (= @p_310 (! (not @p_361) :named @p_362)) :named @p_363)) (not @p_310) @p_362) :rule equiv_pos2)
(step t302 (cl (= @p_309 (! (= e3 e2) :named @p_364))) :rule all_simplify)
(step t303 (cl (= @p_310 (! (not @p_364) :named @p_365))) :rule cong :premises (t302))
(step t304 (cl (= @p_364 @p_366)) :rule cong :premises (t52 t35))
(step t305 (cl (= @p_365 (! (not @p_366) :named @p_367))) :rule cong :premises (t304))
(step t306 (cl (= @p_367 @p_362)) :rule cong :premises (t246))
(step t307 (cl (= @p_365 @p_362)) :rule trans :premises (t305 t306))
(step t308 (cl @p_363) :rule trans :premises (t303 t307))
(step t309 (cl @p_310) :rule and :premises (t180))
(step t310 (cl @p_362) :rule resolution :premises (t301 t308 t309))
(step t311 (cl @p_360) :rule resolution :premises (t253 t271 t290 t300 t310))
(step t312 (cl (= @p_339 @p_359)) :rule trans :premises (t229 t311))
(step t313 (cl (= @p_359 @p_339)) :rule symm :premises (t312))
(step t314 (cl (= @p_552 (! (= @p_218 @p_339) :named @p_555))) :rule cong :premises (t215 t313))
(step t315 (cl (= (! (= @p_555 false) :named @p_556) (! (not @p_555) :named @p_557))) :rule equiv_simplify)
(step t316 (cl @p_556 (! (not @p_557) :named @p_561)) :rule equiv2 :premises (t315))
(step t317 (cl (not @p_561) @p_555) :rule not_not)
(step t318 (cl @p_556 @p_555) :rule resolution :premises (t316 t317))
(step t319 (cl (not (! (= @p_279 @p_557) :named @p_558)) (not @p_279) @p_557) :rule equiv_pos2)
(step t320 (cl (= @p_278 (! (= @p_339 @p_218) :named @p_559))) :rule cong :premises (t93 t84))
(step t321 (cl (= @p_279 (! (not @p_559) :named @p_560))) :rule cong :premises (t320))
(step t322 (cl (= @p_559 @p_555)) :rule all_simplify)
(step t323 (cl (= @p_560 @p_557)) :rule cong :premises (t322))
(step t324 (cl @p_558) :rule trans :premises (t321 t323))
(step t325 (cl @p_288) :rule and :premises (t259))
(step t326 (cl @p_287) :rule and :premises (t325))
(step t327 (cl @p_286) :rule and :premises (t326))
(step t328 (cl @p_285) :rule and :premises (t327))
(step t329 (cl @p_284) :rule and :premises (t328))
(step t330 (cl @p_283) :rule and :premises (t329))
(step t331 (cl @p_282) :rule and :premises (t330))
(step t332 (cl @p_281) :rule and :premises (t331))
(step t333 (cl @p_280) :rule and :premises (t332))
(step t334 (cl @p_279) :rule and :premises (t333))
(step t335 (cl @p_557) :rule resolution :premises (t319 t324 t334))
(step t336 (cl @p_556) :rule resolution :premises (t318 t335))
(step t337 (cl @p_554) :rule trans :premises (t314 t336))
(step t338 (cl @p_553) :rule resolution :premises (t209 t337))
(step t339 (cl @p_551) :rule resolution :premises (t170 t185 t196 t207 t338))
(step t340 (cl (= @p_383 e1)) :rule symm :premises (t339))
(step t341 (cl (= @p_594 @p_590)) :rule cong :premises (t149 t340))
(step t342 (cl @p_598) :rule trans :premises (t148 t341))
(step t343 (cl (not (! (= @p_189 (! (= @p_588 @p_590) :named @p_592)) :named @p_593)) (not @p_189) @p_592) :rule equiv_pos2)
(step t344 (cl (= @p_189 (! (= @p_590 @p_588) :named @p_591))) :rule cong :premises (t97 t87))
(step t345 (cl (= @p_591 @p_592)) :rule all_simplify)
(step t346 (cl @p_593) :rule trans :premises (t344 t345))
(step t347 (cl @p_197) :rule and :premises (t225))
(step t348 (cl @p_196) :rule and :premises (t347))
(step t349 (cl @p_189) :rule symm :premises (t348))
(step t350 (cl @p_592) :rule resolution :premises (t343 t346 t349))
(step t351 (cl @p_591) :rule symm :premises (t350))
(step t352 (cl @p_589) :rule trans :premises (t148 t341 t351))
(step t353 (cl (not (! (= @p_225 (! (= @p_114 (! (op @p_338 @p_339) :named @p_435)) :named @p_437)) :named @p_438)) (not @p_225) @p_437) :rule equiv_pos2)
(step t354 (cl (= @p_225 (! (= e2 @p_224) :named @p_439))) :rule all_simplify)
(step t355 (cl (= @p_224 @p_435)) :rule cong :premises (t41 t93))
(step t356 (cl (= @p_439 @p_437)) :rule cong :premises (t35 t355))
(step t357 (cl @p_438) :rule trans :premises (t354 t356))
(step t358 (cl @p_225) :rule and :premises (t145))
(step t359 (cl @p_437) :rule resolution :premises (t353 t357 t358))
(step t360 (cl (= @p_435 @p_428)) :rule cong :premises (t149 t229))
(step t361 (cl @p_434) :rule trans :premises (t359 t360))
(step t362 (cl (not (! (= @p_173 (! (= @p_355 @p_404) :named @p_415)) :named @p_416)) (not @p_173) @p_415) :rule equiv_pos2)
(step t363 (cl (= @p_173 (! (= @p_404 @p_355) :named @p_414))) :rule cong :premises (t89 t59))
(step t364 (cl (= @p_414 @p_415)) :rule all_simplify)
(step t365 (cl @p_416) :rule trans :premises (t363 t364))
(step t366 (cl @p_193) :rule and :premises (t223))
(step t367 (cl @p_173) :rule symm :premises (t366))
(step t368 (cl @p_415) :rule resolution :premises (t362 t365 t367))
(step t369 (cl @p_414) :rule symm :premises (t368))
(step t370 (cl (not (! (= @p_206 @p_356) :named @p_357)) (not @p_206) @p_356) :rule equiv_pos2)
(step t371 (cl (= @p_206 (! (= e4 @p_205) :named @p_358))) :rule all_simplify)
(step t372 (cl (= @p_205 @p_355)) :rule cong :premises (t39 t40))
(step t373 (cl (= @p_358 @p_356)) :rule cong :premises (t39 t372))
(step t374 (cl @p_357) :rule trans :premises (t371 t373))
(step t375 (cl @p_207) :rule and :premises (a6))
(step t376 (cl @p_206) :rule and :premises (t375))
(step t377 (cl @p_356) :rule resolution :premises (t370 t374 t376))
(step t378 (cl (= @p_355 @p_337)) :rule symm :premises (t377))
(step t379 (cl (= @p_404 @p_337)) :rule trans :premises (t369 t378))
(step t380 (cl @p_413) :rule symm :premises (t379))
(step t381 (cl (not (! (= @p_217 (! (= @p_337 (! (op e1 @p_520) :named @p_521)) :named @p_525)) :named @p_526)) (not @p_217) @p_525) :rule equiv_pos2)
(step t382 (cl (= @p_217 (! (= e4 @p_216) :named @p_527))) :rule all_simplify)
(step t383 (cl (= @p_216 @p_521)) :rule cong :premises (t43 t80))
(step t384 (cl (= @p_527 @p_525)) :rule cong :premises (t39 t383))
(step t385 (cl @p_526) :rule trans :premises (t382 t384))
(step t386 (cl @p_217) :rule and :premises (t210))
(step t387 (cl @p_525) :rule resolution :premises (t381 t385 t386))
(step t388 (cl (not (! (= @p_172 (! (= @p_459 @p_520) :named @p_523)) :named @p_524)) (not @p_172) @p_523) :rule equiv_pos2)
(step t389 (cl (= @p_172 (! (= @p_520 @p_459) :named @p_522))) :rule cong :premises (t80 t57))
(step t390 (cl (= @p_522 @p_523)) :rule all_simplify)
(step t391 (cl @p_524) :rule trans :premises (t389 t390))
(step t392 (cl @p_192) :rule and :premises (t222))
(step t393 (cl @p_185) :rule and :premises (t392))
(step t394 (cl @p_172) :rule symm :premises (t393))
(step t395 (cl @p_523) :rule resolution :premises (t388 t391 t394))
(step t396 (cl @p_522) :rule symm :premises (t395))
(step t397 (cl (not (! (= @p_159 (! (or @p_371 (! (= @p_340 @p_398) :named @p_417) (! (= @p_340 @p_454) :named @p_461) (! (= @p_340 @p_355) :named @p_474) @p_460) :named @p_477)) :named @p_478)) (not @p_159) @p_477) :rule equiv_pos2)
(step t398 (cl (= @p_159 (! (or @p_3 @p_9 @p_16 @p_28 @p_22) :named @p_484))) :rule all_simplify)
(step t399 (cl (= @p_9 (! (= e3 @p_7) :named @p_450))) :rule all_simplify)
(step t400 (cl (= @p_16 (! (= e3 @p_14) :named @p_481))) :rule all_simplify)
(step t401 (cl (= @p_28 (! (= e3 @p_26) :named @p_480))) :rule all_simplify)
(step t402 (cl (= @p_484 (! (or @p_393 @p_450 @p_481 @p_480 @p_479) :named @p_482))) :rule cong :premises (t232 t399 t400 t401 t5))
(step t403 (cl (= @p_159 @p_482)) :rule trans :premises (t398 t402))
(step t404 (cl (= @p_450 @p_417)) :rule cong :premises (t52 t53))
(step t405 (cl (= @p_481 @p_461)) :rule cong :premises (t52 t55))
(step t406 (cl (= @p_480 @p_474)) :rule cong :premises (t52 t59))
(step t407 (cl (= @p_482 (! (or @p_376 @p_417 @p_461 @p_474 @p_460) :named @p_483))) :rule cong :premises (t238 t404 t405 t406 t58))
(step t408 (cl (= @p_417 @p_417)) :rule refl)
(step t409 (cl (= @p_461 @p_461)) :rule refl)
(step t410 (cl (= @p_474 @p_474)) :rule refl)
(step t411 (cl (= @p_483 @p_477)) :rule cong :premises (t243 t408 t409 t410 t104))
(step t412 (cl (= @p_482 @p_477)) :rule trans :premises (t407 t411))
(step t413 (cl @p_478) :rule trans :premises (t403 t412))
(step t414 (cl @p_162) :rule and :premises (a1))
(step t415 (cl @p_161) :rule and :premises (t414))
(step t416 (cl @p_160) :rule and :premises (t415))
(step t417 (cl @p_159) :rule and :premises (t416))
(step t418 (cl @p_477) :rule resolution :premises (t397 t413 t417))
(step t419 (cl @p_371 @p_417 @p_461 @p_474 @p_460) :rule or :premises (t418))
(step t420 (cl @p_474 @p_371 @p_460 @p_461 @p_417) :rule reordering :premises (t419))
(step t421 (cl (not (! (= @p_290 (! (not @p_474) :named @p_475)) :named @p_476)) (not @p_290) @p_475) :rule equiv_pos2)
(step t422 (cl (= @p_289 @p_474)) :rule cong :premises (t78 t59))
(step t423 (cl @p_476) :rule cong :premises (t422))
(step t424 (cl @p_290) :rule and :premises (t262))
(step t425 (cl @p_475) :rule resolution :premises (t421 t423 t424))
(step t426 (cl (= (! (= @p_461 false) :named @p_463) (! (not @p_461) :named @p_462))) :rule equiv_simplify)
(step t427 (cl (not @p_463) @p_462) :rule equiv1 :premises (t426))
(step t428 (cl @p_473) :rule refl)
(step t429 (cl (not (! (= @p_177 (! (= @p_464 @p_428) :named @p_470)) :named @p_471)) (not @p_177) @p_470) :rule equiv_pos2)
(step t430 (cl (= @p_177 (! (= @p_428 @p_464) :named @p_472))) :rule cong :premises (t91 t69))
(step t431 (cl (= @p_472 @p_470)) :rule all_simplify)
(step t432 (cl @p_471) :rule trans :premises (t430 t431))
(step t433 (cl @p_194) :rule and :premises (t224))
(step t434 (cl @p_177) :rule symm :premises (t433))
(step t435 (cl @p_470) :rule resolution :premises (t429 t432 t434))
(step t436 (cl (not (! (= @p_171 (! (= @p_454 @p_351) :named @p_457)) :named @p_458)) (not @p_171) @p_457) :rule equiv_pos2)
(step t437 (cl (= @p_171 (! (= @p_351 @p_454) :named @p_456))) :rule cong :premises (t71 t55))
(step t438 (cl (= @p_456 @p_457)) :rule all_simplify)
(step t439 (cl @p_458) :rule trans :premises (t437 t438))
(step t440 (cl @p_184) :rule and :premises (t221))
(step t441 (cl @p_179) :rule and :premises (t440))
(step t442 (cl @p_171) :rule symm :premises (t441))
(step t443 (cl @p_457) :rule resolution :premises (t436 t439 t442))
(step t444 (cl (not (! (= @p_209 @p_352) :named @p_353)) (not @p_209) @p_352) :rule equiv_pos2)
(step t445 (cl (= @p_209 (! (= e2 @p_208) :named @p_354))) :rule all_simplify)
(step t446 (cl (= @p_208 @p_351)) :rule cong :premises (t35 t36))
(step t447 (cl (= @p_354 @p_352)) :rule cong :premises (t35 t446))
(step t448 (cl @p_353) :rule trans :premises (t445 t447))
(step t449 (cl @p_215) :rule and :premises (t141))
(step t450 (cl @p_214) :rule and :premises (t449))
(step t451 (cl @p_213) :rule and :premises (t450))
(step t452 (cl @p_209) :rule and :premises (t451))
(step t453 (cl @p_352) :rule resolution :premises (t444 t448 t452))
(step t454 (cl (= @p_351 @p_114)) :rule symm :premises (t453))
(step t455 (cl (= @p_454 @p_428)) :rule trans :premises (t443 t454 t359 t360))
(step t456 (cl (= @p_428 @p_454)) :rule symm :premises (t455))
(step t457 (cl (= @p_464 @p_454)) :rule trans :premises (t435 t456))
(step t458 (cl (= @p_454 @p_464)) :rule symm :premises (t457))
(step t459 (cl (= @p_461 (! (= @p_340 @p_464) :named @p_465))) :rule cong :premises (t428 t458))
(step t460 (cl (= (! (= @p_465 false) :named @p_466) (! (not @p_465) :named @p_467))) :rule equiv_simplify)
(step t461 (cl @p_466 (! (not @p_467) :named @p_469)) :rule equiv2 :premises (t460))
(step t462 (cl (not @p_469) @p_465) :rule not_not)
(step t463 (cl @p_466 @p_465) :rule resolution :premises (t461 t462))
(step t464 (cl (not (! (= @p_292 @p_467) :named @p_468)) (not @p_292) @p_467) :rule equiv_pos2)
(step t465 (cl (= @p_291 @p_465)) :rule cong :premises (t78 t69))
(step t466 (cl @p_468) :rule cong :premises (t465))
(step t467 (cl @p_292) :rule and :premises (t263))
(step t468 (cl @p_467) :rule resolution :premises (t464 t466 t467))
(step t469 (cl @p_466) :rule resolution :premises (t463 t468))
(step t470 (cl @p_463) :rule trans :premises (t459 t469))
(step t471 (cl @p_462) :rule resolution :premises (t427 t470))
(step t472 (cl (= (! (= @p_417 false) :named @p_419) (! (not @p_417) :named @p_418))) :rule equiv_simplify)
(step t473 (cl (not @p_419) @p_418) :rule equiv1 :premises (t472))
(step t474 (cl (= @p_398 @p_398)) :rule refl)
(step t475 (cl (= @p_417 (! (= @p_359 @p_398) :named @p_420))) :rule cong :premises (t311 t474))
(step t476 (cl (= (! (= @p_420 false) :named @p_421) (! (not @p_420) :named @p_422))) :rule equiv_simplify)
(step t477 (cl @p_421 (! (not @p_422) :named @p_424)) :rule equiv2 :premises (t476))
(step t478 (cl (not @p_424) @p_420) :rule not_not)
(step t479 (cl @p_421 @p_420) :rule resolution :premises (t477 t478))
(step t480 (cl (not (! (= @p_268 @p_422) :named @p_423)) (not @p_268) @p_422) :rule equiv_pos2)
(step t481 (cl (= @p_267 @p_420)) :rule cong :premises (t63 t53))
(step t482 (cl @p_423) :rule cong :premises (t481))
(step t483 (cl @p_277) :rule and :premises (t258))
(step t484 (cl @p_268) :rule and :premises (t483))
(step t485 (cl @p_422) :rule resolution :premises (t480 t482 t484))
(step t486 (cl @p_421) :rule resolution :premises (t479 t485))
(step t487 (cl @p_419) :rule trans :premises (t475 t486))
(step t488 (cl @p_418) :rule resolution :premises (t473 t487))
(step t489 (cl @p_460) :rule resolution :premises (t420 t425 t300 t471 t488))
(step t490 (cl (= @p_459 @p_340)) :rule symm :premises (t489))
(step t491 (cl (= @p_520 @p_340)) :rule trans :premises (t396 t490))
(step t492 (cl (= @p_521 @p_516)) :rule cong :premises (t43 t491))
(step t493 (cl @p_550) :rule trans :premises (t387 t492))
(step t494 (cl (not (! (= @p_158 (! (or @p_371 (! (= @p_340 @p_403) :named @p_489) (! (= @p_340 @p_351) :named @p_539) (! (= @p_340 @p_404) :named @p_534) @p_533) :named @p_542)) :named @p_543)) (not @p_158) @p_542) :rule equiv_pos2)
(step t495 (cl (= @p_158 (! (or @p_3 @p_36 @p_68 @p_128 @p_98) :named @p_549))) :rule all_simplify)
(step t496 (cl (= @p_36 (! (= e3 @p_34) :named @p_505))) :rule all_simplify)
(step t497 (cl (= @p_68 (! (= e3 @p_66) :named @p_546))) :rule all_simplify)
(step t498 (cl (= @p_128 (! (= e3 @p_126) :named @p_545))) :rule all_simplify)
(step t499 (cl (= @p_549 (! (or @p_393 @p_505 @p_546 @p_545 @p_544) :named @p_547))) :rule cong :premises (t232 t496 t497 t498 t17))
(step t500 (cl (= @p_158 @p_547)) :rule trans :premises (t495 t499))
(step t501 (cl (= @p_505 @p_489)) :rule cong :premises (t52 t61))
(step t502 (cl (= @p_546 @p_539)) :rule cong :premises (t52 t71))
(step t503 (cl (= @p_545 @p_534)) :rule cong :premises (t52 t89))
(step t504 (cl (= @p_547 (! (or @p_376 @p_489 @p_539 @p_534 @p_533) :named @p_548))) :rule cong :premises (t238 t501 t502 t503 t81))
(step t505 (cl (= @p_489 @p_489)) :rule refl)
(step t506 (cl (= @p_539 @p_539)) :rule refl)
(step t507 (cl (= @p_534 @p_534)) :rule refl)
(step t508 (cl (= @p_548 @p_542)) :rule cong :premises (t243 t505 t506 t507 t116))
(step t509 (cl (= @p_547 @p_542)) :rule trans :premises (t504 t508))
(step t510 (cl @p_543) :rule trans :premises (t500 t509))
(step t511 (cl @p_158) :rule and :premises (t416))
(step t512 (cl @p_542) :rule resolution :premises (t494 t510 t511))
(step t513 (cl @p_371 @p_489 @p_539 @p_534 @p_533) :rule or :premises (t512))
(step t514 (cl @p_539 @p_371 @p_534 @p_533 @p_489) :rule reordering :premises (t513))
(step t515 (cl (not (! (= @p_240 (! (not @p_539) :named @p_540)) :named @p_541)) (not @p_240) @p_540) :rule equiv_pos2)
(step t516 (cl (= @p_239 @p_539)) :rule cong :premises (t78 t71))
(step t517 (cl @p_541) :rule cong :premises (t516))
(step t518 (cl @p_240) :rule and :premises (t286))
(step t519 (cl @p_540) :rule resolution :premises (t515 t517 t518))
(step t520 (cl (= (! (= @p_534 false) :named @p_536) (! (not @p_534) :named @p_535))) :rule equiv_simplify)
(step t521 (cl (not @p_536) @p_535) :rule equiv1 :premises (t520))
(step t522 (cl (= @p_534 @p_474)) :rule cong :premises (t428 t369))
(step t523 (cl (= (! (= @p_474 false) :named @p_537) @p_475)) :rule equiv_simplify)
(step t524 (cl @p_537 (! (not @p_475) :named @p_538)) :rule equiv2 :premises (t523))
(step t525 (cl (not @p_538) @p_474) :rule not_not)
(step t526 (cl @p_537 @p_474) :rule resolution :premises (t524 t525))
(step t527 (cl @p_537) :rule resolution :premises (t526 t425))
(step t528 (cl @p_536) :rule trans :premises (t522 t527))
(step t529 (cl @p_535) :rule resolution :premises (t521 t528))
(step t530 (cl (= (! (= @p_489 false) :named @p_491) (! (not @p_489) :named @p_490))) :rule equiv_simplify)
(step t531 (cl (not @p_491) @p_490) :rule equiv1 :premises (t530))
(step t532 (cl (not (! (= @p_170 (! (= @p_398 @p_403) :named @p_410)) :named @p_411)) (not @p_170) @p_410) :rule equiv_pos2)
(step t533 (cl (= @p_170 (! (= @p_403 @p_398) :named @p_412))) :rule cong :premises (t61 t53))
(step t534 (cl (= @p_412 @p_410)) :rule all_simplify)
(step t535 (cl @p_411) :rule trans :premises (t533 t534))
(step t536 (cl @p_178) :rule and :premises (t220))
(step t537 (cl @p_174) :rule and :premises (t536))
(step t538 (cl @p_170) :rule symm :premises (t537))
(step t539 (cl @p_410) :rule resolution :premises (t532 t535 t538))
(step t540 (cl @p_412) :rule symm :premises (t539))
(step t541 (cl (= @p_489 @p_420)) :rule cong :premises (t311 t540))
(step t542 (cl @p_491) :rule trans :premises (t541 t486))
(step t543 (cl @p_490) :rule resolution :premises (t531 t542))
(step t544 (cl @p_533) :rule resolution :premises (t514 t519 t300 t529 t543))
(step t545 (cl (not (! (= @p_181 (! (= @p_218 @p_529) :named @p_531)) :named @p_532)) (not @p_181) @p_531) :rule equiv_pos2)
(step t546 (cl @p_532) :rule cong :premises (t84 t76))
(step t547 (cl @p_191) :rule and :premises (t392))
(step t548 (cl @p_190) :rule and :premises (t547))
(step t549 (cl @p_187) :rule and :premises (t548))
(step t550 (cl @p_181) :rule symm :premises (t549))
(step t551 (cl @p_531) :rule resolution :premises (t545 t546 t550))
(step t552 (cl @p_530) :rule trans :premises (t215 t551))
(step t553 (cl @p_472) :rule symm :premises (t435))
(step t554 (cl @p_528) :rule trans :premises (t359 t360 t553))
(step t555 (cl (not (! (= @p_176 (! (= @p_514 @p_516) :named @p_518)) :named @p_519)) (not @p_176) @p_518) :rule equiv_pos2)
(step t556 (cl (= @p_176 (! (= @p_516 @p_514) :named @p_517))) :rule cong :premises (t82 t67))
(step t557 (cl (= @p_517 @p_518)) :rule all_simplify)
(step t558 (cl @p_519) :rule trans :premises (t556 t557))
(step t559 (cl @p_186) :rule and :premises (t547))
(step t560 (cl @p_176) :rule symm :premises (t559))
(step t561 (cl @p_518) :rule resolution :premises (t555 t558 t560))
(step t562 (cl @p_517) :rule symm :premises (t561))
(step t563 (cl @p_515) :rule trans :premises (t387 t492 t562))
(step t564 (cl (not (! (= @p_211 @p_348) :named @p_349)) (not @p_211) @p_348) :rule equiv_pos2)
(step t565 (cl (= @p_211 (! (= e0 @p_210) :named @p_350))) :rule all_simplify)
(step t566 (cl (= @p_210 @p_347)) :rule cong :premises (t35 t78))
(step t567 (cl (= @p_350 @p_348)) :rule cong :premises (t41 t566))
(step t568 (cl @p_349) :rule trans :premises (t565 t567))
(step t569 (cl @p_212) :rule and :premises (t451))
(step t570 (cl @p_211) :rule and :premises (t569))
(step t571 (cl @p_348) :rule resolution :premises (t564 t568 t570))
(step t572 (cl (not (! (= @p_175 (! (= @p_509 @p_347) :named @p_512)) :named @p_513)) (not @p_175) @p_512) :rule equiv_pos2)
(step t573 (cl (= @p_175 (! (= @p_347 @p_509) :named @p_511))) :rule cong :premises (t73 t65))
(step t574 (cl (= @p_511 @p_512)) :rule all_simplify)
(step t575 (cl @p_513) :rule trans :premises (t573 t574))
(step t576 (cl @p_183) :rule and :premises (t440))
(step t577 (cl @p_180) :rule and :premises (t576))
(step t578 (cl @p_175) :rule symm :premises (t577))
(step t579 (cl @p_512) :rule resolution :premises (t572 t575 t578))
(step t580 (cl @p_511) :rule symm :premises (t579))
(step t581 (cl @p_510) :rule trans :premises (t571 t580))
(step t582 (cl (not (! (= @p_40 (! (or (! (= @p_337 @p_403) :named @p_486) @p_489 (! (= @p_114 @p_403) :named @p_492) (! (= @p_338 @p_403) :named @p_495) @p_485) :named @p_500)) :named @p_501)) (not @p_40) @p_500) :rule equiv_pos2)
(step t583 (cl (= @p_40 (! (or @p_35 @p_36 @p_37 @p_39 @p_38) :named @p_508))) :rule all_simplify)
(step t584 (cl (= @p_35 (! (= e4 @p_34) :named @p_506))) :rule all_simplify)
(step t585 (cl (= @p_37 (! (= e2 @p_34) :named @p_504))) :rule all_simplify)
(step t586 (cl (= @p_39 (! (= e0 @p_34) :named @p_503))) :rule all_simplify)
(step t587 (cl (= @p_508 (! (or @p_506 @p_505 @p_504 @p_503 @p_502) :named @p_507))) :rule cong :premises (t584 t496 t585 t586 t7))
(step t588 (cl (= @p_40 @p_507)) :rule trans :premises (t583 t587))
(step t589 (cl (= @p_506 @p_486)) :rule cong :premises (t39 t61))
(step t590 (cl (= @p_504 @p_492)) :rule cong :premises (t35 t61))
(step t591 (cl (= @p_503 @p_495)) :rule cong :premises (t41 t61))
(step t592 (cl (= @p_507 @p_500)) :rule cong :premises (t589 t501 t590 t591 t62))
(step t593 (cl @p_501) :rule trans :premises (t588 t592))
(step t594 (cl @p_156) :rule and :premises (a0))
(step t595 (cl @p_65) :rule and :premises (t594))
(step t596 (cl @p_40) :rule and :premises (t595))
(step t597 (cl @p_500) :rule resolution :premises (t582 t593 t596))
(step t598 (cl @p_486 @p_489 @p_492 @p_495 @p_485) :rule or :premises (t597))
(step t599 (cl @p_495 @p_485 @p_492 @p_489 @p_486) :rule reordering :premises (t598))
(step t600 (cl (not (! (= @p_256 (! (not @p_495) :named @p_496)) :named @p_497)) (not @p_256) @p_496) :rule equiv_pos2)
(step t601 (cl (= @p_255 (! (= @p_403 @p_338) :named @p_498))) :rule cong :premises (t61 t40))
(step t602 (cl (= @p_256 (! (not @p_498) :named @p_499))) :rule cong :premises (t601))
(step t603 (cl (= @p_498 @p_495)) :rule all_simplify)
(step t604 (cl (= @p_499 @p_496)) :rule cong :premises (t603))
(step t605 (cl @p_497) :rule trans :premises (t602 t604))
(step t606 (cl @p_266) :rule and :premises (t257))
(step t607 (cl @p_256) :rule and :premises (t606))
(step t608 (cl @p_496) :rule resolution :premises (t600 t605 t607))
(step t609 (cl (= (! (= @p_492 false) :named @p_494) (! (not @p_492) :named @p_493))) :rule equiv_simplify)
(step t610 (cl (not @p_494) @p_493) :rule equiv1 :premises (t609))
(step t611 (cl (= @p_492 (! (= @p_428 @p_398) :named @p_429))) :rule cong :premises (t361 t540))
(step t612 (cl (= (! (= @p_429 false) :named @p_430) (! (not @p_429) :named @p_431))) :rule equiv_simplify)
(step t613 (cl @p_430 (! (not @p_431) :named @p_433)) :rule equiv2 :premises (t612))
(step t614 (cl (not @p_433) @p_429) :rule not_not)
(step t615 (cl @p_430 @p_429) :rule resolution :premises (t613 t614))
(step t616 (cl (not (! (= @p_270 @p_431) :named @p_432)) (not @p_270) @p_431) :rule equiv_pos2)
(step t617 (cl (= @p_269 @p_429)) :rule cong :premises (t91 t53))
(step t618 (cl @p_432) :rule cong :premises (t617))
(step t619 (cl @p_276) :rule and :premises (t483))
(step t620 (cl @p_275) :rule and :premises (t619))
(step t621 (cl @p_274) :rule and :premises (t620))
(step t622 (cl @p_273) :rule and :premises (t621))
(step t623 (cl @p_272) :rule and :premises (t622))
(step t624 (cl @p_271) :rule and :premises (t623))
(step t625 (cl @p_270) :rule and :premises (t624))
(step t626 (cl @p_431) :rule resolution :premises (t616 t618 t625))
(step t627 (cl @p_430) :rule resolution :premises (t615 t626))
(step t628 (cl @p_494) :rule trans :premises (t611 t627))
(step t629 (cl @p_493) :rule resolution :premises (t610 t628))
(step t630 (cl (= (! (= @p_486 false) :named @p_488) (! (not @p_486) :named @p_487))) :rule equiv_simplify)
(step t631 (cl (not @p_488) @p_487) :rule equiv1 :premises (t630))
(step t632 (cl (= @p_403 @p_403)) :rule refl)
(step t633 (cl (= @p_486 (! (= @p_404 @p_403) :named @p_405))) :rule cong :premises (t380 t632))
(step t634 (cl (= (! (= @p_405 false) :named @p_406) (! (not @p_405) :named @p_407))) :rule equiv_simplify)
(step t635 (cl @p_406 (! (not @p_407) :named @p_409)) :rule equiv2 :premises (t634))
(step t636 (cl (not @p_409) @p_405) :rule not_not)
(step t637 (cl @p_406 @p_405) :rule resolution :premises (t635 t636))
(step t638 (cl (not (! (= @p_258 @p_407) :named @p_408)) (not @p_258) @p_407) :rule equiv_pos2)
(step t639 (cl (= @p_257 @p_405)) :rule cong :premises (t89 t61))
(step t640 (cl @p_408) :rule cong :premises (t639))
(step t641 (cl @p_265) :rule and :premises (t606))
(step t642 (cl @p_264) :rule and :premises (t641))
(step t643 (cl @p_263) :rule and :premises (t642))
(step t644 (cl @p_262) :rule and :premises (t643))
(step t645 (cl @p_261) :rule and :premises (t644))
(step t646 (cl @p_260) :rule and :premises (t645))
(step t647 (cl @p_259) :rule and :premises (t646))
(step t648 (cl @p_258) :rule and :premises (t647))
(step t649 (cl @p_407) :rule resolution :premises (t638 t640 t648))
(step t650 (cl @p_406) :rule resolution :premises (t637 t649))
(step t651 (cl @p_488) :rule trans :premises (t633 t650))
(step t652 (cl @p_487) :rule resolution :premises (t631 t651))
(step t653 (cl @p_485) :rule resolution :premises (t599 t608 t629 t543 t652))
(step t654 (cl @p_456) :rule symm :premises (t443))
(step t655 (cl @p_455) :rule trans :premises (t453 t654))
(step t656 (cl (not (! (= @p_13 (! (or (! (= @p_337 @p_398) :named @p_400) @p_417 (! (= @p_114 @p_398) :named @p_425) (! (= @p_338 @p_398) :named @p_440) @p_399) :named @p_445)) :named @p_446)) (not @p_13) @p_445) :rule equiv_pos2)
(step t657 (cl (= @p_13 (! (or @p_8 @p_9 @p_10 @p_12 @p_11) :named @p_453))) :rule all_simplify)
(step t658 (cl (= @p_8 (! (= e4 @p_7) :named @p_451))) :rule all_simplify)
(step t659 (cl (= @p_10 (! (= e2 @p_7) :named @p_449))) :rule all_simplify)
(step t660 (cl (= @p_12 (! (= e0 @p_7) :named @p_448))) :rule all_simplify)
(step t661 (cl (= @p_453 (! (or @p_451 @p_450 @p_449 @p_448 @p_447) :named @p_452))) :rule cong :premises (t658 t399 t659 t660 t3))
(step t662 (cl (= @p_13 @p_452)) :rule trans :premises (t657 t661))
(step t663 (cl (= @p_451 @p_400)) :rule cong :premises (t39 t53))
(step t664 (cl (= @p_449 @p_425)) :rule cong :premises (t35 t53))
(step t665 (cl (= @p_448 @p_440)) :rule cong :premises (t41 t53))
(step t666 (cl (= @p_452 @p_445)) :rule cong :premises (t663 t404 t664 t665 t54))
(step t667 (cl @p_446) :rule trans :premises (t662 t666))
(step t668 (cl @p_33) :rule and :premises (a0))
(step t669 (cl @p_32) :rule and :premises (t668))
(step t670 (cl @p_13) :rule and :premises (t669))
(step t671 (cl @p_445) :rule resolution :premises (t656 t667 t670))
(step t672 (cl @p_400 @p_417 @p_425 @p_440 @p_399) :rule or :premises (t671))
(step t673 (cl @p_440 @p_399 @p_425 @p_417 @p_400) :rule reordering :premises (t672))
(step t674 (cl (not (! (= @p_237 (! (not @p_440) :named @p_441)) :named @p_442)) (not @p_237) @p_441) :rule equiv_pos2)
(step t675 (cl (= @p_236 (! (= @p_398 @p_338) :named @p_443))) :rule cong :premises (t53 t40))
(step t676 (cl (= @p_237 (! (not @p_443) :named @p_444))) :rule cong :premises (t675))
(step t677 (cl (= @p_443 @p_440)) :rule all_simplify)
(step t678 (cl (= @p_444 @p_441)) :rule cong :premises (t677))
(step t679 (cl @p_442) :rule trans :premises (t676 t678))
(step t680 (cl @p_238) :rule and :premises (t277))
(step t681 (cl @p_237) :rule and :premises (t680))
(step t682 (cl @p_441) :rule resolution :premises (t674 t679 t681))
(step t683 (cl (= (! (= @p_425 false) :named @p_427) (! (not @p_425) :named @p_426))) :rule equiv_simplify)
(step t684 (cl (not @p_427) @p_426) :rule equiv1 :premises (t683))
(step t685 (cl (= @p_425 @p_429)) :rule cong :premises (t361 t474))
(step t686 (cl @p_427) :rule trans :premises (t685 t627))
(step t687 (cl @p_426) :rule resolution :premises (t684 t686))
(step t688 (cl (= (! (= @p_400 false) :named @p_402) (! (not @p_400) :named @p_401))) :rule equiv_simplify)
(step t689 (cl (not @p_402) @p_401) :rule equiv1 :premises (t688))
(step t690 (cl (= @p_400 @p_405)) :rule cong :premises (t380 t539))
(step t691 (cl @p_402) :rule trans :premises (t690 t650))
(step t692 (cl @p_401) :rule resolution :premises (t689 t691))
(step t693 (cl @p_399) :rule resolution :premises (t673 t682 t687 t488 t692))
(step t694 (cl) :rule resolution :premises (t134 t342 t352 t361 t380 t339 t493 t544 t552 t554 t563 t581 t653 t489 t655 t693 t311 t377 t453 t571 t215 t228))
