unsat
(assume a0 (! (and (! (and (! (and (! (and (! (and (! (and (! (and (and (or (or (or (or (! (= (! (op e0 e0) :named @p_148) e0) :named @p_153) (! (= @p_148 e1) :named @p_152)) (! (= @p_148 e2) :named @p_151)) (! (= @p_148 e3) :named @p_150)) (! (= @p_148 e4) :named @p_149)) (or (or (or (or (! (= (! (op e0 e1) :named @p_142) e0) :named @p_147) (! (= @p_142 e1) :named @p_146)) (! (= @p_142 e2) :named @p_145)) (! (= @p_142 e3) :named @p_144)) (! (= @p_142 e4) :named @p_143))) (! (or (or (or (or (! (= (! (op e0 e2) :named @p_135) e0) :named @p_140) (! (= @p_135 e1) :named @p_139)) (! (= @p_135 e2) :named @p_138)) (! (= @p_135 e3) :named @p_137)) (! (= @p_135 e4) :named @p_136)) :named @p_141)) :named @p_154) (or (or (or (or (! (= (! (op e0 e3) :named @p_129) e0) :named @p_134) (! (= @p_129 e1) :named @p_133)) (! (= @p_129 e2) :named @p_132)) (! (= @p_129 e3) :named @p_131)) (! (= @p_129 e4) :named @p_130))) :named @p_155) (or (or (or (or (! (= (! (op e0 e4) :named @p_123) e0) :named @p_128) (! (= @p_123 e1) :named @p_127)) (! (= @p_123 e2) :named @p_126)) (! (= @p_123 e3) :named @p_125)) (! (= @p_123 e4) :named @p_124))) :named @p_156) (! (and (and (and (and (or (or (or (or (! (= (! (op e1 e0) :named @p_116) e0) :named @p_121) (! (= @p_116 e1) :named @p_120)) (! (= @p_116 e2) :named @p_119)) (! (= @p_116 e3) :named @p_118)) (! (= @p_116 e4) :named @p_117)) (or (or (or (or (! (= (! (op e1 e1) :named @p_110) e0) :named @p_115) (! (= @p_110 e1) :named @p_114)) (! (= @p_110 e2) :named @p_113)) (! (= @p_110 e3) :named @p_112)) (! (= @p_110 e4) :named @p_111))) (or (or (or (or (! (= (! (op e1 e2) :named @p_104) e0) :named @p_109) (! (= @p_104 e1) :named @p_108)) (! (= @p_104 e2) :named @p_107)) (! (= @p_104 e3) :named @p_106)) (! (= @p_104 e4) :named @p_105))) (or (or (or (or (! (= (! (op e1 e3) :named @p_98) e0) :named @p_103) (! (= @p_98 e1) :named @p_102)) (! (= @p_98 e2) :named @p_101)) (! (= @p_98 e3) :named @p_100)) (! (= @p_98 e4) :named @p_99))) (! (or (or (or (or (! (= (! (op e1 e4) :named @p_91) e0) :named @p_96) (! (= @p_91 e1) :named @p_95)) (! (= @p_91 e2) :named @p_94)) (! (= @p_91 e3) :named @p_93)) (! (= @p_91 e4) :named @p_92)) :named @p_97)) :named @p_122)) :named @p_157) (and (and (and (and (or (or (or (or (! (= (! (op e2 e0) :named @p_85) e0) :named @p_90) (! (= @p_85 e1) :named @p_89)) (! (= @p_85 e2) :named @p_88)) (! (= @p_85 e3) :named @p_87)) (! (= @p_85 e4) :named @p_86)) (or (or (or (or (! (= (! (op e2 e1) :named @p_79) e0) :named @p_84) (! (= @p_79 e1) :named @p_83)) (! (= @p_79 e2) :named @p_82)) (! (= @p_79 e3) :named @p_81)) (! (= @p_79 e4) :named @p_80))) (or (or (or (or (! (= (! (op e2 e2) :named @p_73) e0) :named @p_78) (! (= @p_73 e1) :named @p_77)) (! (= @p_73 e2) :named @p_76)) (! (= @p_73 e3) :named @p_75)) (! (= @p_73 e4) :named @p_74))) (or (or (or (or (! (= (! (op e2 e3) :named @p_67) e0) :named @p_72) (! (= @p_67 e1) :named @p_71)) (! (= @p_67 e2) :named @p_70)) (! (= @p_67 e3) :named @p_69)) (! (= @p_67 e4) :named @p_68))) (or (or (or (or (! (= (! (op e2 e4) :named @p_61) e0) :named @p_66) (! (= @p_61 e1) :named @p_65)) (! (= @p_61 e2) :named @p_64)) (! (= @p_61 e3) :named @p_63)) (! (= @p_61 e4) :named @p_62)))) :named @p_158) (and (and (and (and (or (or (or (or (! (= (! (op e3 e0) :named @p_55) e0) :named @p_60) (! (= @p_55 e1) :named @p_59)) (! (= @p_55 e2) :named @p_58)) (! (= @p_55 e3) :named @p_57)) (! (= @p_55 e4) :named @p_56)) (or (or (or (or (! (= (! (op e3 e1) :named @p_49) e0) :named @p_54) (! (= @p_49 e1) :named @p_53)) (! (= @p_49 e2) :named @p_52)) (! (= @p_49 e3) :named @p_51)) (! (= @p_49 e4) :named @p_50))) (or (or (or (or (! (= (! (op e3 e2) :named @p_43) e0) :named @p_48) (! (= @p_43 e1) :named @p_47)) (! (= @p_43 e2) :named @p_46)) (! (= @p_43 e3) :named @p_45)) (! (= @p_43 e4) :named @p_44))) (or (or (or (or (! (= (! (op e3 e3) :named @p_37) e0) :named @p_42) (! (= @p_37 e1) :named @p_41)) (! (= @p_37 e2) :named @p_40)) (! (= @p_37 e3) :named @p_39)) (! (= @p_37 e4) :named @p_38))) (or (or (or (or (! (= (! (op e3 e4) :named @p_31) e0) :named @p_36) (! (= @p_31 e1) :named @p_35)) (! (= @p_31 e2) :named @p_34)) (! (= @p_31 e3) :named @p_33)) (! (= @p_31 e4) :named @p_32)))) :named @p_159) (and (and (and (and (or (or (or (or (! (= (! (op e4 e0) :named @p_25) e0) :named @p_30) (! (= @p_25 e1) :named @p_29)) (! (= @p_25 e2) :named @p_28)) (! (= @p_25 e3) :named @p_27)) (! (= @p_25 e4) :named @p_26)) (or (or (or (or (! (= (! (op e4 e1) :named @p_19) e0) :named @p_24) (! (= @p_19 e1) :named @p_23)) (! (= @p_19 e2) :named @p_22)) (! (= @p_19 e3) :named @p_21)) (! (= @p_19 e4) :named @p_20))) (or (or (or (or (! (= (! (op e4 e2) :named @p_13) e0) :named @p_18) (! (= @p_13 e1) :named @p_17)) (! (= @p_13 e2) :named @p_16)) (! (= @p_13 e3) :named @p_15)) (! (= @p_13 e4) :named @p_14))) (or (or (or (or (! (= (! (op e4 e3) :named @p_7) e0) :named @p_12) (! (= @p_7 e1) :named @p_11)) (! (= @p_7 e2) :named @p_10)) (! (= @p_7 e3) :named @p_9)) (! (= @p_7 e4) :named @p_8))) (or (or (or (or (! (= (! (op e4 e4) :named @p_1) e0) :named @p_6) (! (= @p_1 e1) :named @p_5)) (! (= @p_1 e2) :named @p_4)) (! (= @p_1 e3) :named @p_3)) (! (= @p_1 e4) :named @p_2)))) :named @p_160))
(assume a1 (! (and (! (and (! (and (! (and (! (and (and (and (and (and (or (or (or (or @p_153 @p_147) @p_140) @p_134) @p_128) (or (or (or (or @p_153 @p_121) @p_90) @p_60) @p_30)) (and (or (or (or (or @p_152 @p_146) @p_139) @p_133) @p_127) (or (or (or (or @p_152 @p_120) @p_89) @p_59) @p_29))) (and (or (or (or (or @p_151 @p_145) @p_138) @p_132) @p_126) (or (or (or (or @p_151 @p_119) @p_88) @p_58) @p_28))) (and (or (or (or (or @p_150 @p_144) @p_137) @p_131) @p_125) (or (or (or (or @p_150 @p_118) @p_87) @p_57) @p_27))) (! (and (! (or (or (or (or @p_149 @p_143) @p_136) @p_130) @p_124) :named @p_177) (or (or (or (or @p_149 @p_117) @p_86) @p_56) @p_26)) :named @p_178)) :named @p_179) (! (and (! (and (! (and (! (and (and (or (or (or (or @p_121 @p_115) @p_109) @p_103) @p_96) (or (or (or (or @p_147 @p_115) @p_84) @p_54) @p_24)) (! (and (! (or (or (or (or @p_120 @p_114) @p_108) @p_102) @p_95) :named @p_171) (or (or (or (or @p_146 @p_114) @p_83) @p_53) @p_23)) :named @p_172)) :named @p_173) (and (or (or (or (or @p_119 @p_113) @p_107) @p_101) @p_94) (or (or (or (or @p_145 @p_113) @p_82) @p_52) @p_22))) :named @p_174) (and (or (or (or (or @p_118 @p_112) @p_106) @p_100) @p_93) (or (or (or (or @p_144 @p_112) @p_81) @p_51) @p_21))) :named @p_175) (and (or (or (or (or @p_117 @p_111) @p_105) @p_99) @p_92) (or (or (or (or @p_143 @p_111) @p_80) @p_50) @p_20))) :named @p_176)) :named @p_180) (! (and (! (and (and (and (and (or (or (or (or @p_90 @p_84) @p_78) @p_72) @p_66) (or (or (or (or @p_140 @p_109) @p_78) @p_48) @p_18)) (and (or (or (or (or @p_89 @p_83) @p_77) @p_71) @p_65) (or (or (or (or @p_139 @p_108) @p_77) @p_47) @p_17))) (and (or (or (or (or @p_88 @p_82) @p_76) @p_70) @p_64) (or (or (or (or @p_138 @p_107) @p_76) @p_46) @p_16))) (! (and (or (or (or (or @p_87 @p_81) @p_75) @p_69) @p_63) (! (or (or (or (or @p_137 @p_106) @p_75) @p_45) @p_15) :named @p_167)) :named @p_168)) :named @p_169) (and (or (or (or (or @p_86 @p_80) @p_74) @p_68) @p_62) (or (or (or (or @p_136 @p_105) @p_74) @p_44) @p_14))) :named @p_170)) :named @p_181) (and (and (and (and (and (or (or (or (or @p_60 @p_54) @p_48) @p_42) @p_36) (or (or (or (or @p_134 @p_103) @p_72) @p_42) @p_12)) (and (or (or (or (or @p_59 @p_53) @p_47) @p_41) @p_35) (or (or (or (or @p_133 @p_102) @p_71) @p_41) @p_11))) (and (or (or (or (or @p_58 @p_52) @p_46) @p_40) @p_34) (or (or (or (or @p_132 @p_101) @p_70) @p_40) @p_10))) (and (or (or (or (or @p_57 @p_51) @p_45) @p_39) @p_33) (or (or (or (or @p_131 @p_100) @p_69) @p_39) @p_9))) (and (or (or (or (or @p_56 @p_50) @p_44) @p_38) @p_32) (or (or (or (or @p_130 @p_99) @p_68) @p_38) @p_8)))) :named @p_182) (! (and (! (and (! (and (! (and (and (or (or (or (or @p_30 @p_24) @p_18) @p_12) @p_6) (or (or (or (or @p_128 @p_96) @p_66) @p_36) @p_6)) (! (and (! (or (or (or (or @p_29 @p_23) @p_17) @p_11) @p_5) :named @p_161) (or (or (or (or @p_127 @p_95) @p_65) @p_35) @p_5)) :named @p_162)) :named @p_163) (and (or (or (or (or @p_28 @p_22) @p_16) @p_10) @p_4) (or (or (or (or @p_126 @p_94) @p_64) @p_34) @p_4))) :named @p_164) (and (or (or (or (or @p_27 @p_21) @p_15) @p_9) @p_3) (or (or (or (or @p_125 @p_93) @p_63) @p_33) @p_3))) :named @p_165) (and (or (or (or (or @p_26 @p_20) @p_14) @p_8) @p_2) (or (or (or (or @p_124 @p_92) @p_62) @p_32) @p_2))) :named @p_166)) :named @p_183))
(assume a2 (or (or (or (or (and (and (and (and (! (not @p_153) :named @p_188) (not @p_115)) (not @p_78)) (not @p_42)) (not @p_6)) (and (and (and (and (not @p_152) (! (not @p_114) :named @p_187)) (not @p_77)) (not @p_41)) (not @p_5))) (and (and (and (and (not @p_151) (not @p_113)) (! (not @p_76) :named @p_186)) (not @p_40)) (not @p_4))) (and (and (and (and (not @p_150) (not @p_112)) (not @p_75)) (! (not @p_39) :named @p_185)) (not @p_3))) (and (and (and (and (not @p_149) (not @p_111)) (not @p_74)) (not @p_38)) (! (not @p_2) :named @p_184))))
(assume a3 (and (and (and (and (or (or (or (or (or (or (or (or (and @p_153 @p_153) (and @p_147 @p_121)) (and @p_140 @p_90)) (and @p_134 @p_60)) (and @p_128 @p_30)) (or (or (or (or (and @p_121 @p_147) (and @p_115 @p_115)) (and @p_109 @p_84)) (and @p_103 @p_54)) (and @p_96 @p_24))) (or (or (or (or (and @p_90 @p_140) (and @p_84 @p_109)) (and @p_78 @p_78)) (and @p_72 @p_48)) (and @p_66 @p_18))) (or (or (or (or (and @p_60 @p_134) (and @p_54 @p_103)) (and @p_48 @p_72)) (and @p_42 @p_42)) (and @p_36 @p_12))) (or (or (or (or (and @p_30 @p_128) (and @p_24 @p_96)) (and @p_18 @p_66)) (and @p_12 @p_36)) (and @p_6 @p_6))) (or (or (or (or (or (or (or (or (and @p_152 @p_152) (and @p_146 @p_120)) (and @p_139 @p_89)) (and @p_133 @p_59)) (and @p_127 @p_29)) (or (or (or (or (and @p_120 @p_146) (and @p_114 @p_114)) (and @p_108 @p_83)) (and @p_102 @p_53)) (and @p_95 @p_23))) (or (or (or (or (and @p_89 @p_139) (and @p_83 @p_108)) (and @p_77 @p_77)) (and @p_71 @p_47)) (and @p_65 @p_17))) (or (or (or (or (and @p_59 @p_133) (and @p_53 @p_102)) (and @p_47 @p_71)) (and @p_41 @p_41)) (and @p_35 @p_11))) (or (or (or (or (and @p_29 @p_127) (and @p_23 @p_95)) (and @p_17 @p_65)) (and @p_11 @p_35)) (and @p_5 @p_5)))) (or (or (or (or (or (or (or (or (and @p_151 @p_151) (and @p_145 @p_119)) (and @p_138 @p_88)) (and @p_132 @p_58)) (and @p_126 @p_28)) (or (or (or (or (and @p_119 @p_145) (and @p_113 @p_113)) (and @p_107 @p_82)) (and @p_101 @p_52)) (and @p_94 @p_22))) (or (or (or (or (and @p_88 @p_138) (and @p_82 @p_107)) (and @p_76 @p_76)) (and @p_70 @p_46)) (and @p_64 @p_16))) (or (or (or (or (and @p_58 @p_132) (and @p_52 @p_101)) (and @p_46 @p_70)) (and @p_40 @p_40)) (and @p_34 @p_10))) (or (or (or (or (and @p_28 @p_126) (and @p_22 @p_94)) (and @p_16 @p_64)) (and @p_10 @p_34)) (and @p_4 @p_4)))) (or (or (or (or (or (or (or (or (and @p_150 @p_150) (and @p_144 @p_118)) (and @p_137 @p_87)) (and @p_131 @p_57)) (and @p_125 @p_27)) (or (or (or (or (and @p_118 @p_144) (and @p_112 @p_112)) (and @p_106 @p_81)) (and @p_100 @p_51)) (and @p_93 @p_21))) (or (or (or (or (and @p_87 @p_137) (and @p_81 @p_106)) (and @p_75 @p_75)) (and @p_69 @p_45)) (and @p_63 @p_15))) (or (or (or (or (and @p_57 @p_131) (and @p_51 @p_100)) (and @p_45 @p_69)) (and @p_39 @p_39)) (and @p_33 @p_9))) (or (or (or (or (and @p_27 @p_125) (and @p_21 @p_93)) (and @p_15 @p_63)) (and @p_9 @p_33)) (and @p_3 @p_3)))) (or (or (or (or (or (or (or (or (and @p_149 @p_149) (and @p_143 @p_117)) (and @p_136 @p_86)) (and @p_130 @p_56)) (and @p_124 @p_26)) (or (or (or (or (and @p_117 @p_143) (and @p_111 @p_111)) (and @p_105 @p_80)) (and @p_99 @p_50)) (and @p_92 @p_20))) (or (or (or (or (and @p_86 @p_136) (and @p_80 @p_105)) (and @p_74 @p_74)) (and @p_68 @p_44)) (and @p_62 @p_14))) (or (or (or (or (and @p_56 @p_130) (and @p_50 @p_99)) (and @p_44 @p_68)) (and @p_38 @p_38)) (and @p_32 @p_8))) (or (or (or (or (and @p_26 @p_124) (and @p_20 @p_92)) (and @p_14 @p_62)) (and @p_8 @p_32)) (and @p_2 @p_2)))))
(assume a4 (or (or (or (or (and (and (and (and (= @p_148 @p_148) (= @p_142 @p_116)) (= @p_135 @p_85)) (= @p_129 @p_55)) (= @p_123 @p_25)) (and (and (and (and (= @p_116 @p_142) (= @p_110 @p_110)) (= @p_104 @p_79)) (= @p_98 @p_49)) (= @p_91 @p_19))) (and (and (and (and (= @p_85 @p_135) (= @p_79 @p_104)) (= @p_73 @p_73)) (= @p_67 @p_43)) (= @p_61 @p_13))) (and (and (and (and (= @p_55 @p_129) (= @p_49 @p_98)) (= @p_43 @p_67)) (! (= @p_37 @p_37) :named @p_189)) (= @p_31 @p_7))) (and (and (and (and (= @p_25 @p_123) (= @p_19 @p_91)) (= @p_13 @p_61)) (= @p_7 @p_31)) (= @p_1 @p_1))))
(assume a5 (or (or (or (or @p_153 @p_114) @p_76) @p_39) @p_2))
(assume a6 (and (and (and (and (or (or (or (or (! (= (op @p_148 e0) e0) :named @p_227) (! (= (op @p_142 e0) e1) :named @p_226)) (! (= (! (op @p_135 e0) :named @p_224) e2) :named @p_225)) (! (= (op @p_129 e0) e3) :named @p_223)) (! (= (! (op @p_123 e0) :named @p_221) e4) :named @p_222)) (or (or (or (or (! (= (op @p_116 e1) e0) :named @p_220) (! (= (op @p_110 e1) e1) :named @p_219)) (! (= (! (op @p_104 e1) :named @p_217) e2) :named @p_218)) (! (= (! (op @p_98 e1) :named @p_215) e3) :named @p_216)) (! (= (! (op @p_91 e1) :named @p_213) e4) :named @p_214))) (or (or (or (or (! (= (! (op @p_85 e2) :named @p_211) e0) :named @p_212) (! (= (! (op @p_79 e2) :named @p_209) e1) :named @p_210)) (! (= (op @p_73 e2) e2) :named @p_208)) (! (= (op @p_67 e2) e3) :named @p_207)) (! (= (op @p_61 e2) e4) :named @p_206))) (or (or (or (or (! (= (op @p_55 e3) e0) :named @p_205) (! (= (op @p_49 e3) e1) :named @p_204)) (! (= (op @p_43 e3) e2) :named @p_203)) (! (= (! (op @p_37 e3) :named @p_201) e3) :named @p_202)) (! (= (! (op @p_31 e3) :named @p_199) e4) :named @p_200))) (or (or (or (or (! (= (! (op @p_25 e4) :named @p_197) e0) :named @p_198) (! (= (! (op @p_19 e4) :named @p_195) e1) :named @p_196)) (! (= (! (op @p_13 e4) :named @p_193) e2) :named @p_194)) (! (= (! (op @p_7 e4) :named @p_191) e3) :named @p_192)) (! (= (op @p_1 e4) e4) :named @p_190))))
(assume a7 (and (and (and (and (or (or (or (or (! (= (! (op e0 @p_148) :named @p_233) e0) :named @p_234) (= (op e0 @p_142) e1)) (= (op e0 @p_135) e2)) (= (op e0 @p_129) e3)) (= (op e0 @p_123) e4)) (or (or (or (or (= (op e1 @p_116) e0) (! (= (op e1 @p_110) e1) :named @p_232)) (= (op e1 @p_104) e2)) (= (op e1 @p_98) e3)) (= (op e1 @p_91) e4))) (or (or (or (or (= (op e2 @p_85) e0) (= (op e2 @p_79) e1)) (! (= (op e2 @p_73) e2) :named @p_231)) (= (op e2 @p_67) e3)) (= (op e2 @p_61) e4))) (or (or (or (or (= (op e3 @p_55) e0) (= (op e3 @p_49) e1)) (= (op e3 @p_43) e2)) (! (= (! (op e3 @p_37) :named @p_229) e3) :named @p_230)) (= (op e3 @p_31) e4))) (or (or (or (or (= (op e4 @p_25) e0) (= (op e4 @p_19) e1)) (= (op e4 @p_13) e2)) (= (op e4 @p_7) e3)) (! (= (op e4 @p_1) e4) :named @p_228))))
(assume a8 (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (or @p_188 @p_153) (! (or (! (not @p_147) :named @p_244) @p_121) :named @p_245)) :named @p_246) (or (not @p_140) @p_90)) :named @p_247) (or (not @p_134) @p_60)) :named @p_248) (or (not @p_128) @p_30)) :named @p_249) (! (and (and (and (and (or (not @p_120) @p_146) (or @p_187 @p_114)) (or (not @p_108) @p_83)) (or (not @p_102) @p_53)) (! (or (! (not @p_95) :named @p_241) @p_23) :named @p_242)) :named @p_243)) :named @p_250) (and (and (and (and (or (not @p_88) @p_138) (or (not @p_82) @p_107)) (or @p_186 @p_76)) (or (not @p_70) @p_46)) (or (not @p_64) @p_16))) :named @p_251) (! (and (! (and (! (and (! (and (! (or (! (not @p_57) :named @p_235) @p_131) :named @p_236) (or (not @p_51) @p_100)) :named @p_237) (or (not @p_45) @p_69)) :named @p_238) (or @p_185 @p_39)) :named @p_239) (or (not @p_33) @p_9)) :named @p_240)) :named @p_252) (and (and (and (and (or (not @p_26) @p_124) (or (not @p_20) @p_92)) (or (not @p_14) @p_62)) (or (not @p_8) @p_32)) (or @p_184 @p_2))) :named @p_253))
(assume a9 (! (and (! (and (! (and (! (and @p_234 @p_232) :named @p_254) @p_231) :named @p_255) @p_230) :named @p_256) @p_228) :named @p_257))
(assume a10 (! (and (! (and (! (and (! (and (! (and (! (and (! (and (and @p_227 @p_226) @p_225) :named @p_270) @p_223) :named @p_271) @p_222) :named @p_272) (! (and (! (and (! (and (and @p_220 @p_219) @p_218) :named @p_267) @p_216) :named @p_268) @p_214) :named @p_269)) :named @p_273) (! (and (! (and (! (and (! (and @p_212 @p_210) :named @p_263) @p_208) :named @p_264) @p_207) :named @p_265) @p_206) :named @p_266)) :named @p_274) (! (and (and (and (and @p_205 @p_204) @p_203) @p_202) @p_200) :named @p_262)) :named @p_275) (! (and (! (and (! (and (! (and @p_198 @p_196) :named @p_258) @p_194) :named @p_259) @p_192) :named @p_260) @p_190) :named @p_261)) :named @p_276))
(assume a11 (or (or (or (or (not (= (op @p_148 @p_148) @p_148)) (not (= (op @p_110 @p_110) @p_110))) (not (= (op @p_73 @p_73) @p_73))) (not (= (! (op @p_37 @p_37) :named @p_277) @p_37))) (not (= (op @p_1 @p_1) @p_1))))
(assume a12 (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (not (! (= @p_148 @p_116) :named @p_343)) :named @p_344) (not (= @p_148 @p_85))) :named @p_345) (! (not (! (= @p_148 @p_55) :named @p_341)) :named @p_342)) :named @p_346) (! (not (! (= @p_148 @p_25) :named @p_339)) :named @p_340)) :named @p_347) (not (= @p_116 @p_85))) :named @p_348) (not (= @p_116 @p_55))) :named @p_349) (not (= @p_116 @p_25))) :named @p_350) (not (= @p_85 @p_55))) :named @p_351) (not (= @p_85 @p_25))) :named @p_352) (not (= @p_55 @p_25))) :named @p_353) (and (and (and (and (and (and (and (and (and (not (= @p_142 @p_110)) (not (= @p_142 @p_79))) (not (= @p_142 @p_49))) (not (= @p_142 @p_19))) (not (= @p_110 @p_79))) (not (= @p_110 @p_49))) (not (= @p_110 @p_19))) (not (= @p_79 @p_49))) (not (= @p_79 @p_19))) (not (= @p_49 @p_19)))) :named @p_354) (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (not (! (= @p_135 @p_104) :named @p_328)) :named @p_329) (not (= @p_135 @p_73))) :named @p_330) (not (= @p_135 @p_43))) :named @p_331) (not (= @p_135 @p_13))) :named @p_332) (not (= @p_104 @p_73))) :named @p_333) (not (= @p_104 @p_43))) :named @p_334) (not (= @p_104 @p_13))) :named @p_335) (not (= @p_73 @p_43))) :named @p_336) (not (= @p_73 @p_13))) :named @p_337) (! (not (! (= @p_43 @p_13) :named @p_326)) :named @p_327)) :named @p_338)) :named @p_355) (and (and (and (and (and (and (and (and (and (not (= @p_129 @p_98)) (not (= @p_129 @p_67))) (not (= @p_129 @p_37))) (not (= @p_129 @p_7))) (not (= @p_98 @p_67))) (not (= @p_98 @p_37))) (not (= @p_98 @p_7))) (not (= @p_67 @p_37))) (not (= @p_67 @p_7))) (not (= @p_37 @p_7)))) :named @p_356) (! (and (! (and (and (and (and (and (and (and (and (not (= @p_123 @p_91)) (not (= @p_123 @p_61))) (not (= @p_123 @p_31))) (not (= @p_123 @p_1))) (not (= @p_91 @p_61))) (not (= @p_91 @p_31))) (not (= @p_91 @p_1))) (not (= @p_61 @p_31))) (! (not (! (= @p_61 @p_1) :named @p_322)) :named @p_323)) :named @p_324) (not (= @p_31 @p_1))) :named @p_325)) :named @p_357) (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (not (= @p_148 @p_142)) (! (not (! (= @p_148 @p_135) :named @p_307)) :named @p_308)) :named @p_309) (not (= @p_148 @p_129))) :named @p_310) (not (= @p_148 @p_123))) :named @p_311) (! (not (! (= @p_142 @p_135) :named @p_305)) :named @p_306)) :named @p_312) (not (= @p_142 @p_129))) :named @p_313) (not (= @p_142 @p_123))) :named @p_314) (! (not (! (= @p_135 @p_129) :named @p_303)) :named @p_304)) :named @p_315) (not (= @p_135 @p_123))) :named @p_316) (not (= @p_129 @p_123))) :named @p_317) (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (! (and (not (= @p_116 @p_110)) (! (not (! (= @p_116 @p_104) :named @p_292)) :named @p_293)) :named @p_294) (not (= @p_116 @p_98))) :named @p_295) (! (not (! (= @p_116 @p_91) :named @p_290)) :named @p_291)) :named @p_296) (not (= @p_110 @p_104))) :named @p_297) (not (= @p_110 @p_98))) :named @p_298) (not (= @p_110 @p_91))) :named @p_299) (not (= @p_104 @p_98))) :named @p_300) (! (not (! (= @p_104 @p_91) :named @p_288)) :named @p_289)) :named @p_301) (! (not (! (= @p_98 @p_91) :named @p_286)) :named @p_287)) :named @p_302)) :named @p_318) (! (and (! (and (! (and (! (and (! (and (! (and (and (and (and (not (= @p_85 @p_79)) (not (= @p_85 @p_73))) (not (= @p_85 @p_67))) (not (= @p_85 @p_61))) (! (not (! (= @p_79 @p_73) :named @p_278)) :named @p_279)) :named @p_280) (not (= @p_79 @p_67))) :named @p_281) (not (= @p_79 @p_61))) :named @p_282) (not (= @p_73 @p_67))) :named @p_283) (not (= @p_73 @p_61))) :named @p_284) (not (= @p_67 @p_61))) :named @p_285)) :named @p_319) (and (and (and (and (and (and (and (and (and (not (= @p_55 @p_49)) (not (= @p_55 @p_43))) (not (= @p_55 @p_37))) (not (= @p_55 @p_31))) (not (= @p_49 @p_43))) (not (= @p_49 @p_37))) (not (= @p_49 @p_31))) (not (= @p_43 @p_37))) (not (= @p_43 @p_31))) (not (= @p_37 @p_31)))) :named @p_320) (and (and (and (and (and (and (and (and (and (not (= @p_25 @p_19)) (not (= @p_25 @p_13))) (not (= @p_25 @p_7))) (not (= @p_25 @p_1))) (not (= @p_19 @p_13))) (not (= @p_19 @p_7))) (not (= @p_19 @p_1))) (not (= @p_13 @p_7))) (not (= @p_13 @p_1))) (not (= @p_7 @p_1)))) :named @p_321)) :named @p_358))
(assume a13 (! (and (! (and (! (and (! (and (! (and (! (and (and (and (and (not (= e0 e1)) (not (= e0 e2))) (not (= e0 e3))) (not (= e0 e4))) (! (not (! (= e1 e2) :named @p_363)) :named @p_364)) :named @p_365) (not (= e1 e3))) :named @p_366) (! (not (! (= e1 e4) :named @p_361)) :named @p_362)) :named @p_367) (not (= e2 e3))) :named @p_368) (not (= e2 e4))) :named @p_369) (! (not (! (= e3 e4) :named @p_359)) :named @p_360)) :named @p_370))
(assume a14 (! (and (! (and (! (and (! (= e0 @p_37) :named @p_374) (! (= e1 @p_148) :named @p_373)) :named @p_375) (! (= e2 @p_98) :named @p_372)) :named @p_376) (! (= e4 @p_104) :named @p_371)) :named @p_377))
(assume a15 (! (not (! (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and (and @p_152 @p_147) @p_138) @p_131) @p_124) @p_121) @p_114) @p_105) @p_101) @p_93) @p_88) @p_81) @p_78) @p_68) @p_65) @p_57) @p_50) @p_47) @p_42) @p_34) @p_26) @p_22) @p_15) @p_11) @p_6) :named @p_378)) :named @p_379))
(step t0 (cl (not (! (= @p_379 (! (not (! (and (! (= @p_37 (! (op (! (op @p_277 (! (op @p_277 e3) :named @p_404)) :named @p_405) @p_405) :named @p_416)) :named @p_417) (! (= @p_277 (! (op @p_405 e3) :named @p_478)) :named @p_479) (! (= e3 (! (op @p_405 @p_404) :named @p_494)) :named @p_535) (! (= @p_404 (! (op @p_405 @p_277) :named @p_406)) :named @p_407) (! (= (! (op @p_277 (! (op @p_277 e3) :named @p_404)) :named @p_405) (! (op @p_405 @p_37) :named @p_418)) :named @p_574) (! (= @p_404 (! (op e3 @p_405) :named @p_575)) :named @p_576) (! (= @p_277 (! (op e3 @p_404) :named @p_498)) :named @p_504) (! (= (! (op @p_277 (! (op @p_277 e3) :named @p_404)) :named @p_405) (! (op e3 @p_277) :named @p_581)) :named @p_582) (! (= e3 @p_229) :named @p_395) (! (= @p_277 (! (op @p_404 @p_405) :named @p_483)) :named @p_489) (! (= (! (op @p_277 (! (op @p_277 e3) :named @p_404)) :named @p_405) (! (op @p_404 e3) :named @p_592)) :named @p_593) (! (= @p_37 (! (op @p_404 @p_404) :named @p_546)) :named @p_666) (! (= e3 (! (op @p_404 @p_277) :named @p_412)) :named @p_413) (! (= @p_404 (! (op @p_404 @p_37) :named @p_660)) :named @p_661) (! (= e3 (! (op @p_277 @p_405) :named @p_583)) :named @p_618) (! (= @p_277 (! (op @p_277 @p_277) :named @p_598)) :named @p_599) (! (= @p_37 (! (op @p_277 @p_37) :named @p_384)) :named @p_385) (! (= (! (op @p_277 (! (op @p_277 e3) :named @p_404)) :named @p_405) (! (op @p_37 @p_405) :named @p_421)) :named @p_428) (! (= e3 @p_201) :named @p_396) (! (= @p_404 (! (op @p_37 @p_404) :named @p_458)) :named @p_641) (! (= @p_37 (! (op @p_37 @p_277) :named @p_380)) :named @p_381)) :named @p_690)) :named @p_691)) :named @p_692)) (not @p_379) @p_691) :rule equiv_pos2)
(step t1 (cl (= @p_378 (! (and @p_6 @p_11 @p_15 @p_22 @p_26 @p_34 @p_42 @p_47 @p_50 @p_57 @p_65 @p_68 @p_78 @p_81 @p_88 @p_93 @p_101 @p_105 @p_114 @p_121 @p_124 @p_131 @p_138 @p_152 @p_147) :named @p_709))) :rule all_simplify)
(step t2 (cl (= @p_6 (! (= e0 @p_1) :named @p_703))) :rule all_simplify)
(step t3 (cl (= @p_11 (! (= e1 @p_7) :named @p_531))) :rule all_simplify)
(step t4 (cl (= @p_15 (! (= e3 @p_13) :named @p_571))) :rule all_simplify)
(step t5 (cl (= @p_22 (! (= e2 @p_19) :named @p_702))) :rule all_simplify)
(step t6 (cl (= @p_26 (! (= e4 @p_25) :named @p_701))) :rule all_simplify)
(step t7 (cl (= @p_34 (! (= e2 @p_31) :named @p_700))) :rule all_simplify)
(step t8 (cl (= @p_42 @p_374)) :rule all_simplify)
(step t9 (cl (= @p_47 (! (= e1 @p_43) :named @p_699))) :rule all_simplify)
(step t10 (cl (= @p_50 (! (= e4 @p_49) :named @p_698))) :rule all_simplify)
(step t11 (cl (= @p_57 (! (= e3 @p_55) :named @p_401))) :rule all_simplify)
(step t12 (cl (= @p_65 (! (= e1 @p_61) :named @p_697))) :rule all_simplify)
(step t13 (cl (= @p_68 (! (= e4 @p_67) :named @p_696))) :rule all_simplify)
(step t14 (cl (= @p_78 (! (= e0 @p_73) :named @p_695))) :rule all_simplify)
(step t15 (cl (= @p_81 (! (= e3 @p_79) :named @p_694))) :rule all_simplify)
(step t16 (cl (= @p_88 (! (= e2 @p_85) :named @p_693))) :rule all_simplify)
(step t17 (cl (= @p_93 (! (= e3 @p_91) :named @p_637))) :rule all_simplify)
(step t18 (cl (= @p_101 @p_372)) :rule all_simplify)
(step t19 (cl (= @p_105 @p_371)) :rule all_simplify)
(step t20 (cl (= @p_114 (! (= e1 @p_110) :named @p_612))) :rule all_simplify)
(step t21 (cl (= @p_121 (! (= e0 @p_116) :named @p_389))) :rule all_simplify)
(step t22 (cl (= @p_124 (! (= e4 @p_123) :named @p_470))) :rule all_simplify)
(step t23 (cl (= @p_131 (! (= e3 @p_129) :named @p_400))) :rule all_simplify)
(step t24 (cl (= @p_138 (! (= e2 @p_135) :named @p_657))) :rule all_simplify)
(step t25 (cl (= @p_152 @p_373)) :rule all_simplify)
(step t26 (cl (= @p_147 (! (= e0 @p_142) :named @p_390))) :rule all_simplify)
(step t27 (cl (= @p_709 (! (and @p_703 @p_531 @p_571 @p_702 @p_701 @p_700 @p_374 @p_699 @p_698 @p_401 @p_697 @p_696 @p_695 @p_694 @p_693 @p_637 @p_372 @p_371 @p_612 @p_389 @p_470 @p_400 @p_657 @p_373 @p_390) :named @p_704))) :rule cong :premises (t2 t3 t4 t5 t6 t7 t8 t9 t10 t11 t12 t13 t14 t15 t16 t17 t18 t19 t20 t21 t22 t23 t24 t25 t26))
(step t28 (cl (= @p_378 @p_704)) :rule trans :premises (t1 t27))
(step t29 (cl (= @p_379 (! (not @p_704) :named @p_705))) :rule cong :premises (t28))
(step t30 (cl @p_376) :rule and :premises (a14))
(step t31 (cl @p_375) :rule and :premises (t30))
(step t32 (cl @p_374) :rule and :premises (t31))
(step t33 (cl (not (! (= @p_371 (! (= e4 @p_405) :named @p_426)) :named @p_427)) (not @p_371) @p_426) :rule equiv_pos2)
(step t34 (cl (= e4 e4)) :rule refl)
(step t35 (cl (not (! (= @p_373 (! (= e1 @p_277) :named @p_393)) :named @p_394)) (not @p_373) @p_393) :rule equiv_pos2)
(step t36 (cl (= e1 e1)) :rule refl)
(step t37 (cl (= @p_148 @p_277)) :rule cong :premises (t32 t32))
(step t38 (cl @p_394) :rule cong :premises (t36 t37))
(step t39 (cl @p_373) :rule and :premises (t31))
(step t40 (cl @p_393) :rule resolution :premises (t35 t38 t39))
(step t41 (cl (not (! (= @p_372 (! (= e2 @p_404) :named @p_410)) :named @p_411)) (not @p_372) @p_410) :rule equiv_pos2)
(step t42 (cl (= e2 e2)) :rule refl)
(step t43 (cl (= e3 e3)) :rule refl)
(step t44 (cl (= @p_98 @p_404)) :rule cong :premises (t40 t43))
(step t45 (cl @p_411) :rule cong :premises (t42 t44))
(step t46 (cl @p_372) :rule and :premises (t30))
(step t47 (cl @p_410) :rule resolution :premises (t41 t45 t46))
(step t48 (cl (= @p_104 @p_405)) :rule cong :premises (t40 t47))
(step t49 (cl @p_427) :rule cong :premises (t34 t48))
(step t50 (cl @p_371) :rule and :premises (a14))
(step t51 (cl @p_426) :rule resolution :premises (t33 t49 t50))
(step t52 (cl (= @p_1 @p_416)) :rule cong :premises (t51 t51))
(step t53 (cl (= @p_703 @p_417)) :rule cong :premises (t32 t52))
(step t54 (cl (= @p_7 @p_478)) :rule cong :premises (t51 t43))
(step t55 (cl (= @p_531 @p_479)) :rule cong :premises (t40 t54))
(step t56 (cl (= @p_13 @p_494)) :rule cong :premises (t51 t47))
(step t57 (cl (= @p_571 @p_535)) :rule cong :premises (t43 t56))
(step t58 (cl (= @p_19 @p_406)) :rule cong :premises (t51 t40))
(step t59 (cl (= @p_702 @p_407)) :rule cong :premises (t47 t58))
(step t60 (cl (= @p_25 @p_418)) :rule cong :premises (t51 t32))
(step t61 (cl (= @p_701 @p_574)) :rule cong :premises (t51 t60))
(step t62 (cl (= @p_31 @p_575)) :rule cong :premises (t43 t51))
(step t63 (cl (= @p_700 @p_576)) :rule cong :premises (t47 t62))
(step t64 (cl @p_189) :rule refl)
(step t65 (cl (= @p_374 @p_189)) :rule cong :premises (t32 t64))
(step t66 (cl (= @p_43 @p_498)) :rule cong :premises (t43 t47))
(step t67 (cl (= @p_699 @p_504)) :rule cong :premises (t40 t66))
(step t68 (cl (= @p_49 @p_581)) :rule cong :premises (t43 t40))
(step t69 (cl (= @p_698 @p_582)) :rule cong :premises (t51 t68))
(step t70 (cl (= @p_55 @p_229)) :rule cong :premises (t43 t32))
(step t71 (cl (= @p_401 @p_395)) :rule cong :premises (t43 t70))
(step t72 (cl (= @p_61 @p_483)) :rule cong :premises (t47 t51))
(step t73 (cl (= @p_697 @p_489)) :rule cong :premises (t40 t72))
(step t74 (cl (= @p_67 @p_592)) :rule cong :premises (t47 t43))
(step t75 (cl (= @p_696 @p_593)) :rule cong :premises (t51 t74))
(step t76 (cl (= @p_73 @p_546)) :rule cong :premises (t47 t47))
(step t77 (cl (= @p_695 @p_666)) :rule cong :premises (t32 t76))
(step t78 (cl (= @p_79 @p_412)) :rule cong :premises (t47 t40))
(step t79 (cl (= @p_694 @p_413)) :rule cong :premises (t43 t78))
(step t80 (cl (= @p_85 @p_660)) :rule cong :premises (t47 t32))
(step t81 (cl (= @p_693 @p_661)) :rule cong :premises (t47 t80))
(step t82 (cl (= @p_91 @p_583)) :rule cong :premises (t40 t51))
(step t83 (cl (= @p_637 @p_618)) :rule cong :premises (t43 t82))
(step t84 (cl (= @p_372 (! (= @p_404 @p_404) :named @p_506))) :rule cong :premises (t47 t44))
(step t85 (cl (= @p_371 (! (= @p_405 @p_405) :named @p_420))) :rule cong :premises (t51 t48))
(step t86 (cl (= @p_110 @p_598)) :rule cong :premises (t40 t40))
(step t87 (cl (= @p_612 @p_599)) :rule cong :premises (t40 t86))
(step t88 (cl (= @p_116 @p_384)) :rule cong :premises (t40 t32))
(step t89 (cl (= @p_389 @p_385)) :rule cong :premises (t32 t88))
(step t90 (cl (= @p_123 @p_421)) :rule cong :premises (t32 t51))
(step t91 (cl (= @p_470 @p_428)) :rule cong :premises (t51 t90))
(step t92 (cl (= @p_129 @p_201)) :rule cong :premises (t32 t43))
(step t93 (cl (= @p_400 @p_396)) :rule cong :premises (t43 t92))
(step t94 (cl (= @p_135 @p_458)) :rule cong :premises (t32 t47))
(step t95 (cl (= @p_657 @p_641)) :rule cong :premises (t47 t94))
(step t96 (cl (= @p_373 (! (= @p_277 @p_277) :named @p_522))) :rule cong :premises (t40 t37))
(step t97 (cl (= @p_142 @p_380)) :rule cong :premises (t32 t40))
(step t98 (cl (= @p_390 @p_381)) :rule cong :premises (t32 t97))
(step t99 (cl (= @p_704 (! (and @p_417 @p_479 @p_535 @p_407 @p_574 @p_576 @p_189 @p_504 @p_582 @p_395 @p_489 @p_593 @p_666 @p_413 @p_661 @p_618 @p_506 @p_420 @p_599 @p_385 @p_428 @p_396 @p_641 @p_522 @p_381) :named @p_706))) :rule cong :premises (t53 t55 t57 t59 t61 t63 t65 t67 t69 t71 t73 t75 t77 t79 t81 t83 t84 t85 t87 t89 t91 t93 t95 t96 t98))
(step t100 (cl (= @p_705 (! (not @p_706) :named @p_707))) :rule cong :premises (t99))
(step t101 (cl (= @p_417 @p_417)) :rule refl)
(step t102 (cl (= @p_479 @p_479)) :rule refl)
(step t103 (cl (= @p_535 @p_535)) :rule refl)
(step t104 (cl (= @p_407 @p_407)) :rule refl)
(step t105 (cl (= @p_574 @p_574)) :rule refl)
(step t106 (cl (= @p_576 @p_576)) :rule refl)
(step t107 (cl (= @p_189 true)) :rule all_simplify)
(step t108 (cl (= @p_504 @p_504)) :rule refl)
(step t109 (cl (= @p_582 @p_582)) :rule refl)
(step t110 (cl (= @p_395 @p_395)) :rule refl)
(step t111 (cl (= @p_489 @p_489)) :rule refl)
(step t112 (cl (= @p_593 @p_593)) :rule refl)
(step t113 (cl (= @p_666 @p_666)) :rule refl)
(step t114 (cl (= @p_413 @p_413)) :rule refl)
(step t115 (cl (= @p_661 @p_661)) :rule refl)
(step t116 (cl (= @p_618 @p_618)) :rule refl)
(step t117 (cl (= @p_506 true)) :rule all_simplify)
(step t118 (cl (= @p_420 true)) :rule all_simplify)
(step t119 (cl (= @p_599 @p_599)) :rule refl)
(step t120 (cl (= @p_385 @p_385)) :rule refl)
(step t121 (cl (= @p_428 @p_428)) :rule refl)
(step t122 (cl (= @p_396 @p_396)) :rule refl)
(step t123 (cl (= @p_641 @p_641)) :rule refl)
(step t124 (cl (= @p_522 true)) :rule all_simplify)
(step t125 (cl (= @p_381 @p_381)) :rule refl)
(step t126 (cl (= @p_706 (! (and @p_417 @p_479 @p_535 @p_407 @p_574 @p_576 true @p_504 @p_582 @p_395 @p_489 @p_593 @p_666 @p_413 @p_661 @p_618 true true @p_599 @p_385 @p_428 @p_396 @p_641 true @p_381) :named @p_708))) :rule cong :premises (t101 t102 t103 t104 t105 t106 t107 t108 t109 t110 t111 t112 t113 t114 t115 t116 t117 t118 t119 t120 t121 t122 t123 t124 t125))
(step t127 (cl (= @p_708 @p_690)) :rule all_simplify)
(step t128 (cl (= @p_706 @p_690)) :rule trans :premises (t126 t127))
(step t129 (cl (= @p_707 @p_691)) :rule cong :premises (t128))
(step t130 (cl (= @p_705 @p_691)) :rule trans :premises (t100 t129))
(step t131 (cl @p_692) :rule trans :premises (t29 t130))
(step t132 (cl @p_691) :rule resolution :premises (t0 t131 a15))
(step t133 (cl (! (not @p_417) :named @p_674) (! (not @p_479) :named @p_689) (! (not @p_535) :named @p_688) (! (not @p_407) :named @p_687) (! (not @p_574) :named @p_676) (! (not @p_576) :named @p_686) (! (not @p_504) :named @p_685) (! (not @p_582) :named @p_684) (! (not @p_395) :named @p_397) (! (not @p_489) :named @p_683) (! (not @p_593) :named @p_682) (! (not @p_666) :named @p_673) (! (not @p_413) :named @p_681) (! (not @p_661) :named @p_675) (! (not @p_618) :named @p_680) (! (not @p_599) :named @p_672) (! (not @p_385) :named @p_679) (! (not @p_428) :named @p_678) (! (not @p_396) :named @p_671) (! (not @p_641) :named @p_677) (! (not @p_381) :named @p_386)) :rule not_and :premises (t132))
(step t134 (cl @p_397 @p_671 @p_672 @p_673 @p_674 @p_386 @p_675 @p_676 @p_677 @p_678 @p_679 @p_680 @p_681 @p_682 @p_683 @p_684 @p_685 @p_686 @p_687 @p_688 @p_689) :rule reordering :premises (t133))
(step t135 (cl (not (! (= @p_212 (! (= @p_37 (! (op @p_660 @p_404) :named @p_667)) :named @p_668)) :named @p_669)) (not @p_212) @p_668) :rule equiv_pos2)
(step t136 (cl (= @p_212 (! (= e0 @p_211) :named @p_670))) :rule all_simplify)
(step t137 (cl (= @p_211 @p_667)) :rule cong :premises (t80 t47))
(step t138 (cl (= @p_670 @p_668)) :rule cong :premises (t32 t137))
(step t139 (cl @p_669) :rule trans :premises (t136 t138))
(step t140 (cl @p_275) :rule and :premises (a10))
(step t141 (cl @p_274) :rule and :premises (t140))
(step t142 (cl @p_266) :rule and :premises (t141))
(step t143 (cl @p_265) :rule and :premises (t142))
(step t144 (cl @p_264) :rule and :premises (t143))
(step t145 (cl @p_263) :rule and :premises (t144))
(step t146 (cl @p_212) :rule and :premises (t145))
(step t147 (cl @p_668) :rule resolution :premises (t135 t139 t146))
(step t148 (cl (not (! (= @p_141 (! (or (! (= @p_405 @p_458) :named @p_459) (! (= e3 @p_458) :named @p_555) @p_641 (! (= @p_37 @p_458) :named @p_642) (! (= @p_277 @p_458) :named @p_650)) :named @p_653)) :named @p_654)) (not @p_141) @p_653) :rule equiv_pos2)
(step t149 (cl (= @p_141 (! (or @p_136 @p_137 @p_138 @p_140 @p_139) :named @p_659))) :rule all_simplify)
(step t150 (cl (= @p_136 (! (= e4 @p_135) :named @p_468))) :rule all_simplify)
(step t151 (cl (= @p_137 (! (= e3 @p_135) :named @p_568))) :rule all_simplify)
(step t152 (cl (= @p_140 (! (= e0 @p_135) :named @p_656))) :rule all_simplify)
(step t153 (cl (= @p_139 (! (= e1 @p_135) :named @p_655))) :rule all_simplify)
(step t154 (cl (= @p_659 (! (or @p_468 @p_568 @p_657 @p_656 @p_655) :named @p_658))) :rule cong :premises (t150 t151 t24 t152 t153))
(step t155 (cl (= @p_141 @p_658)) :rule trans :premises (t149 t154))
(step t156 (cl (= @p_468 @p_459)) :rule cong :premises (t51 t94))
(step t157 (cl (= @p_568 @p_555)) :rule cong :premises (t43 t94))
(step t158 (cl (= @p_656 @p_642)) :rule cong :premises (t32 t94))
(step t159 (cl (= @p_655 @p_650)) :rule cong :premises (t40 t94))
(step t160 (cl (= @p_658 @p_653)) :rule cong :premises (t156 t157 t95 t158 t159))
(step t161 (cl @p_654) :rule trans :premises (t155 t160))
(step t162 (cl @p_159) :rule and :premises (a0))
(step t163 (cl @p_158) :rule and :premises (t162))
(step t164 (cl @p_157) :rule and :premises (t163))
(step t165 (cl @p_156) :rule and :premises (t164))
(step t166 (cl @p_155) :rule and :premises (t165))
(step t167 (cl @p_154) :rule and :premises (t166))
(step t168 (cl @p_141) :rule and :premises (t167))
(step t169 (cl @p_653) :rule resolution :premises (t148 t161 t168))
(step t170 (cl @p_459 @p_555 @p_641 @p_642 @p_650) :rule or :premises (t169))
(step t171 (cl @p_459 @p_650 @p_642 @p_641 @p_555) :rule reordering :premises (t170))
(step t172 (cl (not (! (= @p_329 (! (not @p_459) :named @p_460)) :named @p_461)) (not @p_329) @p_460) :rule equiv_pos2)
(step t173 (cl (= @p_328 (! (= @p_458 @p_405) :named @p_462))) :rule cong :premises (t94 t48))
(step t174 (cl (= @p_329 (! (not @p_462) :named @p_463))) :rule cong :premises (t173))
(step t175 (cl (= @p_462 @p_459)) :rule all_simplify)
(step t176 (cl (= @p_463 @p_460)) :rule cong :premises (t175))
(step t177 (cl @p_461) :rule trans :premises (t174 t176))
(step t178 (cl @p_357) :rule and :premises (a12))
(step t179 (cl @p_356) :rule and :premises (t178))
(step t180 (cl @p_355) :rule and :premises (t179))
(step t181 (cl @p_338) :rule and :premises (t180))
(step t182 (cl @p_337) :rule and :premises (t181))
(step t183 (cl @p_336) :rule and :premises (t182))
(step t184 (cl @p_335) :rule and :premises (t183))
(step t185 (cl @p_334) :rule and :premises (t184))
(step t186 (cl @p_333) :rule and :premises (t185))
(step t187 (cl @p_332) :rule and :premises (t186))
(step t188 (cl @p_331) :rule and :premises (t187))
(step t189 (cl @p_330) :rule and :premises (t188))
(step t190 (cl @p_329) :rule and :premises (t189))
(step t191 (cl @p_460) :rule resolution :premises (t172 t177 t190))
(step t192 (cl (not (! (= @p_308 (! (not @p_650) :named @p_651)) :named @p_652)) (not @p_308) @p_651) :rule equiv_pos2)
(step t193 (cl (= @p_307 @p_650)) :rule cong :premises (t37 t94))
(step t194 (cl @p_652) :rule cong :premises (t193))
(step t195 (cl @p_321) :rule and :premises (a12))
(step t196 (cl @p_320) :rule and :premises (t195))
(step t197 (cl @p_319) :rule and :premises (t196))
(step t198 (cl @p_318) :rule and :premises (t197))
(step t199 (cl @p_317) :rule and :premises (t198))
(step t200 (cl @p_316) :rule and :premises (t199))
(step t201 (cl @p_315) :rule and :premises (t200))
(step t202 (cl @p_314) :rule and :premises (t201))
(step t203 (cl @p_313) :rule and :premises (t202))
(step t204 (cl @p_312) :rule and :premises (t203))
(step t205 (cl @p_311) :rule and :premises (t204))
(step t206 (cl @p_310) :rule and :premises (t205))
(step t207 (cl @p_309) :rule and :premises (t206))
(step t208 (cl @p_308) :rule and :premises (t207))
(step t209 (cl @p_651) :rule resolution :premises (t192 t194 t208))
(step t210 (cl (= (! (= @p_642 false) :named @p_644) (! (not @p_642) :named @p_643))) :rule equiv_simplify)
(step t211 (cl (not @p_644) @p_643) :rule equiv1 :premises (t210))
(step t212 (cl (not (! (= @p_234 @p_381) :named @p_382)) (not @p_234) @p_381) :rule equiv_pos2)
(step t213 (cl (= @p_234 (! (= e0 @p_233) :named @p_383))) :rule all_simplify)
(step t214 (cl (= @p_233 @p_380)) :rule cong :premises (t32 t37))
(step t215 (cl (= @p_383 @p_381)) :rule cong :premises (t32 t214))
(step t216 (cl @p_382) :rule trans :premises (t213 t215))
(step t217 (cl @p_256) :rule and :premises (a9))
(step t218 (cl @p_255) :rule and :premises (t217))
(step t219 (cl @p_254) :rule and :premises (t218))
(step t220 (cl @p_234) :rule and :premises (t219))
(step t221 (cl @p_381) :rule resolution :premises (t212 t216 t220))
(step t222 (cl (= @p_458 @p_458)) :rule refl)
(step t223 (cl (= @p_642 (! (= @p_380 @p_458) :named @p_645))) :rule cong :premises (t221 t222))
(step t224 (cl (= (! (= @p_645 false) :named @p_646) (! (not @p_645) :named @p_647))) :rule equiv_simplify)
(step t225 (cl @p_646 (! (not @p_647) :named @p_649)) :rule equiv2 :premises (t224))
(step t226 (cl (not @p_649) @p_645) :rule not_not)
(step t227 (cl @p_646 @p_645) :rule resolution :premises (t225 t226))
(step t228 (cl (not (! (= @p_306 @p_647) :named @p_648)) (not @p_306) @p_647) :rule equiv_pos2)
(step t229 (cl (= @p_305 @p_645)) :rule cong :premises (t97 t94))
(step t230 (cl @p_648) :rule cong :premises (t229))
(step t231 (cl @p_306) :rule and :premises (t204))
(step t232 (cl @p_647) :rule resolution :premises (t228 t230 t231))
(step t233 (cl @p_646) :rule resolution :premises (t227 t232))
(step t234 (cl @p_644) :rule trans :premises (t223 t233))
(step t235 (cl @p_643) :rule resolution :premises (t211 t234))
(step t236 (cl (= (! (= @p_555 false) :named @p_557) (! (not @p_555) :named @p_556))) :rule equiv_simplify)
(step t237 (cl (not @p_557) @p_556) :rule equiv1 :premises (t236))
(step t238 (cl (not (! (= @p_236 (! (or @p_397 @p_396) :named @p_398)) :named @p_399)) (not @p_236) @p_398) :rule equiv_pos2)
(step t239 (cl (= @p_235 (! (not @p_401) :named @p_402))) :rule cong :premises (t11))
(step t240 (cl (= @p_236 (! (or @p_402 @p_400) :named @p_403))) :rule cong :premises (t239 t23))
(step t241 (cl (= @p_402 @p_397)) :rule cong :premises (t71))
(step t242 (cl (= @p_403 @p_398)) :rule cong :premises (t241 t93))
(step t243 (cl @p_399) :rule trans :premises (t240 t242))
(step t244 (cl @p_252) :rule and :premises (a8))
(step t245 (cl @p_240) :rule and :premises (t244))
(step t246 (cl @p_239) :rule and :premises (t245))
(step t247 (cl @p_238) :rule and :premises (t246))
(step t248 (cl @p_237) :rule and :premises (t247))
(step t249 (cl @p_236) :rule and :premises (t248))
(step t250 (cl @p_398) :rule resolution :premises (t238 t243 t249))
(step t251 (cl @p_397 @p_396) :rule or :premises (t250))
(step t252 (cl @p_396 @p_397) :rule reordering :premises (t251))
(step t253 (cl @p_230) :rule and :premises (t217))
(step t254 (cl @p_395) :rule symm :premises (t253))
(step t255 (cl @p_396) :rule resolution :premises (t252 t254))
(step t256 (cl (= @p_555 (! (= @p_201 @p_458) :named @p_558))) :rule cong :premises (t255 t222))
(step t257 (cl (= (! (= @p_558 false) :named @p_559) (! (not @p_558) :named @p_560))) :rule equiv_simplify)
(step t258 (cl @p_559 (! (not @p_560) :named @p_564)) :rule equiv2 :premises (t257))
(step t259 (cl (not @p_564) @p_558) :rule not_not)
(step t260 (cl @p_559 @p_558) :rule resolution :premises (t258 t259))
(step t261 (cl (not (! (= @p_304 @p_560) :named @p_561)) (not @p_304) @p_560) :rule equiv_pos2)
(step t262 (cl (= @p_303 (! (= @p_458 @p_201) :named @p_562))) :rule cong :premises (t94 t92))
(step t263 (cl (= @p_304 (! (not @p_562) :named @p_563))) :rule cong :premises (t262))
(step t264 (cl (= @p_562 @p_558)) :rule all_simplify)
(step t265 (cl (= @p_563 @p_560)) :rule cong :premises (t264))
(step t266 (cl @p_561) :rule trans :premises (t263 t265))
(step t267 (cl @p_304) :rule and :premises (t201))
(step t268 (cl @p_560) :rule resolution :premises (t261 t266 t267))
(step t269 (cl @p_559) :rule resolution :premises (t260 t268))
(step t270 (cl @p_557) :rule trans :premises (t256 t269))
(step t271 (cl @p_556) :rule resolution :premises (t237 t270))
(step t272 (cl @p_641) :rule resolution :premises (t171 t191 t209 t235 t271))
(step t273 (cl (= @p_660 (! (op @p_458 @p_37) :named @p_662))) :rule cong :premises (t272 t64))
(step t274 (cl (not (! (= @p_225 (! (= @p_404 @p_662) :named @p_663)) :named @p_664)) (not @p_225) @p_663) :rule equiv_pos2)
(step t275 (cl (= @p_225 (! (= e2 @p_224) :named @p_665))) :rule all_simplify)
(step t276 (cl (= @p_224 @p_662)) :rule cong :premises (t94 t32))
(step t277 (cl (= @p_665 @p_663)) :rule cong :premises (t47 t276))
(step t278 (cl @p_664) :rule trans :premises (t275 t277))
(step t279 (cl @p_273) :rule and :premises (t141))
(step t280 (cl @p_272) :rule and :premises (t279))
(step t281 (cl @p_271) :rule and :premises (t280))
(step t282 (cl @p_270) :rule and :premises (t281))
(step t283 (cl @p_225) :rule and :premises (t282))
(step t284 (cl @p_663) :rule resolution :premises (t274 t278 t283))
(step t285 (cl (= @p_662 @p_404)) :rule symm :premises (t284))
(step t286 (cl (= @p_660 @p_404)) :rule trans :premises (t273 t285))
(step t287 (cl @p_506) :rule refl)
(step t288 (cl (= @p_667 @p_546)) :rule cong :premises (t286 t287))
(step t289 (cl @p_666) :rule trans :premises (t147 t288))
(step t290 (cl (= @p_458 @p_404)) :rule symm :premises (t272))
(step t291 (cl (= @p_662 @p_660)) :rule cong :premises (t290 t64))
(step t292 (cl @p_661) :rule trans :premises (t284 t291))
(step t293 (cl (not (! (= @p_97 (! (or (! (= @p_405 @p_583) :named @p_630) @p_618 (! (= @p_404 @p_583) :named @p_627) (! (= @p_37 @p_583) :named @p_619) (! (= @p_277 @p_583) :named @p_600)) :named @p_633)) :named @p_634)) (not @p_97) @p_633) :rule equiv_pos2)
(step t294 (cl (= @p_97 (! (or @p_92 @p_93 @p_94 @p_96 @p_95) :named @p_640))) :rule all_simplify)
(step t295 (cl (= @p_92 (! (= e4 @p_91) :named @p_638))) :rule all_simplify)
(step t296 (cl (= @p_94 (! (= e2 @p_91) :named @p_636))) :rule all_simplify)
(step t297 (cl (= @p_96 (! (= e0 @p_91) :named @p_635))) :rule all_simplify)
(step t298 (cl (= @p_95 (! (= e1 @p_91) :named @p_604))) :rule all_simplify)
(step t299 (cl (= @p_640 (! (or @p_638 @p_637 @p_636 @p_635 @p_604) :named @p_639))) :rule cong :premises (t295 t17 t296 t297 t298))
(step t300 (cl (= @p_97 @p_639)) :rule trans :premises (t294 t299))
(step t301 (cl (= @p_638 @p_630)) :rule cong :premises (t51 t82))
(step t302 (cl (= @p_636 @p_627)) :rule cong :premises (t47 t82))
(step t303 (cl (= @p_635 @p_619)) :rule cong :premises (t32 t82))
(step t304 (cl (= @p_604 @p_600)) :rule cong :premises (t40 t82))
(step t305 (cl (= @p_639 @p_633)) :rule cong :premises (t301 t83 t302 t303 t304))
(step t306 (cl @p_634) :rule trans :premises (t300 t305))
(step t307 (cl @p_122) :rule and :premises (t164))
(step t308 (cl @p_97) :rule and :premises (t307))
(step t309 (cl @p_633) :rule resolution :premises (t293 t306 t308))
(step t310 (cl @p_630 @p_618 @p_627 @p_619 @p_600) :rule or :premises (t309))
(step t311 (cl @p_630 @p_627 @p_600 @p_619 @p_618) :rule reordering :premises (t310))
(step t312 (cl (not (! (= @p_289 (! (not @p_630) :named @p_631)) :named @p_632)) (not @p_289) @p_631) :rule equiv_pos2)
(step t313 (cl (= @p_288 @p_630)) :rule cong :premises (t48 t82))
(step t314 (cl @p_632) :rule cong :premises (t313))
(step t315 (cl @p_302) :rule and :premises (t198))
(step t316 (cl @p_301) :rule and :premises (t315))
(step t317 (cl @p_289) :rule and :premises (t316))
(step t318 (cl @p_631) :rule resolution :premises (t312 t314 t317))
(step t319 (cl (not (! (= @p_287 (! (not @p_627) :named @p_628)) :named @p_629)) (not @p_287) @p_628) :rule equiv_pos2)
(step t320 (cl (= @p_286 @p_627)) :rule cong :premises (t44 t82))
(step t321 (cl @p_629) :rule cong :premises (t320))
(step t322 (cl @p_287) :rule and :premises (t315))
(step t323 (cl @p_628) :rule resolution :premises (t319 t321 t322))
(step t324 (cl (not (! (= @p_242 (! (or (! (not @p_600) :named @p_601) (! (= @p_277 @p_406) :named @p_510)) :named @p_602)) :named @p_603)) (not @p_242) @p_602) :rule equiv_pos2)
(step t325 (cl (= @p_241 (! (not @p_604) :named @p_605))) :rule cong :premises (t298))
(step t326 (cl (= @p_23 (! (= e1 @p_19) :named @p_528))) :rule all_simplify)
(step t327 (cl (= @p_242 (! (or @p_605 @p_528) :named @p_606))) :rule cong :premises (t325 t326))
(step t328 (cl (= @p_605 @p_601)) :rule cong :premises (t304))
(step t329 (cl (= @p_528 @p_510)) :rule cong :premises (t40 t58))
(step t330 (cl (= @p_606 @p_602)) :rule cong :premises (t328 t329))
(step t331 (cl @p_603) :rule trans :premises (t327 t330))
(step t332 (cl @p_251) :rule and :premises (t244))
(step t333 (cl @p_250) :rule and :premises (t332))
(step t334 (cl @p_243) :rule and :premises (t333))
(step t335 (cl @p_242) :rule and :premises (t334))
(step t336 (cl @p_602) :rule resolution :premises (t324 t331 t335))
(step t337 (cl @p_601 @p_510) :rule or :premises (t336))
(step t338 (cl @p_510 @p_601) :rule reordering :premises (t337))
(step t339 (cl (= (! (= @p_510 false) :named @p_512) (! (not @p_510) :named @p_511))) :rule equiv_simplify)
(step t340 (cl (not @p_512) @p_511) :rule equiv1 :premises (t339))
(step t341 (cl @p_522) :rule refl)
(step t342 (cl (not (! (= @p_218 @p_407) :named @p_408)) (not @p_218) @p_407) :rule equiv_pos2)
(step t343 (cl (= @p_218 (! (= e2 @p_217) :named @p_409))) :rule all_simplify)
(step t344 (cl (= @p_217 @p_406)) :rule cong :premises (t48 t40))
(step t345 (cl (= @p_409 @p_407)) :rule cong :premises (t47 t344))
(step t346 (cl @p_408) :rule trans :premises (t343 t345))
(step t347 (cl @p_269) :rule and :premises (t279))
(step t348 (cl @p_268) :rule and :premises (t347))
(step t349 (cl @p_267) :rule and :premises (t348))
(step t350 (cl @p_218) :rule and :premises (t349))
(step t351 (cl @p_407) :rule resolution :premises (t342 t346 t350))
(step t352 (cl (= @p_406 @p_404)) :rule symm :premises (t351))
(step t353 (cl (= @p_510 (! (= @p_277 @p_404) :named @p_513))) :rule cong :premises (t341 t352))
(step t354 (cl (= (! (= @p_513 false) :named @p_514) (! (not @p_513) :named @p_515))) :rule equiv_simplify)
(step t355 (cl @p_514 (! (not @p_515) :named @p_521)) :rule equiv2 :premises (t354))
(step t356 (cl (not @p_521) @p_513) :rule not_not)
(step t357 (cl @p_514 @p_513) :rule resolution :premises (t355 t356))
(step t358 (cl (not (! (= @p_364 @p_515) :named @p_516)) (not @p_364) @p_515) :rule equiv_pos2)
(step t359 (cl (= @p_363 (! (= e2 e1) :named @p_517))) :rule all_simplify)
(step t360 (cl (= @p_364 (! (not @p_517) :named @p_518))) :rule cong :premises (t359))
(step t361 (cl (= @p_517 (! (= @p_404 @p_277) :named @p_519))) :rule cong :premises (t47 t40))
(step t362 (cl (= @p_518 (! (not @p_519) :named @p_520))) :rule cong :premises (t361))
(step t363 (cl (= @p_519 @p_513)) :rule all_simplify)
(step t364 (cl (= @p_520 @p_515)) :rule cong :premises (t363))
(step t365 (cl (= @p_518 @p_515)) :rule trans :premises (t362 t364))
(step t366 (cl @p_516) :rule trans :premises (t360 t365))
(step t367 (cl @p_369) :rule and :premises (a13))
(step t368 (cl @p_368) :rule and :premises (t367))
(step t369 (cl @p_367) :rule and :premises (t368))
(step t370 (cl @p_366) :rule and :premises (t369))
(step t371 (cl @p_365) :rule and :premises (t370))
(step t372 (cl @p_364) :rule and :premises (t371))
(step t373 (cl @p_515) :rule resolution :premises (t358 t366 t372))
(step t374 (cl @p_514) :rule resolution :premises (t357 t373))
(step t375 (cl @p_512) :rule trans :premises (t353 t374))
(step t376 (cl @p_511) :rule resolution :premises (t340 t375))
(step t377 (cl @p_601) :rule resolution :premises (t338 t376))
(step t378 (cl (= (! (= @p_619 false) :named @p_621) (! (not @p_619) :named @p_620))) :rule equiv_simplify)
(step t379 (cl (not @p_621) @p_620) :rule equiv1 :premises (t378))
(step t380 (cl (not (! (= @p_245 (! (or @p_386 @p_385) :named @p_387)) :named @p_388)) (not @p_245) @p_387) :rule equiv_pos2)
(step t381 (cl (= @p_244 (! (not @p_390) :named @p_391))) :rule cong :premises (t26))
(step t382 (cl (= @p_245 (! (or @p_391 @p_389) :named @p_392))) :rule cong :premises (t381 t21))
(step t383 (cl (= @p_391 @p_386)) :rule cong :premises (t98))
(step t384 (cl (= @p_392 @p_387)) :rule cong :premises (t383 t89))
(step t385 (cl @p_388) :rule trans :premises (t382 t384))
(step t386 (cl @p_249) :rule and :premises (t333))
(step t387 (cl @p_248) :rule and :premises (t386))
(step t388 (cl @p_247) :rule and :premises (t387))
(step t389 (cl @p_246) :rule and :premises (t388))
(step t390 (cl @p_245) :rule and :premises (t389))
(step t391 (cl @p_387) :rule resolution :premises (t380 t385 t390))
(step t392 (cl @p_386 @p_385) :rule or :premises (t391))
(step t393 (cl @p_385 @p_386) :rule reordering :premises (t392))
(step t394 (cl @p_385) :rule resolution :premises (t393 t221))
(step t395 (cl (= @p_583 @p_583)) :rule refl)
(step t396 (cl (= @p_619 (! (= @p_384 @p_583) :named @p_622))) :rule cong :premises (t394 t395))
(step t397 (cl (= (! (= @p_622 false) :named @p_623) (! (not @p_622) :named @p_624))) :rule equiv_simplify)
(step t398 (cl @p_623 (! (not @p_624) :named @p_626)) :rule equiv2 :premises (t397))
(step t399 (cl (not @p_626) @p_622) :rule not_not)
(step t400 (cl @p_623 @p_622) :rule resolution :premises (t398 t399))
(step t401 (cl (not (! (= @p_291 @p_624) :named @p_625)) (not @p_291) @p_624) :rule equiv_pos2)
(step t402 (cl (= @p_290 @p_622)) :rule cong :premises (t88 t82))
(step t403 (cl @p_625) :rule cong :premises (t402))
(step t404 (cl @p_300) :rule and :premises (t316))
(step t405 (cl @p_299) :rule and :premises (t404))
(step t406 (cl @p_298) :rule and :premises (t405))
(step t407 (cl @p_297) :rule and :premises (t406))
(step t408 (cl @p_296) :rule and :premises (t407))
(step t409 (cl @p_291) :rule and :premises (t408))
(step t410 (cl @p_624) :rule resolution :premises (t401 t403 t409))
(step t411 (cl @p_623) :rule resolution :premises (t400 t410))
(step t412 (cl @p_621) :rule trans :premises (t396 t411))
(step t413 (cl @p_620) :rule resolution :premises (t379 t412))
(step t414 (cl @p_618) :rule resolution :premises (t311 t318 t323 t377 t413))
(step t415 (cl (not (! (= @p_177 (! (or @p_428 (! (= @p_201 @p_405) :named @p_429) @p_459 (! (= @p_277 @p_405) :named @p_451) (! (= @p_405 @p_380) :named @p_441)) :named @p_464)) :named @p_465)) (not @p_177) @p_464) :rule equiv_pos2)
(step t416 (cl (= @p_177 (! (or @p_124 @p_130 @p_136 @p_149 @p_143) :named @p_474))) :rule all_simplify)
(step t417 (cl (= @p_130 (! (= e4 @p_129) :named @p_469))) :rule all_simplify)
(step t418 (cl (= @p_149 (! (= e4 @p_148) :named @p_467))) :rule all_simplify)
(step t419 (cl (= @p_143 (! (= e4 @p_142) :named @p_466))) :rule all_simplify)
(step t420 (cl (= @p_474 (! (or @p_470 @p_469 @p_468 @p_467 @p_466) :named @p_471))) :rule cong :premises (t22 t417 t150 t418 t419))
(step t421 (cl (= @p_177 @p_471)) :rule trans :premises (t416 t420))
(step t422 (cl (= @p_469 (! (= @p_405 @p_201) :named @p_472))) :rule cong :premises (t51 t92))
(step t423 (cl (= @p_467 (! (= @p_405 @p_277) :named @p_456))) :rule cong :premises (t51 t37))
(step t424 (cl (= @p_466 @p_441)) :rule cong :premises (t51 t97))
(step t425 (cl (= @p_471 (! (or @p_428 @p_472 @p_459 @p_456 @p_441) :named @p_473))) :rule cong :premises (t91 t422 t156 t423 t424))
(step t426 (cl (= @p_472 @p_429)) :rule all_simplify)
(step t427 (cl (= @p_459 @p_459)) :rule refl)
(step t428 (cl (= @p_456 @p_451)) :rule all_simplify)
(step t429 (cl (= @p_441 @p_441)) :rule refl)
(step t430 (cl (= @p_473 @p_464)) :rule cong :premises (t121 t426 t427 t428 t429))
(step t431 (cl (= @p_471 @p_464)) :rule trans :premises (t425 t430))
(step t432 (cl @p_465) :rule trans :premises (t421 t431))
(step t433 (cl @p_182) :rule and :premises (a1))
(step t434 (cl @p_181) :rule and :premises (t433))
(step t435 (cl @p_180) :rule and :premises (t434))
(step t436 (cl @p_179) :rule and :premises (t435))
(step t437 (cl @p_178) :rule and :premises (t436))
(step t438 (cl @p_177) :rule and :premises (t437))
(step t439 (cl @p_464) :rule resolution :premises (t415 t432 t438))
(step t440 (cl @p_428 @p_429 @p_459 @p_451 @p_441) :rule or :premises (t439))
(step t441 (cl @p_459 @p_451 @p_441 @p_429 @p_428) :rule reordering :premises (t440))
(step t442 (cl (not (! (= @p_362 (! (not @p_451) :named @p_452)) :named @p_453)) (not @p_362) @p_452) :rule equiv_pos2)
(step t443 (cl (= @p_361 (! (= e4 e1) :named @p_454))) :rule all_simplify)
(step t444 (cl (= @p_362 (! (not @p_454) :named @p_455))) :rule cong :premises (t443))
(step t445 (cl (= @p_454 @p_456)) :rule cong :premises (t51 t40))
(step t446 (cl (= @p_455 (! (not @p_456) :named @p_457))) :rule cong :premises (t445))
(step t447 (cl (= @p_457 @p_452)) :rule cong :premises (t428))
(step t448 (cl (= @p_455 @p_452)) :rule trans :premises (t446 t447))
(step t449 (cl @p_453) :rule trans :premises (t444 t448))
(step t450 (cl @p_362) :rule and :premises (t369))
(step t451 (cl @p_452) :rule resolution :premises (t442 t449 t450))
(step t452 (cl (= (! (= @p_441 false) :named @p_443) (! (not @p_441) :named @p_442))) :rule equiv_simplify)
(step t453 (cl (not @p_443) @p_442) :rule equiv1 :premises (t452))
(step t454 (cl @p_420) :rule refl)
(step t455 (cl (= @p_380 @p_37)) :rule symm :premises (t221))
(step t456 (cl (= @p_380 @p_384)) :rule trans :premises (t455 t394))
(step t457 (cl (= @p_441 (! (= @p_405 @p_384) :named @p_444))) :rule cong :premises (t454 t456))
(step t458 (cl (= (! (= @p_444 false) :named @p_445) (! (not @p_444) :named @p_446))) :rule equiv_simplify)
(step t459 (cl @p_445 (! (not @p_446) :named @p_450)) :rule equiv2 :premises (t458))
(step t460 (cl (not @p_450) @p_444) :rule not_not)
(step t461 (cl @p_445 @p_444) :rule resolution :premises (t459 t460))
(step t462 (cl (not (! (= @p_293 @p_446) :named @p_447)) (not @p_293) @p_446) :rule equiv_pos2)
(step t463 (cl (= @p_292 (! (= @p_384 @p_405) :named @p_448))) :rule cong :premises (t88 t48))
(step t464 (cl (= @p_293 (! (not @p_448) :named @p_449))) :rule cong :premises (t463))
(step t465 (cl (= @p_448 @p_444)) :rule all_simplify)
(step t466 (cl (= @p_449 @p_446)) :rule cong :premises (t465))
(step t467 (cl @p_447) :rule trans :premises (t464 t466))
(step t468 (cl @p_295) :rule and :premises (t408))
(step t469 (cl @p_294) :rule and :premises (t468))
(step t470 (cl @p_293) :rule and :premises (t469))
(step t471 (cl @p_446) :rule resolution :premises (t462 t467 t470))
(step t472 (cl @p_445) :rule resolution :premises (t461 t471))
(step t473 (cl @p_443) :rule trans :premises (t457 t472))
(step t474 (cl @p_442) :rule resolution :premises (t453 t473))
(step t475 (cl (= (! (= @p_429 false) :named @p_431) (! (not @p_429) :named @p_430))) :rule equiv_simplify)
(step t476 (cl (not @p_431) @p_430) :rule equiv1 :premises (t475))
(step t477 (cl @p_202) :rule symm :premises (t255))
(step t478 (cl (= @p_429 (! (= e3 @p_405) :named @p_432))) :rule cong :premises (t477 t454))
(step t479 (cl (= (! (= @p_432 false) :named @p_433) (! (not @p_432) :named @p_434))) :rule equiv_simplify)
(step t480 (cl @p_433 (! (not @p_434) :named @p_440)) :rule equiv2 :premises (t479))
(step t481 (cl (not @p_440) @p_432) :rule not_not)
(step t482 (cl @p_433 @p_432) :rule resolution :premises (t480 t481))
(step t483 (cl (not (! (= @p_360 @p_434) :named @p_435)) (not @p_360) @p_434) :rule equiv_pos2)
(step t484 (cl (= @p_359 (! (= e4 e3) :named @p_436))) :rule all_simplify)
(step t485 (cl (= @p_360 (! (not @p_436) :named @p_437))) :rule cong :premises (t484))
(step t486 (cl (= @p_436 (! (= @p_405 e3) :named @p_438))) :rule cong :premises (t51 t43))
(step t487 (cl (= @p_437 (! (not @p_438) :named @p_439))) :rule cong :premises (t486))
(step t488 (cl (= @p_438 @p_432)) :rule all_simplify)
(step t489 (cl (= @p_439 @p_434)) :rule cong :premises (t488))
(step t490 (cl (= @p_437 @p_434)) :rule trans :premises (t487 t489))
(step t491 (cl @p_435) :rule trans :premises (t485 t490))
(step t492 (cl @p_360) :rule and :premises (a13))
(step t493 (cl @p_434) :rule resolution :premises (t483 t491 t492))
(step t494 (cl @p_433) :rule resolution :premises (t482 t493))
(step t495 (cl @p_431) :rule trans :premises (t478 t494))
(step t496 (cl @p_430) :rule resolution :premises (t476 t495))
(step t497 (cl @p_428) :rule resolution :premises (t441 t191 t451 t474 t496))
(step t498 (cl (not (! (= @p_171 (! (or @p_600 @p_513 @p_451 (! (= @p_277 @p_384) :named @p_607) @p_599) :named @p_610)) :named @p_611)) (not @p_171) @p_610) :rule equiv_pos2)
(step t499 (cl (= @p_171 (! (or @p_95 @p_102 @p_108 @p_120 @p_114) :named @p_617))) :rule all_simplify)
(step t500 (cl (= @p_102 (! (= e1 @p_98) :named @p_615))) :rule all_simplify)
(step t501 (cl (= @p_108 (! (= e1 @p_104) :named @p_614))) :rule all_simplify)
(step t502 (cl (= @p_120 (! (= e1 @p_116) :named @p_613))) :rule all_simplify)
(step t503 (cl (= @p_617 (! (or @p_604 @p_615 @p_614 @p_613 @p_612) :named @p_616))) :rule cong :premises (t298 t500 t501 t502 t20))
(step t504 (cl (= @p_171 @p_616)) :rule trans :premises (t499 t503))
(step t505 (cl (= @p_615 @p_513)) :rule cong :premises (t40 t44))
(step t506 (cl (= @p_614 @p_451)) :rule cong :premises (t40 t48))
(step t507 (cl (= @p_613 @p_607)) :rule cong :premises (t40 t88))
(step t508 (cl (= @p_616 @p_610)) :rule cong :premises (t304 t505 t506 t507 t87))
(step t509 (cl @p_611) :rule trans :premises (t504 t508))
(step t510 (cl @p_176) :rule and :premises (t435))
(step t511 (cl @p_175) :rule and :premises (t510))
(step t512 (cl @p_174) :rule and :premises (t511))
(step t513 (cl @p_173) :rule and :premises (t512))
(step t514 (cl @p_172) :rule and :premises (t513))
(step t515 (cl @p_171) :rule and :premises (t514))
(step t516 (cl @p_610) :rule resolution :premises (t498 t509 t515))
(step t517 (cl @p_600 @p_513 @p_451 @p_607 @p_599) :rule or :premises (t516))
(step t518 (cl @p_607 @p_513 @p_451 @p_599 @p_600) :rule reordering :premises (t517))
(step t519 (cl (not (! (= @p_344 (! (not @p_607) :named @p_608)) :named @p_609)) (not @p_344) @p_608) :rule equiv_pos2)
(step t520 (cl (= @p_343 @p_607)) :rule cong :premises (t37 t88))
(step t521 (cl @p_609) :rule cong :premises (t520))
(step t522 (cl @p_354) :rule and :premises (t180))
(step t523 (cl @p_353) :rule and :premises (t522))
(step t524 (cl @p_352) :rule and :premises (t523))
(step t525 (cl @p_351) :rule and :premises (t524))
(step t526 (cl @p_350) :rule and :premises (t525))
(step t527 (cl @p_349) :rule and :premises (t526))
(step t528 (cl @p_348) :rule and :premises (t527))
(step t529 (cl @p_347) :rule and :premises (t528))
(step t530 (cl @p_346) :rule and :premises (t529))
(step t531 (cl @p_345) :rule and :premises (t530))
(step t532 (cl @p_344) :rule and :premises (t531))
(step t533 (cl @p_608) :rule resolution :premises (t519 t521 t532))
(step t534 (cl @p_599) :rule resolution :premises (t518 t533 t373 t451 t377))
(step t535 (cl (not (! (= @p_200 (! (= @p_405 (! (op @p_575 e3) :named @p_594)) :named @p_595)) :named @p_596)) (not @p_200) @p_595) :rule equiv_pos2)
(step t536 (cl (= @p_200 (! (= e4 @p_199) :named @p_597))) :rule all_simplify)
(step t537 (cl (= @p_199 @p_594)) :rule cong :premises (t62 t43))
(step t538 (cl (= @p_597 @p_595)) :rule cong :premises (t51 t537))
(step t539 (cl @p_596) :rule trans :premises (t536 t538))
(step t540 (cl @p_262) :rule and :premises (t140))
(step t541 (cl @p_200) :rule and :premises (t540))
(step t542 (cl @p_595) :rule resolution :premises (t535 t539 t541))
(step t543 (cl (not (! (= @p_167 (! (or @p_535 (! (= e3 @p_498) :named @p_536) (! (= e3 @p_546) :named @p_547) @p_555 @p_432) :named @p_565)) :named @p_566)) (not @p_167) @p_565) :rule equiv_pos2)
(step t544 (cl (= @p_167 (! (or @p_15 @p_45 @p_75 @p_137 @p_106) :named @p_573))) :rule all_simplify)
(step t545 (cl (= @p_45 (! (= e3 @p_43) :named @p_570))) :rule all_simplify)
(step t546 (cl (= @p_75 (! (= e3 @p_73) :named @p_569))) :rule all_simplify)
(step t547 (cl (= @p_106 (! (= e3 @p_104) :named @p_567))) :rule all_simplify)
(step t548 (cl (= @p_573 (! (or @p_571 @p_570 @p_569 @p_568 @p_567) :named @p_572))) :rule cong :premises (t4 t545 t546 t151 t547))
(step t549 (cl (= @p_167 @p_572)) :rule trans :premises (t544 t548))
(step t550 (cl (= @p_570 @p_536)) :rule cong :premises (t43 t66))
(step t551 (cl (= @p_569 @p_547)) :rule cong :premises (t43 t76))
(step t552 (cl (= @p_567 @p_432)) :rule cong :premises (t43 t48))
(step t553 (cl (= @p_572 @p_565)) :rule cong :premises (t57 t550 t551 t157 t552))
(step t554 (cl @p_566) :rule trans :premises (t549 t553))
(step t555 (cl @p_170) :rule and :premises (t434))
(step t556 (cl @p_169) :rule and :premises (t555))
(step t557 (cl @p_168) :rule and :premises (t556))
(step t558 (cl @p_167) :rule and :premises (t557))
(step t559 (cl @p_565) :rule resolution :premises (t543 t554 t558))
(step t560 (cl @p_535 @p_536 @p_547 @p_555 @p_432) :rule or :premises (t559))
(step t561 (cl @p_432 @p_555 @p_547 @p_536 @p_535) :rule reordering :premises (t560))
(step t562 (cl (= (! (= @p_547 false) :named @p_549) (! (not @p_547) :named @p_548))) :rule equiv_simplify)
(step t563 (cl (not @p_549) @p_548) :rule equiv1 :premises (t562))
(step t564 (cl (not (! (= @p_216 @p_413) :named @p_414)) (not @p_216) @p_413) :rule equiv_pos2)
(step t565 (cl (= @p_216 (! (= e3 @p_215) :named @p_415))) :rule all_simplify)
(step t566 (cl (= @p_215 @p_412)) :rule cong :premises (t44 t40))
(step t567 (cl (= @p_415 @p_413)) :rule cong :premises (t43 t566))
(step t568 (cl @p_414) :rule trans :premises (t565 t567))
(step t569 (cl @p_216) :rule and :premises (t348))
(step t570 (cl @p_413) :rule resolution :premises (t564 t568 t569))
(step t571 (cl (= @p_546 @p_546)) :rule refl)
(step t572 (cl (= @p_547 (! (= @p_412 @p_546) :named @p_550))) :rule cong :premises (t570 t571))
(step t573 (cl (= (! (= @p_550 false) :named @p_551) (! (not @p_550) :named @p_552))) :rule equiv_simplify)
(step t574 (cl @p_551 (! (not @p_552) :named @p_554)) :rule equiv2 :premises (t573))
(step t575 (cl (not @p_554) @p_550) :rule not_not)
(step t576 (cl @p_551 @p_550) :rule resolution :premises (t574 t575))
(step t577 (cl (not (! (= @p_279 @p_552) :named @p_553)) (not @p_279) @p_552) :rule equiv_pos2)
(step t578 (cl (= @p_278 @p_550)) :rule cong :premises (t78 t76))
(step t579 (cl @p_553) :rule cong :premises (t578))
(step t580 (cl @p_285) :rule and :premises (t197))
(step t581 (cl @p_284) :rule and :premises (t580))
(step t582 (cl @p_283) :rule and :premises (t581))
(step t583 (cl @p_282) :rule and :premises (t582))
(step t584 (cl @p_281) :rule and :premises (t583))
(step t585 (cl @p_280) :rule and :premises (t584))
(step t586 (cl @p_279) :rule and :premises (t585))
(step t587 (cl @p_552) :rule resolution :premises (t577 t579 t586))
(step t588 (cl @p_551) :rule resolution :premises (t576 t587))
(step t589 (cl @p_549) :rule trans :premises (t572 t588))
(step t590 (cl @p_548) :rule resolution :premises (t563 t589))
(step t591 (cl (= (! (= @p_536 false) :named @p_538) (! (not @p_536) :named @p_537))) :rule equiv_simplify)
(step t592 (cl (not @p_538) @p_537) :rule equiv1 :premises (t591))
(step t593 (cl (= @p_498 (! (op @p_412 @p_404) :named @p_505))) :rule cong :premises (t570 t287))
(step t594 (cl (not (! (= @p_210 (! (= @p_277 @p_505) :named @p_507)) :named @p_508)) (not @p_210) @p_507) :rule equiv_pos2)
(step t595 (cl (= @p_210 (! (= e1 @p_209) :named @p_509))) :rule all_simplify)
(step t596 (cl (= @p_209 @p_505)) :rule cong :premises (t78 t47))
(step t597 (cl (= @p_509 @p_507)) :rule cong :premises (t40 t596))
(step t598 (cl @p_508) :rule trans :premises (t595 t597))
(step t599 (cl @p_210) :rule and :premises (t145))
(step t600 (cl @p_507) :rule resolution :premises (t594 t598 t599))
(step t601 (cl (= @p_505 @p_277)) :rule symm :premises (t600))
(step t602 (cl (= @p_498 @p_277)) :rule trans :premises (t593 t601))
(step t603 (cl (= @p_536 (! (= @p_229 @p_277) :named @p_539))) :rule cong :premises (t254 t602))
(step t604 (cl (= (! (= @p_539 false) :named @p_540) (! (not @p_539) :named @p_541))) :rule equiv_simplify)
(step t605 (cl @p_540 (! (not @p_541) :named @p_545)) :rule equiv2 :premises (t604))
(step t606 (cl (not @p_545) @p_539) :rule not_not)
(step t607 (cl @p_540 @p_539) :rule resolution :premises (t605 t606))
(step t608 (cl (not (! (= @p_342 @p_541) :named @p_542)) (not @p_342) @p_541) :rule equiv_pos2)
(step t609 (cl (= @p_341 (! (= @p_277 @p_229) :named @p_543))) :rule cong :premises (t37 t70))
(step t610 (cl (= @p_342 (! (not @p_543) :named @p_544))) :rule cong :premises (t609))
(step t611 (cl (= @p_543 @p_539)) :rule all_simplify)
(step t612 (cl (= @p_544 @p_541)) :rule cong :premises (t611))
(step t613 (cl @p_542) :rule trans :premises (t610 t612))
(step t614 (cl @p_342) :rule and :premises (t530))
(step t615 (cl @p_541) :rule resolution :premises (t608 t613 t614))
(step t616 (cl @p_540) :rule resolution :premises (t607 t615))
(step t617 (cl @p_538) :rule trans :premises (t603 t616))
(step t618 (cl @p_537) :rule resolution :premises (t592 t617))
(step t619 (cl @p_535) :rule resolution :premises (t561 t493 t271 t590 t618))
(step t620 (cl (= @p_575 (! (op @p_494 @p_405) :named @p_577))) :rule cong :premises (t619 t454))
(step t621 (cl (not (! (= @p_194 (! (= @p_404 @p_577) :named @p_578)) :named @p_579)) (not @p_194) @p_578) :rule equiv_pos2)
(step t622 (cl (= @p_194 (! (= e2 @p_193) :named @p_580))) :rule all_simplify)
(step t623 (cl (= @p_193 @p_577)) :rule cong :premises (t56 t51))
(step t624 (cl (= @p_580 @p_578)) :rule cong :premises (t47 t623))
(step t625 (cl @p_579) :rule trans :premises (t622 t624))
(step t626 (cl @p_261) :rule and :premises (a10))
(step t627 (cl @p_260) :rule and :premises (t626))
(step t628 (cl @p_259) :rule and :premises (t627))
(step t629 (cl @p_194) :rule and :premises (t628))
(step t630 (cl @p_578) :rule resolution :premises (t621 t625 t629))
(step t631 (cl (= @p_577 @p_404)) :rule symm :premises (t630))
(step t632 (cl (= @p_575 @p_404)) :rule trans :premises (t620 t631))
(step t633 (cl (= @p_594 @p_592)) :rule cong :premises (t632 t43))
(step t634 (cl @p_593) :rule trans :premises (t542 t633))
(step t635 (cl (not (! (= @p_196 (! (= @p_277 (! (op @p_406 @p_405) :named @p_490)) :named @p_491)) :named @p_492)) (not @p_196) @p_491) :rule equiv_pos2)
(step t636 (cl (= @p_196 (! (= e1 @p_195) :named @p_493))) :rule all_simplify)
(step t637 (cl (= @p_195 @p_490)) :rule cong :premises (t58 t51))
(step t638 (cl (= @p_493 @p_491)) :rule cong :premises (t40 t637))
(step t639 (cl @p_492) :rule trans :premises (t636 t638))
(step t640 (cl @p_258) :rule and :premises (t628))
(step t641 (cl @p_196) :rule and :premises (t640))
(step t642 (cl @p_491) :rule resolution :premises (t635 t639 t641))
(step t643 (cl (= @p_490 @p_483)) :rule cong :premises (t352 t454))
(step t644 (cl @p_489) :rule trans :premises (t642 t643))
(step t645 (cl (not (! (= @p_214 (! (= @p_405 (! (op @p_583 @p_277) :named @p_584)) :named @p_589)) :named @p_590)) (not @p_214) @p_589) :rule equiv_pos2)
(step t646 (cl (= @p_214 (! (= e4 @p_213) :named @p_591))) :rule all_simplify)
(step t647 (cl (= @p_213 @p_584)) :rule cong :premises (t82 t40))
(step t648 (cl (= @p_591 @p_589)) :rule cong :premises (t51 t647))
(step t649 (cl @p_590) :rule trans :premises (t646 t648))
(step t650 (cl @p_214) :rule and :premises (t347))
(step t651 (cl @p_589) :rule resolution :premises (t645 t649 t650))
(step t652 (cl (not (! (= @p_161 (! (or (! (= @p_277 @p_416) :named @p_480) @p_479 (! (= @p_277 @p_494) :named @p_495) (! (= @p_277 @p_418) :named @p_523) @p_510) :named @p_526)) :named @p_527)) (not @p_161) @p_526) :rule equiv_pos2)
(step t653 (cl (= @p_161 (! (or @p_5 @p_11 @p_17 @p_29 @p_23) :named @p_534))) :rule all_simplify)
(step t654 (cl (= @p_5 (! (= e1 @p_1) :named @p_532))) :rule all_simplify)
(step t655 (cl (= @p_17 (! (= e1 @p_13) :named @p_530))) :rule all_simplify)
(step t656 (cl (= @p_29 (! (= e1 @p_25) :named @p_529))) :rule all_simplify)
(step t657 (cl (= @p_534 (! (or @p_532 @p_531 @p_530 @p_529 @p_528) :named @p_533))) :rule cong :premises (t654 t3 t655 t656 t326))
(step t658 (cl (= @p_161 @p_533)) :rule trans :premises (t653 t657))
(step t659 (cl (= @p_532 @p_480)) :rule cong :premises (t40 t52))
(step t660 (cl (= @p_530 @p_495)) :rule cong :premises (t40 t56))
(step t661 (cl (= @p_529 @p_523)) :rule cong :premises (t40 t60))
(step t662 (cl (= @p_533 @p_526)) :rule cong :premises (t659 t55 t660 t661 t329))
(step t663 (cl @p_527) :rule trans :premises (t658 t662))
(step t664 (cl @p_166) :rule and :premises (a1))
(step t665 (cl @p_165) :rule and :premises (t664))
(step t666 (cl @p_164) :rule and :premises (t665))
(step t667 (cl @p_163) :rule and :premises (t666))
(step t668 (cl @p_162) :rule and :premises (t667))
(step t669 (cl @p_161) :rule and :premises (t668))
(step t670 (cl @p_526) :rule resolution :premises (t652 t663 t669))
(step t671 (cl @p_480 @p_479 @p_495 @p_523 @p_510) :rule or :premises (t670))
(step t672 (cl @p_523 @p_510 @p_495 @p_479 @p_480) :rule reordering :premises (t671))
(step t673 (cl (not (! (= @p_340 (! (not @p_523) :named @p_524)) :named @p_525)) (not @p_340) @p_524) :rule equiv_pos2)
(step t674 (cl (= @p_339 @p_523)) :rule cong :premises (t37 t60))
(step t675 (cl @p_525) :rule cong :premises (t674))
(step t676 (cl @p_340) :rule and :premises (t529))
(step t677 (cl @p_524) :rule resolution :premises (t673 t675 t676))
(step t678 (cl (= (! (= @p_495 false) :named @p_497) (! (not @p_495) :named @p_496))) :rule equiv_simplify)
(step t679 (cl (not @p_497) @p_496) :rule equiv1 :premises (t678))
(step t680 (cl (= @p_412 e3)) :rule symm :premises (t570))
(step t681 (cl (= @p_505 @p_498)) :rule cong :premises (t680 t287))
(step t682 (cl @p_504) :rule trans :premises (t600 t681))
(step t683 (cl (= @p_494 @p_494)) :rule refl)
(step t684 (cl (= @p_495 (! (= @p_498 @p_494) :named @p_499))) :rule cong :premises (t682 t683))
(step t685 (cl (= (! (= @p_499 false) :named @p_500) (! (not @p_499) :named @p_501))) :rule equiv_simplify)
(step t686 (cl @p_500 (! (not @p_501) :named @p_503)) :rule equiv2 :premises (t685))
(step t687 (cl (not @p_503) @p_499) :rule not_not)
(step t688 (cl @p_500 @p_499) :rule resolution :premises (t686 t687))
(step t689 (cl (not (! (= @p_327 @p_501) :named @p_502)) (not @p_327) @p_501) :rule equiv_pos2)
(step t690 (cl (= @p_326 @p_499)) :rule cong :premises (t66 t56))
(step t691 (cl @p_502) :rule cong :premises (t690))
(step t692 (cl @p_327) :rule and :premises (t181))
(step t693 (cl @p_501) :rule resolution :premises (t689 t691 t692))
(step t694 (cl @p_500) :rule resolution :premises (t688 t693))
(step t695 (cl @p_497) :rule trans :premises (t684 t694))
(step t696 (cl @p_496) :rule resolution :premises (t679 t695))
(step t697 (cl (= (! (= @p_480 false) :named @p_482) (! (not @p_480) :named @p_481))) :rule equiv_simplify)
(step t698 (cl (not @p_482) @p_481) :rule equiv1 :premises (t697))
(step t699 (cl (= @p_416 @p_416)) :rule refl)
(step t700 (cl (= @p_480 (! (= @p_483 @p_416) :named @p_484))) :rule cong :premises (t644 t699))
(step t701 (cl (= (! (= @p_484 false) :named @p_485) (! (not @p_484) :named @p_486))) :rule equiv_simplify)
(step t702 (cl @p_485 (! (not @p_486) :named @p_488)) :rule equiv2 :premises (t701))
(step t703 (cl (not @p_488) @p_484) :rule not_not)
(step t704 (cl @p_485 @p_484) :rule resolution :premises (t702 t703))
(step t705 (cl (not (! (= @p_323 @p_486) :named @p_487)) (not @p_323) @p_486) :rule equiv_pos2)
(step t706 (cl (= @p_322 @p_484)) :rule cong :premises (t72 t52))
(step t707 (cl @p_487) :rule cong :premises (t706))
(step t708 (cl @p_325) :rule and :premises (t178))
(step t709 (cl @p_324) :rule and :premises (t708))
(step t710 (cl @p_323) :rule and :premises (t709))
(step t711 (cl @p_486) :rule resolution :premises (t705 t707 t710))
(step t712 (cl @p_485) :rule resolution :premises (t704 t711))
(step t713 (cl @p_482) :rule trans :premises (t700 t712))
(step t714 (cl @p_481) :rule resolution :premises (t698 t713))
(step t715 (cl @p_479) :rule resolution :premises (t672 t677 t376 t696 t714))
(step t716 (cl (= @p_583 (! (op @p_478 @p_405) :named @p_585))) :rule cong :premises (t715 t454))
(step t717 (cl (not (! (= @p_192 (! (= e3 @p_585) :named @p_586)) :named @p_587)) (not @p_192) @p_586) :rule equiv_pos2)
(step t718 (cl (= @p_192 (! (= e3 @p_191) :named @p_588))) :rule all_simplify)
(step t719 (cl (= @p_191 @p_585)) :rule cong :premises (t54 t51))
(step t720 (cl (= @p_588 @p_586)) :rule cong :premises (t43 t719))
(step t721 (cl @p_587) :rule trans :premises (t718 t720))
(step t722 (cl @p_192) :rule and :premises (t627))
(step t723 (cl @p_586) :rule resolution :premises (t717 t721 t722))
(step t724 (cl (= @p_585 e3)) :rule symm :premises (t723))
(step t725 (cl (= @p_583 e3)) :rule trans :premises (t716 t724))
(step t726 (cl (= @p_584 @p_581)) :rule cong :premises (t725 t341))
(step t727 (cl @p_582) :rule trans :premises (t651 t726))
(step t728 (cl (= @p_494 e3)) :rule symm :premises (t619))
(step t729 (cl (= @p_577 @p_575)) :rule cong :premises (t728 t454))
(step t730 (cl @p_576) :rule trans :premises (t630 t729))
(step t731 (cl (not (! (= @p_222 (! (= @p_405 (! (op @p_421 @p_37) :named @p_422)) :named @p_423)) :named @p_424)) (not @p_222) @p_423) :rule equiv_pos2)
(step t732 (cl (= @p_222 (! (= e4 @p_221) :named @p_425))) :rule all_simplify)
(step t733 (cl (= @p_221 @p_422)) :rule cong :premises (t90 t32))
(step t734 (cl (= @p_425 @p_423)) :rule cong :premises (t51 t733))
(step t735 (cl @p_424) :rule trans :premises (t732 t734))
(step t736 (cl @p_222) :rule and :premises (t280))
(step t737 (cl @p_423) :rule resolution :premises (t731 t735 t736))
(step t738 (cl (= @p_421 @p_405)) :rule symm :premises (t497))
(step t739 (cl (= @p_422 @p_418)) :rule cong :premises (t738 t64))
(step t740 (cl @p_574) :rule trans :premises (t737 t739))
(step t741 (cl (not (! (= @p_198 (! (= @p_37 (! (op @p_418 @p_405) :named @p_419)) :named @p_475)) :named @p_476)) (not @p_198) @p_475) :rule equiv_pos2)
(step t742 (cl (= @p_198 (! (= e0 @p_197) :named @p_477))) :rule all_simplify)
(step t743 (cl (= @p_197 @p_419)) :rule cong :premises (t60 t51))
(step t744 (cl (= @p_477 @p_475)) :rule cong :premises (t32 t743))
(step t745 (cl @p_476) :rule trans :premises (t742 t744))
(step t746 (cl @p_198) :rule and :premises (t640))
(step t747 (cl @p_475) :rule resolution :premises (t741 t745 t746))
(step t748 (cl (= @p_418 @p_422)) :rule cong :premises (t497 t64))
(step t749 (cl (= @p_422 @p_405)) :rule symm :premises (t737))
(step t750 (cl (= @p_418 @p_405)) :rule trans :premises (t748 t749))
(step t751 (cl (= @p_419 @p_416)) :rule cong :premises (t750 t454))
(step t752 (cl @p_417) :rule trans :premises (t747 t751))
(step t753 (cl) :rule resolution :premises (t134 t289 t292 t272 t414 t497 t534 t634 t644 t727 t682 t730 t740 t619 t715 t752 t570 t351 t255 t254 t394 t221))
