unsat
(assume a0 (! (and (distinct c_0 c_1) (or (! (not (p3 c_0)) :named @p_70) (! (not (! (p1 c_0 c_0) :named @p_80)) :named @p_81) (! (= c_0 c_0) :named @p_83) @p_80 @p_70 @p_70 @p_83 @p_81 (! (not (! (p4 c_0) :named @p_77)) :named @p_78) @p_83 (! (not (! (p14 c_0 c_0 c_0) :named @p_101)) :named @p_114)) (or (! (not (p3 c_0)) :named @p_70) (! (not (! (p1 c_0 c_0) :named @p_80)) :named @p_81) (! (= c_0 c_1) :named @p_74) (! (p1 c_1 c_0) :named @p_75) @p_70 (! (not (p3 c_1)) :named @p_62) @p_74 @p_81 (! (not (! (p4 c_0) :named @p_77)) :named @p_78) @p_83 (! (not (! (p14 c_0 c_0 c_1) :named @p_99)) :named @p_111)) (or (! (not (p3 c_0)) :named @p_70) (! (not (! (p1 c_0 c_1) :named @p_71)) :named @p_72) (! (= c_0 c_0) :named @p_83) @p_71 @p_70 @p_70 @p_83 @p_72 (! (not (! (p4 c_1) :named @p_65)) :named @p_66) @p_83 (! (not (! (p14 c_0 c_0 c_0) :named @p_101)) :named @p_114)) (or (! (not (p3 c_0)) :named @p_70) (! (not (! (p1 c_0 c_1) :named @p_71)) :named @p_72) (! (= c_0 c_1) :named @p_74) (! (p1 c_1 c_1) :named @p_63) @p_70 (! (not (p3 c_1)) :named @p_62) @p_74 @p_72 (! (not (! (p4 c_1) :named @p_65)) :named @p_66) @p_83 (! (not (! (p14 c_0 c_0 c_1) :named @p_99)) :named @p_111)) (or (! (not (p3 c_0)) :named @p_70) (! (not @p_75) :named @p_76) (! (= c_1 c_0) :named @p_82) @p_80 @p_62 @p_70 @p_83 @p_81 (! (not (! (p4 c_0) :named @p_77)) :named @p_78) @p_82 (! (not (! (p14 c_1 c_0 c_0) :named @p_94)) :named @p_113)) (or (! (not (p3 c_0)) :named @p_70) (! (not @p_75) :named @p_76) (! (= c_1 c_1) :named @p_69) (! (p1 c_1 c_0) :named @p_75) @p_62 (! (not (p3 c_1)) :named @p_62) @p_74 @p_81 (! (not (! (p4 c_0) :named @p_77)) :named @p_78) @p_82 (! (not (! (p14 c_1 c_0 c_1) :named @p_92)) :named @p_109)) (or (! (not (p3 c_0)) :named @p_70) (! (not @p_63) :named @p_64) (! (= c_1 c_0) :named @p_82) @p_71 @p_62 @p_70 @p_83 @p_72 (! (not (! (p4 c_1) :named @p_65)) :named @p_66) @p_82 (! (not (! (p14 c_1 c_0 c_0) :named @p_94)) :named @p_113)) (or @p_70 @p_64 @p_69 @p_63 @p_62 @p_62 @p_74 @p_72 @p_66 @p_82 @p_109) (or @p_62 (! (not (! (p1 c_0 c_0) :named @p_80)) :named @p_81) (! (= c_0 c_0) :named @p_83) @p_80 @p_70 @p_70 @p_82 @p_76 (! (not (! (p4 c_0) :named @p_77)) :named @p_78) @p_74 (! (not (! (p14 c_0 c_1 c_0) :named @p_97)) :named @p_107)) (or @p_62 (! (not (! (p1 c_0 c_0) :named @p_80)) :named @p_81) (! (= c_0 c_1) :named @p_74) (! (p1 c_1 c_0) :named @p_75) @p_70 (! (not (p3 c_1)) :named @p_62) @p_69 @p_76 (! (not (! (p4 c_0) :named @p_77)) :named @p_78) @p_74 (! (not (! (p14 c_0 c_1 c_1) :named @p_96)) :named @p_104)) (or @p_62 @p_72 @p_83 @p_71 @p_70 @p_70 @p_82 @p_64 @p_66 @p_74 @p_107) (or @p_62 @p_72 @p_74 @p_63 @p_70 @p_62 @p_69 @p_64 @p_66 @p_74 @p_104) (or @p_62 (! (not @p_75) :named @p_76) (! (= c_1 c_0) :named @p_82) @p_80 @p_62 @p_70 @p_82 @p_76 (! (not (! (p4 c_0) :named @p_77)) :named @p_78) @p_69 (! (not (! (p14 c_1 c_1 c_0) :named @p_91)) :named @p_105)) (or @p_62 (! (not @p_75) :named @p_76) (! (= c_1 c_1) :named @p_69) (! (p1 c_1 c_0) :named @p_75) @p_62 (! (not (p3 c_1)) :named @p_62) @p_69 @p_76 (! (not (! (p4 c_0) :named @p_77)) :named @p_78) @p_69 (! (not (! (p14 c_1 c_1 c_1) :named @p_89)) :named @p_102)) (or @p_62 @p_64 @p_82 @p_71 @p_62 @p_70 @p_82 @p_64 @p_66 @p_69 @p_105) (or @p_62 @p_64 @p_69 @p_63 @p_62 @p_62 @p_69 @p_64 @p_66 @p_69 @p_102) (or (! (not (p9 c_0)) :named @p_61) (p3 (! (f12 c_0) :named @p_31))) (or (! (not (p9 c_1)) :named @p_60) (p3 (! (f12 c_1) :named @p_30))) (! (or @p_78 (! (p1 (! (f6 c_0) :named @p_26) c_0) :named @p_147)) :named @p_148) (! (or @p_66 (! (p1 (! (f6 c_1) :named @p_22) c_1) :named @p_145)) :named @p_146) (or @p_70 (p4 (! (f2 c_0 c_0) :named @p_35)) @p_83 @p_70) (or @p_70 (p4 (! (f2 c_0 c_1) :named @p_34)) @p_74 @p_62) (or @p_62 (p4 (! (f2 c_1 c_0) :named @p_33)) @p_82 @p_70) (or @p_62 (p4 (! (f2 c_1 c_1) :named @p_32)) @p_69 @p_62) (or @p_78 @p_70 @p_83 @p_70 @p_83 @p_81 @p_81 @p_81 @p_78 @p_81) (or @p_78 @p_70 @p_83 @p_62 @p_74 @p_76 @p_81 @p_76 @p_78 @p_81) (or @p_78 @p_70 @p_82 @p_70 @p_83 @p_81 @p_72 @p_72 @p_66 @p_81) (or @p_78 @p_70 @p_82 @p_62 @p_74 @p_76 @p_72 @p_64 @p_66 @p_81) (or @p_78 @p_62 @p_83 @p_70 @p_82 @p_81 @p_76 @p_81 @p_78 @p_76) (or @p_78 @p_62 @p_83 @p_62 @p_69 @p_76 @p_76 @p_76 @p_78 @p_76) (or @p_78 @p_62 @p_82 @p_70 @p_82 @p_81 @p_64 @p_72 @p_66 @p_76) (or @p_78 @p_62 @p_82 @p_62 @p_69 @p_76 @p_64 @p_64 @p_66 @p_76) (or @p_66 @p_70 @p_74 @p_70 @p_83 @p_72 @p_81 @p_81 @p_78 @p_72) (or @p_66 @p_70 @p_74 @p_62 @p_74 @p_64 @p_81 @p_76 @p_78 @p_72) (or @p_66 @p_70 @p_69 @p_70 @p_83 @p_72 @p_72 @p_72 @p_66 @p_72) (or @p_66 @p_70 @p_69 @p_62 @p_74 @p_64 @p_72 @p_64 @p_66 @p_72) (or @p_66 @p_62 @p_74 @p_70 @p_82 @p_72 @p_76 @p_81 @p_78 @p_64) (or @p_66 @p_62 @p_74 @p_62 @p_69 @p_64 @p_76 @p_76 @p_78 @p_64) (or @p_66 @p_62 @p_69 @p_70 @p_82 @p_72 @p_64 @p_72 @p_66 @p_64) (or @p_66 @p_62 @p_69 @p_62 @p_69 @p_64 @p_64 @p_64 @p_66 @p_64) (or @p_83 @p_70 (! (p1 c_0 @p_35) :named @p_117) @p_70) (or @p_74 @p_62 (p1 c_1 @p_34) @p_70) (or @p_82 @p_70 (p1 c_0 @p_33) @p_62) (or @p_69 @p_62 (! (p1 c_1 @p_32) :named @p_116) @p_62) (or @p_83 @p_70 (p4 (! (f17 c_0 c_0 c_0) :named @p_59)) @p_83 @p_70 @p_114 @p_70 @p_83) (or @p_83 @p_70 (p4 (! (f17 c_0 c_1 c_0) :named @p_57)) @p_74 @p_62 @p_107 @p_70 @p_82) (or @p_74 @p_62 (p4 (! (f17 c_0 c_0 c_1) :named @p_58)) @p_83 @p_70 @p_111 @p_70 @p_74) (or @p_74 @p_62 (p4 (! (f17 c_0 c_1 c_1) :named @p_56)) @p_74 @p_62 @p_104 @p_70 @p_69) (or @p_82 @p_70 (p4 (! (f17 c_1 c_0 c_0) :named @p_55)) @p_82 @p_70 @p_113 @p_62 @p_83) (or @p_82 @p_70 (p4 (! (f17 c_1 c_1 c_0) :named @p_53)) @p_69 @p_62 @p_105 @p_62 @p_82) (or @p_69 @p_62 (p4 (! (f17 c_1 c_0 c_1) :named @p_54)) @p_82 @p_70 @p_109 @p_62 @p_74) (or @p_69 @p_62 (p4 (! (f17 c_1 c_1 c_1) :named @p_52)) @p_69 @p_62 @p_102 @p_62 @p_69) (or @p_83 @p_83 @p_83 @p_101 @p_70 @p_81 @p_81 @p_70 @p_78 @p_70 @p_81) (or @p_83 @p_83 @p_83 @p_101 @p_70 @p_72 @p_72 @p_70 @p_66 @p_70 @p_72) (or @p_83 @p_82 @p_74 @p_97 @p_70 @p_76 @p_81 @p_70 @p_78 @p_62 @p_81) (or @p_83 @p_82 @p_74 @p_97 @p_70 @p_64 @p_72 @p_70 @p_66 @p_62 @p_72) (or @p_74 @p_74 @p_83 @p_99 @p_70 @p_81 @p_81 @p_62 @p_78 @p_70 @p_76) (or @p_74 @p_74 @p_83 @p_99 @p_70 @p_72 @p_72 @p_62 @p_66 @p_70 @p_64) (or @p_74 @p_69 @p_74 @p_96 @p_70 @p_76 @p_81 @p_62 @p_78 @p_62 @p_76) (or @p_74 @p_69 @p_74 @p_96 @p_70 @p_64 @p_72 @p_62 @p_66 @p_62 @p_64) (or @p_82 @p_83 @p_82 @p_94 @p_62 @p_81 @p_76 @p_70 @p_78 @p_70 @p_81) (or @p_82 @p_83 @p_82 @p_94 @p_62 @p_72 @p_64 @p_70 @p_66 @p_70 @p_72) (or @p_82 @p_82 @p_69 @p_91 @p_62 @p_76 @p_76 @p_70 @p_78 @p_62 @p_81) (or @p_82 @p_82 @p_69 @p_91 @p_62 @p_64 @p_64 @p_70 @p_66 @p_62 @p_72) (or @p_69 @p_74 @p_82 @p_92 @p_62 @p_81 @p_76 @p_62 @p_78 @p_70 @p_76) (or @p_69 @p_74 @p_82 @p_92 @p_62 @p_72 @p_64 @p_62 @p_66 @p_70 @p_64) (or @p_69 @p_69 @p_69 @p_89 @p_62 @p_76 @p_76 @p_62 @p_78 @p_62 @p_76) (or @p_69 @p_69 @p_69 @p_89 @p_62 @p_64 @p_64 @p_62 @p_66 @p_62 @p_64) (or (! (not (p9 c_0)) :named @p_61) (! (not (p10 c_0 c_0)) :named @p_86) @p_70 (p3 (! (f16 c_0 c_0 c_0) :named @p_51)) @p_83 @p_61 @p_86) (or (! (not (p9 c_0)) :named @p_61) (! (not (p10 c_0 c_1)) :named @p_73) @p_70 (p3 (! (f16 c_0 c_1 c_0) :named @p_49)) @p_74 @p_60 @p_86) (or (! (not (p9 c_0)) :named @p_61) (! (not (p10 c_1 c_0)) :named @p_84) @p_62 (p3 (! (f16 c_0 c_0 c_1) :named @p_50)) @p_83 @p_61 @p_84) (or (! (not (p9 c_0)) :named @p_61) (! (not (p10 c_1 c_1)) :named @p_67) @p_62 (p3 (! (f16 c_0 c_1 c_1) :named @p_48)) @p_74 @p_60 @p_84) (or @p_60 @p_86 @p_70 (p3 (! (f16 c_1 c_0 c_0) :named @p_47)) @p_82 @p_61 @p_73) (or @p_60 @p_73 @p_70 (p3 (! (f16 c_1 c_1 c_0) :named @p_45)) @p_69 @p_60 @p_73) (or @p_60 @p_84 @p_62 (p3 (! (f16 c_1 c_0 c_1) :named @p_46)) @p_82 @p_61 @p_67) (or @p_60 @p_67 @p_62 (p3 (! (f16 c_1 c_1 c_1) :named @p_44)) @p_69 @p_60 @p_67) (! (or @p_78 (! (not (! (p1 (! (f7 c_0) :named @p_18) c_0) :named @p_142)) :named @p_143)) :named @p_144) (! (or @p_66 (! (not (! (p1 (! (f7 c_1) :named @p_14) c_1) :named @p_139)) :named @p_140)) :named @p_141) (or (! (not (p3 c_0)) :named @p_70) @p_83 (! (p10 c_0 (! (f13 c_0 c_0 c_0) :named @p_43)) :named @p_100) @p_70 @p_101 @p_83 @p_70 @p_83) (or @p_70 @p_83 (p10 c_1 (! (f13 c_1 c_0 c_0) :named @p_39)) @p_62 @p_94 @p_82 @p_70 @p_82) (or (! (not (p3 c_0)) :named @p_70) @p_82 (! (p10 c_0 (! (f13 c_0 c_1 c_0) :named @p_41)) :named @p_137) @p_70 @p_97 @p_74 @p_62 @p_83) (or (! (not (p3 c_0)) :named @p_70) @p_82 (! (p10 c_1 (! (f13 c_1 c_1 c_0) :named @p_37)) :named @p_90) @p_62 @p_91 @p_69 @p_62 @p_82) (or @p_62 @p_74 (! (p10 c_0 (! (f13 c_0 c_0 c_1) :named @p_42)) :named @p_98) @p_70 @p_99 @p_83 @p_70 @p_74) (or @p_62 @p_74 (! (p10 c_1 (! (f13 c_1 c_0 c_1) :named @p_38)) :named @p_138) @p_62 @p_92 @p_82 @p_70 @p_69) (or @p_62 @p_69 (p10 c_0 (! (f13 c_0 c_1 c_1) :named @p_40)) @p_70 @p_96 @p_74 @p_62 @p_74) (or @p_62 @p_69 (! (p10 c_1 (! (f13 c_1 c_1 c_1) :named @p_36)) :named @p_88) @p_62 @p_89 @p_69 @p_62 @p_69) (p9 c18) (or @p_83 @p_70 @p_83 @p_83 @p_101 @p_70 @p_100 @p_70) (or @p_83 @p_62 (! (= c_1 c_0) :named @p_82) @p_82 @p_94 @p_70 (! (p10 c_0 @p_39) :named @p_93) @p_70) (or @p_74 @p_70 @p_83 @p_74 @p_99 @p_70 (p10 c_1 @p_42) @p_62) (or @p_74 @p_62 @p_82 @p_69 @p_92 @p_70 @p_138 @p_62) (or @p_82 @p_70 @p_74 @p_83 @p_97 @p_62 @p_137 @p_70) (or @p_82 @p_62 @p_69 @p_82 @p_91 @p_62 (p10 c_0 @p_37) @p_70) (or @p_69 @p_70 (! (= c_0 c_1) :named @p_74) @p_74 @p_96 (! (not (p3 c_1)) :named @p_62) (! (p10 c_1 @p_40) :named @p_95) @p_62) (or @p_69 @p_62 @p_69 @p_69 @p_89 @p_62 @p_88 @p_62) (or (! (not (p3 c_0)) :named @p_70) @p_83 @p_70 (! (not (p10 c_0 c18)) :named @p_136) @p_101 @p_70 @p_136 @p_136 @p_83 @p_83) (or (! (not (p3 c_0)) :named @p_70) @p_83 @p_62 (! (not (p10 c_1 c18)) :named @p_135) @p_99 @p_70 @p_136 @p_136 @p_74 @p_74) (or @p_70 @p_82 @p_70 @p_136 @p_94 @p_62 @p_136 @p_135 @p_82 @p_83) (or @p_70 @p_82 @p_62 @p_135 @p_92 @p_62 @p_136 @p_135 @p_69 @p_74) (or @p_62 @p_74 @p_70 @p_136 @p_97 @p_70 @p_135 @p_136 @p_83 @p_82) (or @p_62 @p_74 @p_62 @p_135 @p_96 @p_70 @p_135 @p_136 @p_74 @p_69) (or @p_62 @p_69 @p_70 @p_136 @p_91 @p_62 @p_135 @p_135 @p_82 @p_82) (or @p_62 @p_69 @p_62 @p_135 @p_89 @p_62 @p_135 @p_135 @p_69 @p_69) (or @p_61 (p10 (! (f11 c_0) :named @p_13) c_0)) (or @p_60 (p10 (! (f11 c_1) :named @p_12) c_1)) (! (or @p_78 (! (p1 (! (f5 c_0) :named @p_8) c_0) :named @p_133)) :named @p_134) (! (or @p_66 (! (p1 (! (f5 c_1) :named @p_4) c_1) :named @p_131)) :named @p_132) (or @p_86 @p_70 (! (p10 @p_51 c_0) :named @p_122) @p_83 @p_86 @p_61 @p_61) (or @p_86 @p_70 (p10 @p_47 c_1) @p_82 @p_73 @p_60 @p_61) (or @p_73 @p_70 (p10 @p_49 c_0) @p_74 @p_86 @p_61 @p_60) (or @p_73 @p_70 (! (p10 @p_45 c_1) :named @p_121) @p_69 @p_73 @p_60 @p_60) (or @p_84 @p_62 (! (p10 @p_50 c_0) :named @p_120) @p_83 @p_84 @p_61 @p_61) (or @p_84 @p_62 (p10 @p_46 c_1) @p_82 @p_67 @p_60 @p_61) (or @p_67 @p_62 (p10 @p_48 c_0) @p_74 @p_84 @p_61 @p_60) (or @p_67 @p_62 (! (p10 @p_44 c_1) :named @p_119) @p_69 @p_67 @p_60 @p_60) (! (or @p_78 (! (not (! (= @p_8 @p_26) :named @p_128)) :named @p_129)) :named @p_130) (! (or @p_66 (! (not (! (= @p_4 @p_22) :named @p_125)) :named @p_126)) :named @p_127) (or @p_78 (p3 @p_8)) (or @p_66 (p3 @p_4)) (or @p_83 @p_70 (! (p1 c_0 @p_59) :named @p_115) @p_70 @p_114 @p_83 @p_70 @p_83) (or @p_83 @p_70 (! (p1 c_0 @p_55) :named @p_124) @p_70 @p_113 @p_82 @p_62 @p_82) (or @p_74 @p_70 (! (p1 c_0 @p_58) :named @p_112) @p_62 @p_111 @p_83 @p_70 @p_74) (or @p_74 @p_70 (p1 c_0 @p_54) @p_62 @p_109 @p_82 @p_62 @p_69) (or @p_82 @p_62 (p1 c_1 @p_57) @p_70 @p_107 @p_74 @p_70 @p_83) (or @p_82 @p_62 (! (p1 c_1 @p_53) :named @p_106) @p_70 @p_105 @p_69 @p_62 @p_82) (or @p_69 @p_62 (! (p1 c_1 @p_56) :named @p_123) @p_62 @p_104 @p_74 @p_70 @p_74) (or @p_69 @p_62 (! (p1 c_1 @p_52) :named @p_103) @p_62 @p_102 @p_69 @p_62 @p_69) (or @p_70 @p_70 @p_83 @p_83 (p9 @p_43) @p_83 @p_70 @p_101) (or @p_70 @p_70 @p_82 @p_83 (p9 @p_41) @p_74 @p_62 @p_97) (or @p_70 @p_62 @p_74 @p_74 (p9 @p_42) @p_83 @p_70 @p_99) (or @p_70 @p_62 @p_69 @p_74 (p9 @p_40) @p_74 @p_62 @p_96) (or @p_62 @p_70 @p_83 @p_82 (p9 @p_39) @p_82 @p_70 @p_94) (or @p_62 @p_70 @p_82 @p_82 (p9 @p_37) @p_69 @p_62 @p_91) (or @p_62 @p_62 @p_74 @p_69 (p9 @p_38) @p_82 @p_70 @p_92) (or @p_62 @p_62 @p_69 @p_69 (p9 @p_36) @p_69 @p_62 @p_89) (or @p_115 @p_114 @p_70 @p_83 @p_70 @p_70 @p_83 @p_83) (or (! (p1 c_0 @p_57) :named @p_108) @p_107 @p_70 @p_74 @p_62 @p_70 @p_82 @p_83) (or @p_124 @p_113 @p_62 @p_82 @p_70 @p_70 @p_83 @p_82) (or (p1 c_0 @p_53) @p_105 @p_62 @p_69 @p_62 @p_70 @p_82 @p_82) (or (p1 c_1 @p_58) @p_111 @p_70 @p_83 @p_70 @p_62 @p_74 @p_74) (or @p_123 @p_104 @p_70 @p_74 @p_62 @p_62 @p_69 @p_74) (or (! (p1 c_1 @p_54) :named @p_110) @p_109 @p_62 @p_82 @p_70 (! (not (p3 c_1)) :named @p_62) @p_74 @p_69) (or @p_103 @p_102 @p_62 @p_69 @p_62 @p_62 @p_69 @p_69) (or @p_86 @p_83 @p_61 @p_61 @p_122 @p_86 @p_70) (or @p_86 @p_82 @p_61 @p_60 (p10 @p_47 c_0) @p_73 @p_70) (or @p_73 @p_74 @p_60 @p_61 (p10 @p_49 c_1) @p_86 @p_70) (or @p_73 @p_69 @p_60 @p_60 @p_121 @p_73 @p_70) (or @p_84 @p_83 @p_61 @p_61 @p_120 @p_84 @p_62) (or @p_84 @p_82 @p_61 @p_60 (p10 @p_46 c_0) @p_67 @p_62) (or @p_67 @p_74 @p_60 @p_61 (p10 @p_48 c_1) @p_84 @p_62) (or @p_67 @p_69 @p_60 @p_60 @p_119 @p_67 @p_62) (! (p4 c8) :named @p_118) (or @p_70 @p_83 @p_117 @p_70) (or @p_70 @p_82 (p1 c_1 @p_33) @p_62) (or @p_62 @p_74 (p1 c_0 @p_34) @p_70) (or @p_62 @p_69 @p_116 @p_62) (or @p_83 @p_83 @p_70 @p_115 @p_114 @p_83 @p_70 @p_70) (or @p_83 @p_82 @p_70 (p1 c_1 @p_55) @p_113 @p_82 @p_70 @p_62) (or @p_74 @p_74 @p_70 @p_112 @p_111 @p_83 @p_62 @p_70) (or @p_74 @p_69 @p_70 @p_110 @p_109 @p_82 @p_62 @p_62) (or @p_82 @p_83 @p_62 @p_108 @p_107 @p_74 @p_70 @p_70) (or @p_82 @p_82 @p_62 @p_106 @p_105 @p_69 @p_70 @p_62) (or @p_69 @p_74 @p_62 (p1 c_0 @p_56) @p_104 @p_74 @p_62 @p_70) (or @p_69 @p_69 @p_62 @p_103 @p_102 @p_69 @p_62 @p_62) (or @p_101 @p_70 @p_86 @p_86 @p_61 @p_86 @p_86 @p_83 @p_70 @p_61 @p_86 @p_83 @p_70 @p_83 @p_83 @p_86) (or @p_101 @p_70 @p_86 @p_73 @p_61 @p_86 @p_73 @p_74 @p_70 @p_60 @p_73 @p_83 @p_70 @p_83 @p_83 @p_86) (or @p_101 @p_70 @p_73 @p_86 @p_60 @p_73 @p_86 @p_82 @p_70 @p_61 @p_86 @p_83 @p_70 @p_83 @p_83 @p_73) (or @p_101 @p_70 @p_73 @p_73 @p_60 @p_73 @p_73 @p_69 @p_70 @p_60 @p_73 @p_83 @p_70 @p_83 @p_83 @p_73) (or @p_99 @p_70 @p_84 @p_86 @p_61 @p_86 @p_84 @p_83 @p_62 @p_61 @p_86 @p_74 @p_70 @p_74 @p_83 @p_86) (or @p_99 @p_70 @p_84 @p_73 @p_61 @p_86 @p_67 @p_74 @p_62 @p_60 @p_73 @p_74 @p_70 @p_74 @p_83 @p_86) (or @p_99 @p_70 @p_67 @p_86 @p_60 @p_73 @p_84 @p_82 @p_62 @p_61 @p_86 @p_74 @p_70 @p_74 @p_83 @p_73) (or @p_99 @p_70 @p_67 @p_73 @p_60 @p_73 @p_67 @p_69 @p_62 @p_60 @p_73 @p_74 @p_70 @p_74 @p_83 @p_73) (or @p_97 @p_62 @p_86 @p_86 @p_61 @p_86 @p_86 @p_83 @p_70 @p_61 @p_84 @p_82 @p_70 @p_83 @p_74 @p_84) (or @p_97 @p_62 @p_86 @p_73 @p_61 @p_86 @p_73 @p_74 @p_70 @p_60 @p_67 @p_82 @p_70 @p_83 @p_74 @p_84) (or @p_97 @p_62 @p_73 @p_86 @p_60 @p_73 @p_86 @p_82 @p_70 @p_61 @p_84 @p_82 @p_70 @p_83 @p_74 @p_67) (or @p_97 @p_62 @p_73 @p_73 @p_60 @p_73 @p_73 @p_69 @p_70 @p_60 @p_67 @p_82 @p_70 @p_83 @p_74 @p_67) (or @p_96 @p_62 @p_84 @p_86 @p_61 @p_86 @p_84 @p_83 @p_62 @p_61 @p_84 @p_69 @p_70 @p_74 @p_74 @p_84) (or @p_96 @p_62 @p_84 @p_73 @p_61 @p_86 @p_67 @p_74 @p_62 @p_60 @p_67 @p_69 @p_70 @p_74 @p_74 @p_84) (or @p_96 @p_62 @p_67 @p_86 @p_60 @p_73 @p_84 @p_82 @p_62 @p_61 @p_84 @p_69 @p_70 @p_74 @p_74 @p_67) (or @p_96 @p_62 @p_67 @p_73 @p_60 @p_73 @p_67 @p_69 @p_62 @p_60 @p_67 @p_69 @p_70 @p_74 @p_74 @p_67) (or @p_94 @p_70 @p_86 @p_84 @p_61 @p_84 @p_86 @p_83 @p_70 @p_61 @p_86 @p_83 @p_62 @p_82 @p_82 @p_86) (or @p_94 @p_70 @p_86 @p_67 @p_61 @p_84 @p_73 @p_74 @p_70 @p_60 @p_73 @p_83 @p_62 @p_82 @p_82 @p_86) (or @p_94 @p_70 @p_73 @p_84 @p_60 @p_67 @p_86 @p_82 @p_70 @p_61 @p_86 @p_83 @p_62 @p_82 @p_82 @p_73) (or @p_94 @p_70 @p_73 @p_67 @p_60 @p_67 @p_73 @p_69 @p_70 @p_60 @p_73 @p_83 @p_62 @p_82 @p_82 @p_73) (or @p_92 @p_70 @p_84 @p_84 @p_61 @p_84 @p_84 @p_83 @p_62 @p_61 @p_86 @p_74 @p_62 @p_69 @p_82 @p_86) (or @p_92 @p_70 @p_84 @p_67 @p_61 @p_84 @p_67 @p_74 @p_62 @p_60 @p_73 @p_74 @p_62 @p_69 @p_82 @p_86) (or @p_92 @p_70 @p_67 @p_84 @p_60 @p_67 @p_84 @p_82 @p_62 @p_61 @p_86 @p_74 @p_62 @p_69 @p_82 @p_73) (or @p_92 @p_70 @p_67 @p_67 @p_60 @p_67 @p_67 @p_69 @p_62 @p_60 @p_73 @p_74 @p_62 @p_69 @p_82 @p_73) (or @p_91 @p_62 @p_86 @p_84 @p_61 @p_84 @p_86 @p_83 @p_70 @p_61 @p_84 @p_82 @p_62 @p_82 @p_69 @p_84) (or @p_91 @p_62 @p_86 @p_67 @p_61 @p_84 @p_73 @p_74 @p_70 @p_60 @p_67 @p_82 @p_62 @p_82 @p_69 @p_84) (or @p_91 @p_62 @p_73 @p_84 @p_60 @p_67 @p_86 @p_82 @p_70 @p_61 @p_84 @p_82 @p_62 @p_82 @p_69 @p_67) (or @p_91 @p_62 @p_73 @p_67 @p_60 @p_67 @p_73 @p_69 @p_70 @p_60 @p_67 @p_82 @p_62 @p_82 @p_69 @p_67) (or @p_89 @p_62 @p_84 @p_84 @p_61 @p_84 @p_84 @p_83 @p_62 @p_61 @p_84 @p_69 @p_62 @p_69 @p_69 @p_84) (or @p_89 @p_62 @p_84 @p_67 @p_61 @p_84 @p_67 @p_74 @p_62 @p_60 @p_67 @p_69 @p_62 @p_69 @p_69 @p_84) (or @p_89 @p_62 @p_67 @p_84 @p_60 @p_67 @p_84 @p_82 @p_62 @p_61 @p_84 @p_69 @p_62 @p_69 @p_69 @p_67) (or @p_89 @p_62 @p_67 @p_67 @p_60 @p_67 @p_67 @p_69 @p_62 @p_60 @p_67 @p_69 @p_62 @p_69 @p_69 @p_67) (or @p_83 @p_70 @p_83 @p_101 @p_70 @p_70 @p_83 @p_100) (or @p_83 @p_70 @p_74 @p_99 @p_70 @p_62 @p_74 @p_98) (or @p_74 @p_70 @p_82 @p_97 @p_62 @p_70 @p_83 (p10 c_1 @p_41)) (or @p_74 @p_70 @p_69 @p_96 @p_62 @p_62 @p_74 @p_95) (or @p_82 @p_62 @p_83 @p_94 @p_70 @p_70 @p_82 @p_93) (or @p_82 @p_62 @p_74 @p_92 @p_70 @p_62 @p_69 (p10 c_0 @p_38)) (or @p_69 @p_62 @p_82 @p_91 @p_62 @p_70 @p_82 @p_90) (or @p_69 @p_62 @p_69 @p_89 @p_62 @p_62 @p_69 @p_88) (or (p3 @p_18) @p_78) (or (p3 @p_14) @p_66) (or (not (= c_0 @p_51)) @p_86 @p_86 @p_61 @p_70 @p_83 @p_61) (or (not (= c_0 @p_49)) @p_86 @p_73 @p_61 @p_70 @p_74 @p_60) (or (not (= c_0 @p_47)) @p_73 @p_86 @p_60 @p_70 @p_82 @p_61) (or (not (= c_0 @p_45)) @p_73 @p_73 @p_60 @p_70 @p_69 @p_60) (or (not (= c_1 @p_50)) @p_84 @p_84 @p_61 @p_62 @p_83 @p_61) (or (not (= c_1 @p_48)) @p_84 @p_67 @p_61 @p_62 @p_74 @p_60) (or (not (= c_1 @p_46)) @p_67 @p_84 @p_60 @p_62 @p_82 @p_61) (or (not (= c_1 @p_44)) @p_67 @p_67 @p_60 @p_62 @p_69 @p_60) (or (p3 @p_26) @p_78) (or (p3 @p_22) @p_66) (or (! (not (p9 c_0)) :named @p_61) (! (not (! (p1 c_0 c_0) :named @p_80)) :named @p_81) (! (= c_0 c_0) :named @p_83) @p_86 (! (p15 c_0 c_0) :named @p_87) @p_86 @p_70 @p_78 @p_81 @p_70) (or @p_61 @p_81 @p_82 @p_86 @p_87 @p_84 @p_70 @p_78 @p_76 @p_62) (or (! (not (p9 c_0)) :named @p_61) (! (not (! (p1 c_0 c_1) :named @p_71)) :named @p_72) (! (= c_0 c_0) :named @p_83) @p_86 (! (p15 c_1 c_0) :named @p_85) @p_86 @p_70 @p_66 @p_72 @p_70) (or @p_61 @p_72 @p_82 @p_86 @p_85 @p_84 @p_70 @p_66 @p_64 @p_62) (or @p_61 @p_76 @p_74 @p_84 @p_87 @p_86 @p_62 @p_78 @p_81 @p_70) (or @p_61 @p_76 @p_69 @p_84 @p_87 @p_84 @p_62 @p_78 @p_76 @p_62) (or @p_61 @p_64 @p_74 @p_84 @p_85 @p_86 @p_62 @p_66 @p_72 @p_70) (or @p_61 @p_64 @p_69 @p_84 @p_85 @p_84 @p_62 @p_66 @p_64 @p_62) (or (! (not (p9 c_1)) :named @p_60) (! (not (! (p1 c_0 c_0) :named @p_80)) :named @p_81) (! (= c_0 c_0) :named @p_83) @p_73 (! (p15 c_0 c_1) :named @p_79) @p_73 @p_70 @p_78 @p_81 @p_70) (or @p_60 @p_81 @p_82 @p_73 @p_79 @p_67 @p_70 @p_78 @p_76 @p_62) (or (! (not (p9 c_1)) :named @p_60) (! (not (! (p1 c_0 c_1) :named @p_71)) :named @p_72) (! (= c_0 c_0) :named @p_83) @p_73 (! (p15 c_1 c_1) :named @p_68) @p_73 @p_70 @p_66 @p_72 @p_70) (or @p_60 @p_72 @p_82 @p_73 @p_68 @p_67 @p_70 @p_66 @p_64 @p_62) (or @p_60 @p_76 @p_74 @p_67 @p_79 @p_73 @p_62 @p_78 @p_81 @p_70) (or @p_60 @p_76 @p_69 @p_67 @p_79 @p_67 @p_62 @p_78 @p_76 @p_62) (or @p_60 @p_64 @p_74 @p_67 @p_68 @p_73 @p_62 @p_66 @p_72 @p_70) (or @p_60 @p_64 @p_69 @p_67 @p_68 @p_67 @p_62 @p_66 @p_64 @p_62) (or @p_61 (p3 @p_13)) (or @p_60 (p3 @p_12)) (or @p_61 (not (p10 @p_31 c_0))) (or @p_60 (not (p10 @p_30 c_1))) (or (= @p_59 c_0) (= @p_59 c_1)) (or (= @p_58 c_0) (= @p_58 c_1)) (or (= @p_57 c_0) (= @p_57 c_1)) (or (= @p_56 c_0) (= @p_56 c_1)) (or (= @p_55 c_0) (= @p_55 c_1)) (or (= @p_54 c_0) (= @p_54 c_1)) (or (= @p_53 c_0) (= @p_53 c_1)) (or (= @p_52 c_0) (= @p_52 c_1)) (or (= @p_51 c_0) (= @p_51 c_1)) (or (= @p_50 c_0) (= @p_50 c_1)) (or (= @p_49 c_0) (= @p_49 c_1)) (or (= @p_48 c_0) (= @p_48 c_1)) (or (= @p_47 c_0) (= @p_47 c_1)) (or (= @p_46 c_0) (= @p_46 c_1)) (or (= @p_45 c_0) (= @p_45 c_1)) (or (= @p_44 c_0) (= @p_44 c_1)) (or (= @p_43 c_0) (= @p_43 c_1)) (or (= @p_42 c_0) (= @p_42 c_1)) (or (= @p_41 c_0) (= @p_41 c_1)) (or (= @p_40 c_0) (= @p_40 c_1)) (or (= @p_39 c_0) (= @p_39 c_1)) (or (= @p_38 c_0) (= @p_38 c_1)) (or (= @p_37 c_0) (= @p_37 c_1)) (or (= @p_36 c_0) (= @p_36 c_1)) (or (= @p_35 c_0) (= @p_35 c_1)) (or (= @p_34 c_0) (= @p_34 c_1)) (or (= @p_33 c_0) (= @p_33 c_1)) (or (= @p_32 c_0) (= @p_32 c_1)) (or (= @p_31 c_0) (= @p_31 c_1)) (or (= @p_30 c_0) (= @p_30 c_1)) (! (or (! (= @p_26 c_0) :named @p_28) (! (= @p_26 c_1) :named @p_27)) :named @p_29) (! (or (! (= @p_22 c_0) :named @p_24) (! (= @p_22 c_1) :named @p_23)) :named @p_25) (! (or (! (= @p_18 c_0) :named @p_20) (! (= @p_18 c_1) :named @p_19)) :named @p_21) (! (or (! (= @p_14 c_0) :named @p_16) (! (= @p_14 c_1) :named @p_15)) :named @p_17) (or (= @p_13 c_0) (= @p_13 c_1)) (or (= @p_12 c_0) (= @p_12 c_1)) (! (or (! (= @p_8 c_0) :named @p_10) (! (= @p_8 c_1) :named @p_9)) :named @p_11) (! (or (! (= @p_4 c_0) :named @p_6) (! (= @p_4 c_1) :named @p_5)) :named @p_7) (or (= c18 c_0) (= c18 c_1)) (! (or (! (= c8 c_0) :named @p_2) (! (= c8 c_1) :named @p_1)) :named @p_3)) :named @p_149))
(step t0 (cl (not (! (= (! (or (! (not @p_140) :named @p_246) (! (not (! (= c_1 @p_14) :named @p_239)) :named @p_303) (! (not (! (= c_1 @p_4) :named @p_270)) :named @p_271) (! (not @p_131) :named @p_288)) :named @p_320) (! (or @p_139 @p_303 @p_271 @p_288) :named @p_319)) :named @p_324)) (not @p_320) @p_319) :rule equiv_pos2)
(step t1 (cl (= @p_246 @p_139)) :rule all_simplify)
(step t2 (cl (= @p_303 @p_303)) :rule refl)
(step t3 (cl (= @p_271 @p_271)) :rule refl)
(step t4 (cl (= @p_288 @p_288)) :rule refl)
(step t5 (cl @p_324) :rule cong :premises (t1 t2 t3 t4))
(step t6 (cl (! (=> (! (and @p_140 @p_239 @p_270 @p_131) :named @p_321) false) :named @p_323) @p_321) :rule implies_neg1)
(anchor :step t7)
(assume t7.a0 @p_140)
(assume t7.a1 @p_239)
(assume t7.a2 @p_270)
(assume t7.a3 @p_131)
(step t7.t0 (cl (! (not (! (= (! (= true false) :named @p_202) false) :named @p_205)) :named @p_206) (! (not @p_202) :named @p_207) false) :rule equiv_pos2)
(step t7.t1 (cl @p_205) :rule all_simplify)
(step t7.t2 (cl (! (= (! (= @p_131 true) :named @p_298) @p_131) :named @p_299)) :rule equiv_simplify)
(step t7.t3 (cl @p_298 @p_288) :rule equiv2 :premises (t7.t2))
(step t7.t4 (cl @p_298) :rule resolution :premises (t7.t3 t7.a3))
(step t7.t5 (cl (! (= true @p_131) :named @p_297)) :rule symm :premises (t7.t4))
(step t7.t6 (cl @p_5) :rule symm :premises (t7.a2))
(step t7.t7 (cl @p_69) :rule refl)
(step t7.t8 (cl (= @p_131 @p_63)) :rule cong :premises (t7.t6 t7.t7))
(step t7.t9 (cl @p_15) :rule symm :premises (t7.a1))
(step t7.t10 (cl @p_239) :rule symm :premises (t7.t9))
(step t7.t11 (cl (! (= @p_63 @p_139) :named @p_310)) :rule cong :premises (t7.t10 t7.t7))
(step t7.t12 (cl (! (= (! (= @p_139 false) :named @p_255) @p_140) :named @p_256)) :rule equiv_simplify)
(step t7.t13 (cl @p_255 @p_246) :rule equiv2 :premises (t7.t12))
(step t7.t14 (cl (! (not @p_246) :named @p_251) @p_139) :rule not_not)
(step t7.t15 (cl @p_255 @p_139) :rule resolution :premises (t7.t13 t7.t14))
(step t7.t16 (cl @p_255) :rule resolution :premises (t7.t15 t7.a0))
(step t7.t17 (cl @p_202) :rule trans :premises (t7.t5 t7.t8 t7.t11 t7.t16))
(step t7.t18 (cl false) :rule resolution :premises (t7.t0 t7.t1 t7.t17))
(step t7 (cl @p_246 @p_303 @p_271 @p_288 false) :rule subproof :discharge (t7.a0 t7.a1 t7.a2 t7.a3))
(step t8 (cl (! (not @p_321) :named @p_322) @p_140) :rule and_pos)
(step t9 (cl @p_322 @p_239) :rule and_pos)
(step t10 (cl @p_322 @p_270) :rule and_pos)
(step t11 (cl @p_322 @p_131) :rule and_pos)
(step t12 (cl false @p_322 @p_322 @p_322 @p_322) :rule resolution :premises (t7 t8 t9 t10 t11))
(step t13 (cl @p_322 @p_322 @p_322 @p_322 false) :rule reordering :premises (t12))
(step t14 (cl @p_322 false) :rule contraction :premises (t13))
(step t15 (cl @p_323 false) :rule resolution :premises (t6 t14))
(step t16 (cl @p_323 (! (not false) :named @p_170)) :rule implies_neg2)
(step t17 (cl @p_323 @p_323) :rule resolution :premises (t15 t16))
(step t18 (cl @p_323) :rule contraction :premises (t17))
(step t19 (cl (= @p_323 @p_322)) :rule implies_simplify)
(step t20 (cl (not @p_323) @p_322) :rule equiv1 :premises (t19))
(step t21 (cl @p_322) :rule resolution :premises (t18 t20))
(step t22 (cl @p_246 @p_303 @p_271 @p_288) :rule not_and :premises (t21))
(step t23 (cl @p_320 @p_251) :rule or_neg)
(step t24 (cl @p_320 (! (not @p_303) :named @p_306)) :rule or_neg)
(step t25 (cl @p_320 (! (not @p_271) :named @p_275)) :rule or_neg)
(step t26 (cl @p_320 (! (not @p_288) :named @p_292)) :rule or_neg)
(step t27 (cl @p_320 @p_320 @p_320 @p_320) :rule resolution :premises (t22 t23 t24 t25 t26))
(step t28 (cl @p_320) :rule contraction :premises (t27))
(step t29 (cl @p_319) :rule resolution :premises (t0 t5 t28))
(step t30 (cl @p_139 @p_303 @p_271 @p_288) :rule or :premises (t29))
(step t31 (cl @p_139 @p_288 @p_303 @p_271) :rule reordering :premises (t30))
(step t32 (cl (not (! (= @p_7 (! (or (! (= c_0 @p_4) :named @p_285) @p_270) :named @p_286)) :named @p_287)) (not @p_7) @p_286) :rule equiv_pos2)
(step t33 (cl (= @p_6 @p_285)) :rule all_simplify)
(step t34 (cl (= @p_5 @p_270)) :rule all_simplify)
(step t35 (cl @p_287) :rule cong :premises (t33 t34))
(step t36 (cl @p_7) :rule and :premises (a0))
(step t37 (cl @p_286) :rule resolution :premises (t32 t35 t36))
(step t38 (cl @p_285 @p_270) :rule or :premises (t37))
(step t39 (cl (not (! (= (! (or (! (not (! (not (! (= @p_22 @p_4) :named @p_265)) :named @p_266)) :named @p_273) (! (not (! (= c_0 @p_22) :named @p_243)) :named @p_244) (! (not @p_285) :named @p_289)) :named @p_313) (! (or @p_265 @p_244 @p_289) :named @p_312)) :named @p_318)) (not @p_313) @p_312) :rule equiv_pos2)
(step t40 (cl (= @p_273 @p_265)) :rule all_simplify)
(step t41 (cl (= @p_244 @p_244)) :rule refl)
(step t42 (cl (= @p_289 @p_289)) :rule refl)
(step t43 (cl @p_318) :rule cong :premises (t40 t41 t42))
(step t44 (cl (! (and @p_266 @p_243) :named @p_314) @p_273 @p_244) :rule and_neg)
(step t45 (cl (! (=> @p_314 @p_289) :named @p_316) @p_314) :rule implies_neg1)
(anchor :step t46)
(assume t46.a0 @p_266)
(assume t46.a1 @p_243)
(step t46.t0 (cl (= (! (= @p_285 false) :named @p_317) @p_289)) :rule equiv_simplify)
(step t46.t1 (cl (not @p_317) @p_289) :rule equiv1 :premises (t46.t0))
(step t46.t2 (cl @p_24) :rule symm :premises (t46.a1))
(step t46.t3 (cl @p_243) :rule symm :premises (t46.t2))
(step t46.t4 (cl (= @p_4 @p_4)) :rule refl)
(step t46.t5 (cl (= @p_285 @p_265)) :rule cong :premises (t46.t3 t46.t4))
(step t46.t6 (cl (! (= (! (= @p_265 false) :named @p_282) @p_266) :named @p_283)) :rule equiv_simplify)
(step t46.t7 (cl @p_282 @p_273) :rule equiv2 :premises (t46.t6))
(step t46.t8 (cl (! (not @p_273) :named @p_277) @p_265) :rule not_not)
(step t46.t9 (cl @p_282 @p_265) :rule resolution :premises (t46.t7 t46.t8))
(step t46.t10 (cl @p_282) :rule resolution :premises (t46.t9 t46.a0))
(step t46.t11 (cl @p_317) :rule trans :premises (t46.t5 t46.t10))
(step t46.t12 (cl @p_289) :rule resolution :premises (t46.t1 t46.t11))
(step t46 (cl @p_273 @p_244 @p_289) :rule subproof :discharge (t46.a0 t46.a1))
(step t47 (cl (! (not @p_314) :named @p_315) @p_266) :rule and_pos)
(step t48 (cl @p_315 @p_243) :rule and_pos)
(step t49 (cl @p_289 @p_315 @p_315) :rule resolution :premises (t46 t47 t48))
(step t50 (cl @p_315 @p_315 @p_289) :rule reordering :premises (t49))
(step t51 (cl @p_315 @p_289) :rule contraction :premises (t50))
(step t52 (cl @p_316 @p_289) :rule resolution :premises (t45 t51))
(step t53 (cl @p_316 (! (not @p_289) :named @p_293)) :rule implies_neg2)
(step t54 (cl @p_316 @p_316) :rule resolution :premises (t52 t53))
(step t55 (cl @p_316) :rule contraction :premises (t54))
(step t56 (cl @p_315 @p_289) :rule implies :premises (t55))
(step t57 (cl @p_273 @p_244 @p_289) :rule resolution :premises (t44 t56))
(step t58 (cl @p_313 @p_277) :rule or_neg)
(step t59 (cl @p_313 (! (not @p_244) :named @p_249)) :rule or_neg)
(step t60 (cl @p_313 @p_293) :rule or_neg)
(step t61 (cl @p_313 @p_313 @p_313) :rule resolution :premises (t57 t58 t59 t60))
(step t62 (cl @p_313) :rule contraction :premises (t61))
(step t63 (cl @p_312) :rule resolution :premises (t39 t43 t62))
(step t64 (cl @p_265 @p_244 @p_289) :rule or :premises (t63))
(step t65 (cl (not (! (= @p_127 (! (or @p_66 @p_266) :named @p_267)) :named @p_268)) (not @p_127) @p_267) :rule equiv_pos2)
(step t66 (cl (= @p_66 @p_66)) :rule refl)
(step t67 (cl (= @p_125 @p_265)) :rule all_simplify)
(step t68 (cl (= @p_126 @p_266)) :rule cong :premises (t67))
(step t69 (cl @p_268) :rule cong :premises (t66 t68))
(step t70 (cl @p_127) :rule and :premises (a0))
(step t71 (cl @p_267) :rule resolution :premises (t65 t69 t70))
(step t72 (cl @p_66 @p_266) :rule or :premises (t71))
(step t73 (cl (! (and @p_118 @p_1) :named @p_235) (! (not @p_118) :named @p_228) (! (not @p_1) :named @p_234)) :rule and_neg)
(step t74 (cl (! (=> @p_235 @p_65) :named @p_237) @p_235) :rule implies_neg1)
(anchor :step t75)
(assume t75.a0 @p_118)
(assume t75.a1 @p_1)
(step t75.t0 (cl (= (! (= @p_65 true) :named @p_238) @p_65)) :rule equiv_simplify)
(step t75.t1 (cl (not @p_238) @p_65) :rule equiv1 :premises (t75.t0))
(step t75.t2 (cl (= c_1 c8)) :rule symm :premises (t75.a1))
(step t75.t3 (cl (= @p_65 @p_118)) :rule cong :premises (t75.t2))
(step t75.t4 (cl (= (! (= @p_118 true) :named @p_233) @p_118)) :rule equiv_simplify)
(step t75.t5 (cl @p_233 @p_228) :rule equiv2 :premises (t75.t4))
(step t75.t6 (cl @p_118) :rule and :premises (a0))
(step t75.t7 (cl @p_233) :rule resolution :premises (t75.t5 t75.t6))
(step t75.t8 (cl @p_238) :rule trans :premises (t75.t3 t75.t7))
(step t75.t9 (cl @p_65) :rule resolution :premises (t75.t1 t75.t8))
(step t75 (cl @p_228 @p_234 @p_65) :rule subproof :discharge (t75.a0 t75.a1))
(step t76 (cl (! (not @p_235) :named @p_236) @p_118) :rule and_pos)
(step t77 (cl @p_236 @p_1) :rule and_pos)
(step t78 (cl @p_65 @p_236 @p_236) :rule resolution :premises (t75 t76 t77))
(step t79 (cl @p_236 @p_236 @p_65) :rule reordering :premises (t78))
(step t80 (cl @p_236 @p_65) :rule contraction :premises (t79))
(step t81 (cl @p_237 @p_65) :rule resolution :premises (t74 t80))
(step t82 (cl @p_237 @p_66) :rule implies_neg2)
(step t83 (cl @p_237 @p_237) :rule resolution :premises (t81 t82))
(step t84 (cl @p_237) :rule contraction :premises (t83))
(step t85 (cl @p_236 @p_65) :rule implies :premises (t84))
(step t86 (cl @p_228 @p_234 @p_65) :rule resolution :premises (t73 t85))
(step t87 (cl @p_65 @p_228 @p_234) :rule reordering :premises (t86))
(step t88 (cl @p_118) :rule and :premises (a0))
(step t89 (cl @p_3) :rule and :premises (a0))
(step t90 (cl @p_2 @p_1) :rule or :premises (t89))
(step t91 (cl (! (and @p_118 @p_2) :named @p_229) @p_228 (! (not @p_2) :named @p_150)) :rule and_neg)
(step t92 (cl (! (=> @p_229 @p_77) :named @p_231) @p_229) :rule implies_neg1)
(anchor :step t93)
(assume t93.a0 @p_118)
(assume t93.a1 @p_2)
(step t93.t0 (cl (= (! (= @p_77 true) :named @p_232) @p_77)) :rule equiv_simplify)
(step t93.t1 (cl (not @p_232) @p_77) :rule equiv1 :premises (t93.t0))
(step t93.t2 (cl (= c_0 c8)) :rule symm :premises (t93.a1))
(step t93.t3 (cl (= @p_77 @p_118)) :rule cong :premises (t93.t2))
(step t93.t4 (cl (= @p_233 @p_118)) :rule equiv_simplify)
(step t93.t5 (cl @p_233 @p_228) :rule equiv2 :premises (t93.t4))
(step t93.t6 (cl @p_233) :rule resolution :premises (t93.t5 t88))
(step t93.t7 (cl @p_232) :rule trans :premises (t93.t3 t93.t6))
(step t93.t8 (cl @p_77) :rule resolution :premises (t93.t1 t93.t7))
(step t93 (cl @p_228 @p_150 @p_77) :rule subproof :discharge (t93.a0 t93.a1))
(step t94 (cl (! (not @p_229) :named @p_230) @p_118) :rule and_pos)
(step t95 (cl @p_230 @p_2) :rule and_pos)
(step t96 (cl @p_77 @p_230 @p_230) :rule resolution :premises (t93 t94 t95))
(step t97 (cl @p_230 @p_230 @p_77) :rule reordering :premises (t96))
(step t98 (cl @p_230 @p_77) :rule contraction :premises (t97))
(step t99 (cl @p_231 @p_77) :rule resolution :premises (t92 t98))
(step t100 (cl @p_231 @p_78) :rule implies_neg2)
(step t101 (cl @p_231 @p_231) :rule resolution :premises (t99 t100))
(step t102 (cl @p_231) :rule contraction :premises (t101))
(step t103 (cl @p_230 @p_77) :rule implies :premises (t102))
(step t104 (cl @p_228 @p_150 @p_77) :rule resolution :premises (t91 t103))
(step t105 (cl @p_77 @p_228 @p_150) :rule reordering :premises (t104))
(step t106 (cl (not (! (= (! (or (! (not @p_133) :named @p_155) (! (not (! (= c_1 @p_8) :named @p_209)) :named @p_212) (! (not (! (= c_1 @p_18) :named @p_182)) :named @p_194) (! (not @p_143) :named @p_162)) :named @p_223) (! (or @p_155 @p_212 @p_194 @p_142) :named @p_222)) :named @p_227)) (not @p_223) @p_222) :rule equiv_pos2)
(step t107 (cl (= @p_155 @p_155)) :rule refl)
(step t108 (cl (= @p_212 @p_212)) :rule refl)
(step t109 (cl (= @p_194 @p_194)) :rule refl)
(step t110 (cl (= @p_162 @p_142)) :rule all_simplify)
(step t111 (cl @p_227) :rule cong :premises (t107 t108 t109 t110))
(step t112 (cl (! (=> (! (and @p_133 @p_209 @p_182 @p_143) :named @p_224) false) :named @p_226) @p_224) :rule implies_neg1)
(anchor :step t113)
(assume t113.a0 @p_133)
(assume t113.a1 @p_209)
(assume t113.a2 @p_182)
(assume t113.a3 @p_143)
(step t113.t0 (cl (! (not (! (= (! (= false true) :named @p_171) false) :named @p_178)) :named @p_179) (! (not @p_171) :named @p_180) false) :rule equiv_pos2)
(step t113.t1 (cl @p_178) :rule all_simplify)
(step t113.t2 (cl (! (= (! (= @p_142 false) :named @p_176) @p_143) :named @p_177)) :rule equiv_simplify)
(step t113.t3 (cl @p_176 @p_162) :rule equiv2 :premises (t113.t2))
(step t113.t4 (cl (! (not @p_162) :named @p_164) @p_142) :rule not_not)
(step t113.t5 (cl @p_176 @p_142) :rule resolution :premises (t113.t3 t113.t4))
(step t113.t6 (cl @p_176) :rule resolution :premises (t113.t5 t113.a3))
(step t113.t7 (cl (! (= false @p_142) :named @p_175)) :rule symm :premises (t113.t6))
(step t113.t8 (cl @p_19) :rule symm :premises (t113.a2))
(step t113.t9 (cl @p_83) :rule refl)
(step t113.t10 (cl (= @p_142 @p_75)) :rule cong :premises (t113.t8 t113.t9))
(step t113.t11 (cl @p_9) :rule symm :premises (t113.a1))
(step t113.t12 (cl @p_209) :rule symm :premises (t113.t11))
(step t113.t13 (cl (= @p_75 @p_133)) :rule cong :premises (t113.t12 t113.t9))
(step t113.t14 (cl (! (= (! (= @p_133 true) :named @p_172) @p_133) :named @p_173)) :rule equiv_simplify)
(step t113.t15 (cl @p_172 @p_155) :rule equiv2 :premises (t113.t14))
(step t113.t16 (cl @p_172) :rule resolution :premises (t113.t15 t113.a0))
(step t113.t17 (cl @p_171) :rule trans :premises (t113.t7 t113.t10 t113.t13 t113.t16))
(step t113.t18 (cl false) :rule resolution :premises (t113.t0 t113.t1 t113.t17))
(step t113 (cl @p_155 @p_212 @p_194 @p_162 false) :rule subproof :discharge (t113.a0 t113.a1 t113.a2 t113.a3))
(step t114 (cl (! (not @p_224) :named @p_225) @p_133) :rule and_pos)
(step t115 (cl @p_225 @p_209) :rule and_pos)
(step t116 (cl @p_225 @p_182) :rule and_pos)
(step t117 (cl @p_225 @p_143) :rule and_pos)
(step t118 (cl false @p_225 @p_225 @p_225 @p_225) :rule resolution :premises (t113 t114 t115 t116 t117))
(step t119 (cl @p_225 @p_225 @p_225 @p_225 false) :rule reordering :premises (t118))
(step t120 (cl @p_225 false) :rule contraction :premises (t119))
(step t121 (cl @p_226 false) :rule resolution :premises (t112 t120))
(step t122 (cl @p_226 @p_170) :rule implies_neg2)
(step t123 (cl @p_226 @p_226) :rule resolution :premises (t121 t122))
(step t124 (cl @p_226) :rule contraction :premises (t123))
(step t125 (cl (= @p_226 @p_225)) :rule implies_simplify)
(step t126 (cl (not @p_226) @p_225) :rule equiv1 :premises (t125))
(step t127 (cl @p_225) :rule resolution :premises (t124 t126))
(step t128 (cl @p_155 @p_212 @p_194 @p_162) :rule not_and :premises (t127))
(step t129 (cl @p_223 (! (not @p_155) :named @p_166)) :rule or_neg)
(step t130 (cl @p_223 (not @p_212)) :rule or_neg)
(step t131 (cl @p_223 (! (not @p_194) :named @p_198)) :rule or_neg)
(step t132 (cl @p_223 @p_164) :rule or_neg)
(step t133 (cl @p_223 @p_223 @p_223 @p_223) :rule resolution :premises (t128 t129 t130 t131 t132))
(step t134 (cl @p_223) :rule contraction :premises (t133))
(step t135 (cl @p_222) :rule resolution :premises (t106 t111 t134))
(step t136 (cl @p_155 @p_212 @p_194 @p_142) :rule or :premises (t135))
(step t137 (cl @p_142 @p_155 @p_194 @p_212) :rule reordering :premises (t136))
(step t138 (cl (not (! (= @p_21 (! (or (! (= c_0 @p_18) :named @p_159) @p_182) :named @p_183)) :named @p_184)) (not @p_21) @p_183) :rule equiv_pos2)
(step t139 (cl (= @p_20 @p_159)) :rule all_simplify)
(step t140 (cl (= @p_19 @p_182)) :rule all_simplify)
(step t141 (cl @p_184) :rule cong :premises (t139 t140))
(step t142 (cl @p_21) :rule and :premises (a0))
(step t143 (cl @p_183) :rule resolution :premises (t138 t141 t142))
(step t144 (cl @p_159 @p_182) :rule or :premises (t143))
(step t145 (cl (not (! (= (! (or (! (not @p_147) :named @p_156) (! (not (! (= c_0 @p_26) :named @p_185)) :named @p_186) (! (not @p_159) :named @p_160) @p_162) :named @p_217) (! (or @p_156 @p_186 @p_160 @p_142) :named @p_216)) :named @p_221)) (not @p_217) @p_216) :rule equiv_pos2)
(step t146 (cl (= @p_156 @p_156)) :rule refl)
(step t147 (cl (= @p_186 @p_186)) :rule refl)
(step t148 (cl (= @p_160 @p_160)) :rule refl)
(step t149 (cl @p_221) :rule cong :premises (t146 t147 t148 t110))
(step t150 (cl (! (=> (! (and @p_147 @p_185 @p_159 @p_143) :named @p_218) false) :named @p_220) @p_218) :rule implies_neg1)
(anchor :step t151)
(assume t151.a0 @p_147)
(assume t151.a1 @p_185)
(assume t151.a2 @p_159)
(assume t151.a3 @p_143)
(step t151.t0 (cl @p_179 @p_180 false) :rule equiv_pos2)
(step t151.t1 (cl @p_178) :rule all_simplify)
(step t151.t2 (cl @p_177) :rule equiv_simplify)
(step t151.t3 (cl @p_176 @p_162) :rule equiv2 :premises (t151.t2))
(step t151.t4 (cl @p_164 @p_142) :rule not_not)
(step t151.t5 (cl @p_176 @p_142) :rule resolution :premises (t151.t3 t151.t4))
(step t151.t6 (cl @p_176) :rule resolution :premises (t151.t5 t151.a3))
(step t151.t7 (cl @p_175) :rule symm :premises (t151.t6))
(step t151.t8 (cl @p_20) :rule symm :premises (t151.a2))
(step t151.t9 (cl @p_83) :rule refl)
(step t151.t10 (cl (! (= @p_142 @p_80) :named @p_174)) :rule cong :premises (t151.t8 t151.t9))
(step t151.t11 (cl @p_28) :rule symm :premises (t151.a1))
(step t151.t12 (cl @p_185) :rule symm :premises (t151.t11))
(step t151.t13 (cl (= @p_80 @p_147)) :rule cong :premises (t151.t12 t151.t9))
(step t151.t14 (cl (! (= (! (= @p_147 true) :named @p_203) @p_147) :named @p_204)) :rule equiv_simplify)
(step t151.t15 (cl @p_203 @p_156) :rule equiv2 :premises (t151.t14))
(step t151.t16 (cl @p_203) :rule resolution :premises (t151.t15 t151.a0))
(step t151.t17 (cl @p_171) :rule trans :premises (t151.t7 t151.t10 t151.t13 t151.t16))
(step t151.t18 (cl false) :rule resolution :premises (t151.t0 t151.t1 t151.t17))
(step t151 (cl @p_156 @p_186 @p_160 @p_162 false) :rule subproof :discharge (t151.a0 t151.a1 t151.a2 t151.a3))
(step t152 (cl (! (not @p_218) :named @p_219) @p_147) :rule and_pos)
(step t153 (cl @p_219 @p_185) :rule and_pos)
(step t154 (cl @p_219 @p_159) :rule and_pos)
(step t155 (cl @p_219 @p_143) :rule and_pos)
(step t156 (cl false @p_219 @p_219 @p_219 @p_219) :rule resolution :premises (t151 t152 t153 t154 t155))
(step t157 (cl @p_219 @p_219 @p_219 @p_219 false) :rule reordering :premises (t156))
(step t158 (cl @p_219 false) :rule contraction :premises (t157))
(step t159 (cl @p_220 false) :rule resolution :premises (t150 t158))
(step t160 (cl @p_220 @p_170) :rule implies_neg2)
(step t161 (cl @p_220 @p_220) :rule resolution :premises (t159 t160))
(step t162 (cl @p_220) :rule contraction :premises (t161))
(step t163 (cl (= @p_220 @p_219)) :rule implies_simplify)
(step t164 (cl (not @p_220) @p_219) :rule equiv1 :premises (t163))
(step t165 (cl @p_219) :rule resolution :premises (t162 t164))
(step t166 (cl @p_156 @p_186 @p_160 @p_162) :rule not_and :premises (t165))
(step t167 (cl @p_217 (! (not @p_156) :named @p_197)) :rule or_neg)
(step t168 (cl @p_217 (not @p_186)) :rule or_neg)
(step t169 (cl @p_217 (! (not @p_160) :named @p_165)) :rule or_neg)
(step t170 (cl @p_217 @p_164) :rule or_neg)
(step t171 (cl @p_217 @p_217 @p_217 @p_217) :rule resolution :premises (t166 t167 t168 t169 t170))
(step t172 (cl @p_217) :rule contraction :premises (t171))
(step t173 (cl @p_216) :rule resolution :premises (t145 t149 t172))
(step t174 (cl @p_156 @p_186 @p_160 @p_142) :rule or :premises (t173))
(step t175 (cl @p_142 @p_156 @p_186 @p_160) :rule reordering :premises (t174))
(step t176 (cl (not (! (= @p_29 (! (or @p_185 (! (= c_1 @p_26) :named @p_190)) :named @p_191)) :named @p_192)) (not @p_29) @p_191) :rule equiv_pos2)
(step t177 (cl (= @p_28 @p_185)) :rule all_simplify)
(step t178 (cl (= @p_27 @p_190)) :rule all_simplify)
(step t179 (cl @p_192) :rule cong :premises (t177 t178))
(step t180 (cl @p_29) :rule and :premises (a0))
(step t181 (cl @p_191) :rule resolution :premises (t176 t179 t180))
(step t182 (cl @p_185 @p_190) :rule or :premises (t181))
(step t183 (cl (! (and @p_209 @p_190) :named @p_213) @p_212 (! (not @p_190) :named @p_193)) :rule and_neg)
(step t184 (cl (! (=> @p_213 (! (= @p_26 @p_8) :named @p_151)) :named @p_215) @p_213) :rule implies_neg1)
(anchor :step t185)
(assume t185.a0 @p_209)
(assume t185.a1 @p_190)
(step t185.t0 (cl @p_27) :rule symm :premises (t185.a1))
(step t185.t1 (cl @p_9) :rule symm :premises (t185.a0))
(step t185.t2 (cl @p_209) :rule symm :premises (t185.t1))
(step t185.t3 (cl @p_151) :rule trans :premises (t185.t0 t185.t2))
(step t185 (cl @p_212 @p_193 @p_151) :rule subproof :discharge (t185.a0 t185.a1))
(step t186 (cl (! (not @p_213) :named @p_214) @p_209) :rule and_pos)
(step t187 (cl @p_214 @p_190) :rule and_pos)
(step t188 (cl @p_151 @p_214 @p_214) :rule resolution :premises (t185 t186 t187))
(step t189 (cl @p_214 @p_214 @p_151) :rule reordering :premises (t188))
(step t190 (cl @p_214 @p_151) :rule contraction :premises (t189))
(step t191 (cl @p_215 @p_151) :rule resolution :premises (t184 t190))
(step t192 (cl @p_215 (! (not @p_151) :named @p_152)) :rule implies_neg2)
(step t193 (cl @p_215 @p_215) :rule resolution :premises (t191 t192))
(step t194 (cl @p_215) :rule contraction :premises (t193))
(step t195 (cl @p_214 @p_151) :rule implies :premises (t194))
(step t196 (cl @p_212 @p_193 @p_151) :rule resolution :premises (t183 t195))
(step t197 (cl @p_151 @p_193 @p_212) :rule reordering :premises (t196))
(step t198 (cl @p_142 @p_155 @p_212 @p_142 @p_156 @p_151 @p_212) :rule resolution :premises (t137 t144 t175 t182 t197))
(step t199 (cl @p_142 @p_155 @p_212 @p_156 @p_151) :rule contraction :premises (t198))
(step t200 (cl (not (! (= @p_11 (! (or (! (= c_0 @p_8) :named @p_157) @p_209) :named @p_210)) :named @p_211)) (not @p_11) @p_210) :rule equiv_pos2)
(step t201 (cl (= @p_10 @p_157)) :rule all_simplify)
(step t202 (cl (= @p_9 @p_209)) :rule all_simplify)
(step t203 (cl @p_211) :rule cong :premises (t201 t202))
(step t204 (cl @p_11) :rule and :premises (a0))
(step t205 (cl @p_210) :rule resolution :premises (t200 t203 t204))
(step t206 (cl @p_157 @p_209) :rule or :premises (t205))
(step t207 (cl (not (! (= (! (or @p_162 @p_194 @p_193 @p_156) :named @p_196) (! (or @p_142 @p_194 @p_193 @p_156) :named @p_195)) :named @p_208)) (not @p_196) @p_195) :rule equiv_pos2)
(step t208 (cl (= @p_193 @p_193)) :rule refl)
(step t209 (cl @p_208) :rule cong :premises (t110 t109 t208 t146))
(step t210 (cl (! (=> (! (and @p_143 @p_182 @p_190 @p_147) :named @p_199) false) :named @p_201) @p_199) :rule implies_neg1)
(anchor :step t211)
(assume t211.a0 @p_143)
(assume t211.a1 @p_182)
(assume t211.a2 @p_190)
(assume t211.a3 @p_147)
(step t211.t0 (cl @p_206 @p_207 false) :rule equiv_pos2)
(step t211.t1 (cl @p_205) :rule all_simplify)
(step t211.t2 (cl @p_204) :rule equiv_simplify)
(step t211.t3 (cl @p_203 @p_156) :rule equiv2 :premises (t211.t2))
(step t211.t4 (cl @p_203) :rule resolution :premises (t211.t3 t211.a3))
(step t211.t5 (cl (= true @p_147)) :rule symm :premises (t211.t4))
(step t211.t6 (cl @p_27) :rule symm :premises (t211.a2))
(step t211.t7 (cl @p_83) :rule refl)
(step t211.t8 (cl (= @p_147 @p_75)) :rule cong :premises (t211.t6 t211.t7))
(step t211.t9 (cl @p_19) :rule symm :premises (t211.a1))
(step t211.t10 (cl @p_182) :rule symm :premises (t211.t9))
(step t211.t11 (cl (= @p_75 @p_142)) :rule cong :premises (t211.t10 t211.t7))
(step t211.t12 (cl @p_177) :rule equiv_simplify)
(step t211.t13 (cl @p_176 @p_162) :rule equiv2 :premises (t211.t12))
(step t211.t14 (cl @p_164 @p_142) :rule not_not)
(step t211.t15 (cl @p_176 @p_142) :rule resolution :premises (t211.t13 t211.t14))
(step t211.t16 (cl @p_176) :rule resolution :premises (t211.t15 t211.a0))
(step t211.t17 (cl @p_202) :rule trans :premises (t211.t5 t211.t8 t211.t11 t211.t16))
(step t211.t18 (cl false) :rule resolution :premises (t211.t0 t211.t1 t211.t17))
(step t211 (cl @p_162 @p_194 @p_193 @p_156 false) :rule subproof :discharge (t211.a0 t211.a1 t211.a2 t211.a3))
(step t212 (cl (! (not @p_199) :named @p_200) @p_143) :rule and_pos)
(step t213 (cl @p_200 @p_182) :rule and_pos)
(step t214 (cl @p_200 @p_190) :rule and_pos)
(step t215 (cl @p_200 @p_147) :rule and_pos)
(step t216 (cl false @p_200 @p_200 @p_200 @p_200) :rule resolution :premises (t211 t212 t213 t214 t215))
(step t217 (cl @p_200 @p_200 @p_200 @p_200 false) :rule reordering :premises (t216))
(step t218 (cl @p_200 false) :rule contraction :premises (t217))
(step t219 (cl @p_201 false) :rule resolution :premises (t210 t218))
(step t220 (cl @p_201 @p_170) :rule implies_neg2)
(step t221 (cl @p_201 @p_201) :rule resolution :premises (t219 t220))
(step t222 (cl @p_201) :rule contraction :premises (t221))
(step t223 (cl (= @p_201 @p_200)) :rule implies_simplify)
(step t224 (cl (not @p_201) @p_200) :rule equiv1 :premises (t223))
(step t225 (cl @p_200) :rule resolution :premises (t222 t224))
(step t226 (cl @p_162 @p_194 @p_193 @p_156) :rule not_and :premises (t225))
(step t227 (cl @p_196 @p_164) :rule or_neg)
(step t228 (cl @p_196 @p_198) :rule or_neg)
(step t229 (cl @p_196 (not @p_193)) :rule or_neg)
(step t230 (cl @p_196 @p_197) :rule or_neg)
(step t231 (cl @p_196 @p_196 @p_196 @p_196) :rule resolution :premises (t226 t227 t228 t229 t230))
(step t232 (cl @p_196) :rule contraction :premises (t231))
(step t233 (cl @p_195) :rule resolution :premises (t207 t209 t232))
(step t234 (cl @p_142 @p_194 @p_193 @p_156) :rule or :premises (t233))
(step t235 (cl @p_142 @p_156 @p_193 @p_194) :rule reordering :premises (t234))
(step t236 (cl @p_185 @p_190) :rule or :premises (t181))
(step t237 (cl (! (and @p_157 @p_185) :named @p_187) (! (not @p_157) :named @p_158) @p_186) :rule and_neg)
(step t238 (cl (! (=> @p_187 @p_151) :named @p_189) @p_187) :rule implies_neg1)
(anchor :step t239)
(assume t239.a0 @p_157)
(assume t239.a1 @p_185)
(step t239.t0 (cl @p_28) :rule symm :premises (t239.a1))
(step t239.t1 (cl @p_10) :rule symm :premises (t239.a0))
(step t239.t2 (cl @p_157) :rule symm :premises (t239.t1))
(step t239.t3 (cl @p_151) :rule trans :premises (t239.t0 t239.t2))
(step t239 (cl @p_158 @p_186 @p_151) :rule subproof :discharge (t239.a0 t239.a1))
(step t240 (cl (! (not @p_187) :named @p_188) @p_157) :rule and_pos)
(step t241 (cl @p_188 @p_185) :rule and_pos)
(step t242 (cl @p_151 @p_188 @p_188) :rule resolution :premises (t239 t240 t241))
(step t243 (cl @p_188 @p_188 @p_151) :rule reordering :premises (t242))
(step t244 (cl @p_188 @p_151) :rule contraction :premises (t243))
(step t245 (cl @p_189 @p_151) :rule resolution :premises (t238 t244))
(step t246 (cl @p_189 @p_152) :rule implies_neg2)
(step t247 (cl @p_189 @p_189) :rule resolution :premises (t245 t246))
(step t248 (cl @p_189) :rule contraction :premises (t247))
(step t249 (cl @p_188 @p_151) :rule implies :premises (t248))
(step t250 (cl @p_158 @p_186 @p_151) :rule resolution :premises (t237 t249))
(step t251 (cl @p_151 @p_186 @p_158) :rule reordering :premises (t250))
(step t252 (cl @p_159 @p_182) :rule or :premises (t143))
(step t253 (cl (not (! (= (! (or @p_155 @p_158 @p_160 @p_162) :named @p_163) (! (or @p_155 @p_158 @p_160 @p_142) :named @p_161)) :named @p_181)) (not @p_163) @p_161) :rule equiv_pos2)
(step t254 (cl (= @p_158 @p_158)) :rule refl)
(step t255 (cl @p_181) :rule cong :premises (t107 t254 t148 t110))
(step t256 (cl (! (=> (! (and @p_133 @p_157 @p_159 @p_143) :named @p_167) false) :named @p_169) @p_167) :rule implies_neg1)
(anchor :step t257)
(assume t257.a0 @p_133)
(assume t257.a1 @p_157)
(assume t257.a2 @p_159)
(assume t257.a3 @p_143)
(step t257.t0 (cl @p_179 @p_180 false) :rule equiv_pos2)
(step t257.t1 (cl @p_178) :rule all_simplify)
(step t257.t2 (cl @p_177) :rule equiv_simplify)
(step t257.t3 (cl @p_176 @p_162) :rule equiv2 :premises (t257.t2))
(step t257.t4 (cl @p_164 @p_142) :rule not_not)
(step t257.t5 (cl @p_176 @p_142) :rule resolution :premises (t257.t3 t257.t4))
(step t257.t6 (cl @p_176) :rule resolution :premises (t257.t5 t257.a3))
(step t257.t7 (cl @p_175) :rule symm :premises (t257.t6))
(step t257.t8 (cl @p_20) :rule symm :premises (t257.a2))
(step t257.t9 (cl @p_83) :rule refl)
(step t257.t10 (cl @p_174) :rule cong :premises (t257.t8 t257.t9))
(step t257.t11 (cl @p_10) :rule symm :premises (t257.a1))
(step t257.t12 (cl @p_157) :rule symm :premises (t257.t11))
(step t257.t13 (cl (= @p_80 @p_133)) :rule cong :premises (t257.t12 t257.t9))
(step t257.t14 (cl @p_173) :rule equiv_simplify)
(step t257.t15 (cl @p_172 @p_155) :rule equiv2 :premises (t257.t14))
(step t257.t16 (cl @p_172) :rule resolution :premises (t257.t15 t257.a0))
(step t257.t17 (cl @p_171) :rule trans :premises (t257.t7 t257.t10 t257.t13 t257.t16))
(step t257.t18 (cl false) :rule resolution :premises (t257.t0 t257.t1 t257.t17))
(step t257 (cl @p_155 @p_158 @p_160 @p_162 false) :rule subproof :discharge (t257.a0 t257.a1 t257.a2 t257.a3))
(step t258 (cl (! (not @p_167) :named @p_168) @p_133) :rule and_pos)
(step t259 (cl @p_168 @p_157) :rule and_pos)
(step t260 (cl @p_168 @p_159) :rule and_pos)
(step t261 (cl @p_168 @p_143) :rule and_pos)
(step t262 (cl false @p_168 @p_168 @p_168 @p_168) :rule resolution :premises (t257 t258 t259 t260 t261))
(step t263 (cl @p_168 @p_168 @p_168 @p_168 false) :rule reordering :premises (t262))
(step t264 (cl @p_168 false) :rule contraction :premises (t263))
(step t265 (cl @p_169 false) :rule resolution :premises (t256 t264))
(step t266 (cl @p_169 @p_170) :rule implies_neg2)
(step t267 (cl @p_169 @p_169) :rule resolution :premises (t265 t266))
(step t268 (cl @p_169) :rule contraction :premises (t267))
(step t269 (cl (= @p_169 @p_168)) :rule implies_simplify)
(step t270 (cl (not @p_169) @p_168) :rule equiv1 :premises (t269))
(step t271 (cl @p_168) :rule resolution :premises (t268 t270))
(step t272 (cl @p_155 @p_158 @p_160 @p_162) :rule not_and :premises (t271))
(step t273 (cl @p_163 @p_166) :rule or_neg)
(step t274 (cl @p_163 (not @p_158)) :rule or_neg)
(step t275 (cl @p_163 @p_165) :rule or_neg)
(step t276 (cl @p_163 @p_164) :rule or_neg)
(step t277 (cl @p_163 @p_163 @p_163 @p_163) :rule resolution :premises (t272 t273 t274 t275 t276))
(step t278 (cl @p_163) :rule contraction :premises (t277))
(step t279 (cl @p_161) :rule resolution :premises (t253 t255 t278))
(step t280 (cl @p_155 @p_158 @p_160 @p_142) :rule or :premises (t279))
(step t281 (cl @p_142 @p_155 @p_160 @p_158) :rule reordering :premises (t280))
(step t282 (cl @p_142 @p_156 @p_151 @p_158 @p_142 @p_155 @p_158) :rule resolution :premises (t235 t236 t251 t252 t281))
(step t283 (cl @p_142 @p_156 @p_151 @p_158 @p_155) :rule contraction :premises (t282))
(step t284 (cl @p_142 @p_151 @p_156 @p_155 @p_158) :rule reordering :premises (t283))
(step t285 (cl @p_142 @p_155 @p_156 @p_151 @p_142 @p_151 @p_156 @p_155) :rule resolution :premises (t199 t206 t284))
(step t286 (cl @p_142 @p_155 @p_156 @p_151) :rule contraction :premises (t285))
(step t287 (cl @p_148) :rule and :premises (a0))
(step t288 (cl @p_78 @p_147) :rule or :premises (t287))
(step t289 (cl @p_144) :rule and :premises (a0))
(step t290 (cl @p_78 @p_143) :rule or :premises (t289))
(step t291 (cl (not (! (= @p_130 (! (or @p_78 @p_152) :named @p_153)) :named @p_154)) (not @p_130) @p_153) :rule equiv_pos2)
(step t292 (cl (= @p_78 @p_78)) :rule refl)
(step t293 (cl (= @p_128 @p_151)) :rule all_simplify)
(step t294 (cl (= @p_129 @p_152)) :rule cong :premises (t293))
(step t295 (cl @p_154) :rule cong :premises (t292 t294))
(step t296 (cl @p_130) :rule and :premises (a0))
(step t297 (cl @p_153) :rule resolution :premises (t291 t295 t296))
(step t298 (cl @p_78 @p_152) :rule or :premises (t297))
(step t299 (cl @p_134) :rule and :premises (a0))
(step t300 (cl @p_78 @p_133) :rule or :premises (t299))
(step t301 (cl @p_78 @p_78 @p_78 @p_78) :rule resolution :premises (t286 t288 t290 t298 t300))
(step t302 (cl @p_78) :rule contraction :premises (t301))
(step t303 (cl @p_150) :rule resolution :premises (t105 t302 t88))
(step t304 (cl @p_1) :rule resolution :premises (t90 t303))
(step t305 (cl @p_65) :rule resolution :premises (t87 t88 t304))
(step t306 (cl @p_266) :rule resolution :premises (t72 t305))
(step t307 (cl (not (! (= @p_25 (! (or @p_243 (! (= c_1 @p_22) :named @p_262)) :named @p_263)) :named @p_264)) (not @p_25) @p_263) :rule equiv_pos2)
(step t308 (cl (= @p_24 @p_243)) :rule all_simplify)
(step t309 (cl (= @p_23 @p_262)) :rule all_simplify)
(step t310 (cl @p_264) :rule cong :premises (t308 t309))
(step t311 (cl @p_25) :rule and :premises (a0))
(step t312 (cl @p_263) :rule resolution :premises (t307 t310 t311))
(step t313 (cl @p_243 @p_262) :rule or :premises (t312))
(step t314 (cl (not (! (= (! (or @p_246 @p_303 (! (not @p_262) :named @p_269) (! (not @p_145) :named @p_242)) :named @p_305) (! (or @p_139 @p_303 @p_269 @p_242) :named @p_304)) :named @p_311)) (not @p_305) @p_304) :rule equiv_pos2)
(step t315 (cl (= @p_269 @p_269)) :rule refl)
(step t316 (cl (= @p_242 @p_242)) :rule refl)
(step t317 (cl @p_311) :rule cong :premises (t1 t2 t315 t316))
(step t318 (cl (! (=> (! (and @p_140 @p_239 @p_262 @p_145) :named @p_307) false) :named @p_309) @p_307) :rule implies_neg1)
(anchor :step t319)
(assume t319.a0 @p_140)
(assume t319.a1 @p_239)
(assume t319.a2 @p_262)
(assume t319.a3 @p_145)
(step t319.t0 (cl @p_206 @p_207 false) :rule equiv_pos2)
(step t319.t1 (cl @p_205) :rule all_simplify)
(step t319.t2 (cl (! (= (! (= @p_145 true) :named @p_259) @p_145) :named @p_260)) :rule equiv_simplify)
(step t319.t3 (cl @p_259 @p_242) :rule equiv2 :premises (t319.t2))
(step t319.t4 (cl @p_259) :rule resolution :premises (t319.t3 t319.a3))
(step t319.t5 (cl (! (= true @p_145) :named @p_258)) :rule symm :premises (t319.t4))
(step t319.t6 (cl @p_23) :rule symm :premises (t319.a2))
(step t319.t7 (cl @p_69) :rule refl)
(step t319.t8 (cl (= @p_145 @p_63)) :rule cong :premises (t319.t6 t319.t7))
(step t319.t9 (cl @p_15) :rule symm :premises (t319.a1))
(step t319.t10 (cl @p_239) :rule symm :premises (t319.t9))
(step t319.t11 (cl @p_310) :rule cong :premises (t319.t10 t319.t7))
(step t319.t12 (cl @p_256) :rule equiv_simplify)
(step t319.t13 (cl @p_255 @p_246) :rule equiv2 :premises (t319.t12))
(step t319.t14 (cl @p_251 @p_139) :rule not_not)
(step t319.t15 (cl @p_255 @p_139) :rule resolution :premises (t319.t13 t319.t14))
(step t319.t16 (cl @p_255) :rule resolution :premises (t319.t15 t319.a0))
(step t319.t17 (cl @p_202) :rule trans :premises (t319.t5 t319.t8 t319.t11 t319.t16))
(step t319.t18 (cl false) :rule resolution :premises (t319.t0 t319.t1 t319.t17))
(step t319 (cl @p_246 @p_303 @p_269 @p_242 false) :rule subproof :discharge (t319.a0 t319.a1 t319.a2 t319.a3))
(step t320 (cl (! (not @p_307) :named @p_308) @p_140) :rule and_pos)
(step t321 (cl @p_308 @p_239) :rule and_pos)
(step t322 (cl @p_308 @p_262) :rule and_pos)
(step t323 (cl @p_308 @p_145) :rule and_pos)
(step t324 (cl false @p_308 @p_308 @p_308 @p_308) :rule resolution :premises (t319 t320 t321 t322 t323))
(step t325 (cl @p_308 @p_308 @p_308 @p_308 false) :rule reordering :premises (t324))
(step t326 (cl @p_308 false) :rule contraction :premises (t325))
(step t327 (cl @p_309 false) :rule resolution :premises (t318 t326))
(step t328 (cl @p_309 @p_170) :rule implies_neg2)
(step t329 (cl @p_309 @p_309) :rule resolution :premises (t327 t328))
(step t330 (cl @p_309) :rule contraction :premises (t329))
(step t331 (cl (= @p_309 @p_308)) :rule implies_simplify)
(step t332 (cl (not @p_309) @p_308) :rule equiv1 :premises (t331))
(step t333 (cl @p_308) :rule resolution :premises (t330 t332))
(step t334 (cl @p_246 @p_303 @p_269 @p_242) :rule not_and :premises (t333))
(step t335 (cl @p_305 @p_251) :rule or_neg)
(step t336 (cl @p_305 @p_306) :rule or_neg)
(step t337 (cl @p_305 (! (not @p_269) :named @p_276)) :rule or_neg)
(step t338 (cl @p_305 (! (not @p_242) :named @p_248)) :rule or_neg)
(step t339 (cl @p_305 @p_305 @p_305 @p_305) :rule resolution :premises (t334 t335 t336 t337 t338))
(step t340 (cl @p_305) :rule contraction :premises (t339))
(step t341 (cl @p_304) :rule resolution :premises (t314 t317 t340))
(step t342 (cl @p_139 @p_303 @p_269 @p_242) :rule or :premises (t341))
(step t343 (cl @p_139 @p_242 @p_269 @p_303) :rule reordering :premises (t342))
(step t344 (cl @p_141) :rule and :premises (a0))
(step t345 (cl @p_66 @p_140) :rule or :premises (t344))
(step t346 (cl @p_140) :rule resolution :premises (t345 t305))
(step t347 (cl @p_146) :rule and :premises (a0))
(step t348 (cl @p_66 @p_145) :rule or :premises (t347))
(step t349 (cl @p_145) :rule resolution :premises (t348 t305))
(step t350 (cl (not (! (= @p_17 (! (or (! (= c_0 @p_14) :named @p_240) @p_239) :named @p_301)) :named @p_302)) (not @p_17) @p_301) :rule equiv_pos2)
(step t351 (cl (= @p_16 @p_240)) :rule all_simplify)
(step t352 (cl (= @p_15 @p_239)) :rule all_simplify)
(step t353 (cl @p_302) :rule cong :premises (t351 t352))
(step t354 (cl @p_17) :rule and :premises (a0))
(step t355 (cl @p_301) :rule resolution :premises (t350 t353 t354))
(step t356 (cl @p_240 @p_239) :rule or :premises (t355))
(step t357 (cl (not (! (= (! (or @p_246 (! (not @p_240) :named @p_241) @p_289 @p_288) :named @p_291) (! (or @p_139 @p_241 @p_289 @p_288) :named @p_290)) :named @p_300)) (not @p_291) @p_290) :rule equiv_pos2)
(step t358 (cl (= @p_241 @p_241)) :rule refl)
(step t359 (cl @p_300) :rule cong :premises (t1 t358 t42 t4))
(step t360 (cl (! (=> (! (and @p_140 @p_240 @p_285 @p_131) :named @p_294) false) :named @p_296) @p_294) :rule implies_neg1)
(anchor :step t361)
(assume t361.a0 @p_140)
(assume t361.a1 @p_240)
(assume t361.a2 @p_285)
(assume t361.a3 @p_131)
(step t361.t0 (cl @p_206 @p_207 false) :rule equiv_pos2)
(step t361.t1 (cl @p_205) :rule all_simplify)
(step t361.t2 (cl @p_299) :rule equiv_simplify)
(step t361.t3 (cl @p_298 @p_288) :rule equiv2 :premises (t361.t2))
(step t361.t4 (cl @p_298) :rule resolution :premises (t361.t3 t361.a3))
(step t361.t5 (cl @p_297) :rule symm :premises (t361.t4))
(step t361.t6 (cl @p_6) :rule symm :premises (t361.a2))
(step t361.t7 (cl @p_69) :rule refl)
(step t361.t8 (cl (= @p_131 @p_71)) :rule cong :premises (t361.t6 t361.t7))
(step t361.t9 (cl @p_16) :rule symm :premises (t361.a1))
(step t361.t10 (cl @p_240) :rule symm :premises (t361.t9))
(step t361.t11 (cl (! (= @p_71 @p_139) :named @p_257)) :rule cong :premises (t361.t10 t361.t7))
(step t361.t12 (cl @p_256) :rule equiv_simplify)
(step t361.t13 (cl @p_255 @p_246) :rule equiv2 :premises (t361.t12))
(step t361.t14 (cl @p_251 @p_139) :rule not_not)
(step t361.t15 (cl @p_255 @p_139) :rule resolution :premises (t361.t13 t361.t14))
(step t361.t16 (cl @p_255) :rule resolution :premises (t361.t15 t361.a0))
(step t361.t17 (cl @p_202) :rule trans :premises (t361.t5 t361.t8 t361.t11 t361.t16))
(step t361.t18 (cl false) :rule resolution :premises (t361.t0 t361.t1 t361.t17))
(step t361 (cl @p_246 @p_241 @p_289 @p_288 false) :rule subproof :discharge (t361.a0 t361.a1 t361.a2 t361.a3))
(step t362 (cl (! (not @p_294) :named @p_295) @p_140) :rule and_pos)
(step t363 (cl @p_295 @p_240) :rule and_pos)
(step t364 (cl @p_295 @p_285) :rule and_pos)
(step t365 (cl @p_295 @p_131) :rule and_pos)
(step t366 (cl false @p_295 @p_295 @p_295 @p_295) :rule resolution :premises (t361 t362 t363 t364 t365))
(step t367 (cl @p_295 @p_295 @p_295 @p_295 false) :rule reordering :premises (t366))
(step t368 (cl @p_295 false) :rule contraction :premises (t367))
(step t369 (cl @p_296 false) :rule resolution :premises (t360 t368))
(step t370 (cl @p_296 @p_170) :rule implies_neg2)
(step t371 (cl @p_296 @p_296) :rule resolution :premises (t369 t370))
(step t372 (cl @p_296) :rule contraction :premises (t371))
(step t373 (cl (= @p_296 @p_295)) :rule implies_simplify)
(step t374 (cl (not @p_296) @p_295) :rule equiv1 :premises (t373))
(step t375 (cl @p_295) :rule resolution :premises (t372 t374))
(step t376 (cl @p_246 @p_241 @p_289 @p_288) :rule not_and :premises (t375))
(step t377 (cl @p_291 @p_251) :rule or_neg)
(step t378 (cl @p_291 (! (not @p_241) :named @p_250)) :rule or_neg)
(step t379 (cl @p_291 @p_293) :rule or_neg)
(step t380 (cl @p_291 @p_292) :rule or_neg)
(step t381 (cl @p_291 @p_291 @p_291 @p_291) :rule resolution :premises (t376 t377 t378 t379 t380))
(step t382 (cl @p_291) :rule contraction :premises (t381))
(step t383 (cl @p_290) :rule resolution :premises (t357 t359 t382))
(step t384 (cl @p_139 @p_241 @p_289 @p_288) :rule or :premises (t383))
(step t385 (cl @p_139 @p_288 @p_241 @p_289) :rule reordering :premises (t384))
(step t386 (cl @p_132) :rule and :premises (a0))
(step t387 (cl @p_66 @p_131) :rule or :premises (t386))
(step t388 (cl @p_131) :rule resolution :premises (t387 t305))
(step t389 (cl @p_285 @p_270) :rule or :premises (t37))
(step t390 (cl (not (! (= (! (or @p_273 @p_269 @p_271) :named @p_274) (! (or @p_265 @p_269 @p_271) :named @p_272)) :named @p_284)) (not @p_274) @p_272) :rule equiv_pos2)
(step t391 (cl @p_284) :rule cong :premises (t40 t315 t3))
(step t392 (cl (! (and @p_266 @p_262) :named @p_278) @p_273 @p_269) :rule and_neg)
(step t393 (cl (! (=> @p_278 @p_271) :named @p_280) @p_278) :rule implies_neg1)
(anchor :step t394)
(assume t394.a0 @p_266)
(assume t394.a1 @p_262)
(step t394.t0 (cl (= (! (= @p_270 false) :named @p_281) @p_271)) :rule equiv_simplify)
(step t394.t1 (cl (not @p_281) @p_271) :rule equiv1 :premises (t394.t0))
(step t394.t2 (cl @p_23) :rule symm :premises (t394.a1))
(step t394.t3 (cl @p_262) :rule symm :premises (t394.t2))
(step t394.t4 (cl (= @p_4 @p_4)) :rule refl)
(step t394.t5 (cl (= @p_270 @p_265)) :rule cong :premises (t394.t3 t394.t4))
(step t394.t6 (cl @p_283) :rule equiv_simplify)
(step t394.t7 (cl @p_282 @p_273) :rule equiv2 :premises (t394.t6))
(step t394.t8 (cl @p_277 @p_265) :rule not_not)
(step t394.t9 (cl @p_282 @p_265) :rule resolution :premises (t394.t7 t394.t8))
(step t394.t10 (cl @p_282) :rule resolution :premises (t394.t9 t394.a0))
(step t394.t11 (cl @p_281) :rule trans :premises (t394.t5 t394.t10))
(step t394.t12 (cl @p_271) :rule resolution :premises (t394.t1 t394.t11))
(step t394 (cl @p_273 @p_269 @p_271) :rule subproof :discharge (t394.a0 t394.a1))
(step t395 (cl (! (not @p_278) :named @p_279) @p_266) :rule and_pos)
(step t396 (cl @p_279 @p_262) :rule and_pos)
(step t397 (cl @p_271 @p_279 @p_279) :rule resolution :premises (t394 t395 t396))
(step t398 (cl @p_279 @p_279 @p_271) :rule reordering :premises (t397))
(step t399 (cl @p_279 @p_271) :rule contraction :premises (t398))
(step t400 (cl @p_280 @p_271) :rule resolution :premises (t393 t399))
(step t401 (cl @p_280 @p_275) :rule implies_neg2)
(step t402 (cl @p_280 @p_280) :rule resolution :premises (t400 t401))
(step t403 (cl @p_280) :rule contraction :premises (t402))
(step t404 (cl @p_279 @p_271) :rule implies :premises (t403))
(step t405 (cl @p_273 @p_269 @p_271) :rule resolution :premises (t392 t404))
(step t406 (cl @p_274 @p_277) :rule or_neg)
(step t407 (cl @p_274 @p_276) :rule or_neg)
(step t408 (cl @p_274 @p_275) :rule or_neg)
(step t409 (cl @p_274 @p_274 @p_274) :rule resolution :premises (t405 t406 t407 t408))
(step t410 (cl @p_274) :rule contraction :premises (t409))
(step t411 (cl @p_272) :rule resolution :premises (t390 t391 t410))
(step t412 (cl @p_265 @p_269 @p_271) :rule or :premises (t411))
(step t413 (cl @p_243 @p_262) :rule or :premises (t312))
(step t414 (cl (not (! (= (! (or @p_246 @p_241 @p_244 @p_242) :named @p_247) (! (or @p_139 @p_241 @p_244 @p_242) :named @p_245)) :named @p_261)) (not @p_247) @p_245) :rule equiv_pos2)
(step t415 (cl @p_261) :rule cong :premises (t1 t358 t41 t316))
(step t416 (cl (! (=> (! (and @p_140 @p_240 @p_243 @p_145) :named @p_252) false) :named @p_254) @p_252) :rule implies_neg1)
(anchor :step t417)
(assume t417.a0 @p_140)
(assume t417.a1 @p_240)
(assume t417.a2 @p_243)
(assume t417.a3 @p_145)
(step t417.t0 (cl @p_206 @p_207 false) :rule equiv_pos2)
(step t417.t1 (cl @p_205) :rule all_simplify)
(step t417.t2 (cl @p_260) :rule equiv_simplify)
(step t417.t3 (cl @p_259 @p_242) :rule equiv2 :premises (t417.t2))
(step t417.t4 (cl @p_259) :rule resolution :premises (t417.t3 t417.a3))
(step t417.t5 (cl @p_258) :rule symm :premises (t417.t4))
(step t417.t6 (cl @p_24) :rule symm :premises (t417.a2))
(step t417.t7 (cl @p_69) :rule refl)
(step t417.t8 (cl (= @p_145 @p_71)) :rule cong :premises (t417.t6 t417.t7))
(step t417.t9 (cl @p_16) :rule symm :premises (t417.a1))
(step t417.t10 (cl @p_240) :rule symm :premises (t417.t9))
(step t417.t11 (cl @p_257) :rule cong :premises (t417.t10 t417.t7))
(step t417.t12 (cl @p_256) :rule equiv_simplify)
(step t417.t13 (cl @p_255 @p_246) :rule equiv2 :premises (t417.t12))
(step t417.t14 (cl @p_251 @p_139) :rule not_not)
(step t417.t15 (cl @p_255 @p_139) :rule resolution :premises (t417.t13 t417.t14))
(step t417.t16 (cl @p_255) :rule resolution :premises (t417.t15 t417.a0))
(step t417.t17 (cl @p_202) :rule trans :premises (t417.t5 t417.t8 t417.t11 t417.t16))
(step t417.t18 (cl false) :rule resolution :premises (t417.t0 t417.t1 t417.t17))
(step t417 (cl @p_246 @p_241 @p_244 @p_242 false) :rule subproof :discharge (t417.a0 t417.a1 t417.a2 t417.a3))
(step t418 (cl (! (not @p_252) :named @p_253) @p_140) :rule and_pos)
(step t419 (cl @p_253 @p_240) :rule and_pos)
(step t420 (cl @p_253 @p_243) :rule and_pos)
(step t421 (cl @p_253 @p_145) :rule and_pos)
(step t422 (cl false @p_253 @p_253 @p_253 @p_253) :rule resolution :premises (t417 t418 t419 t420 t421))
(step t423 (cl @p_253 @p_253 @p_253 @p_253 false) :rule reordering :premises (t422))
(step t424 (cl @p_253 false) :rule contraction :premises (t423))
(step t425 (cl @p_254 false) :rule resolution :premises (t416 t424))
(step t426 (cl @p_254 @p_170) :rule implies_neg2)
(step t427 (cl @p_254 @p_254) :rule resolution :premises (t425 t426))
(step t428 (cl @p_254) :rule contraction :premises (t427))
(step t429 (cl (= @p_254 @p_253)) :rule implies_simplify)
(step t430 (cl (not @p_254) @p_253) :rule equiv1 :premises (t429))
(step t431 (cl @p_253) :rule resolution :premises (t428 t430))
(step t432 (cl @p_246 @p_241 @p_244 @p_242) :rule not_and :premises (t431))
(step t433 (cl @p_247 @p_251) :rule or_neg)
(step t434 (cl @p_247 @p_250) :rule or_neg)
(step t435 (cl @p_247 @p_249) :rule or_neg)
(step t436 (cl @p_247 @p_248) :rule or_neg)
(step t437 (cl @p_247 @p_247 @p_247 @p_247) :rule resolution :premises (t432 t433 t434 t435 t436))
(step t438 (cl @p_247) :rule contraction :premises (t437))
(step t439 (cl @p_245) :rule resolution :premises (t414 t415 t438))
(step t440 (cl @p_139 @p_241 @p_244 @p_242) :rule or :premises (t439))
(step t441 (cl @p_139 @p_242 @p_244 @p_241) :rule reordering :premises (t440))
(step t442 (cl @p_241 @p_241) :rule resolution :premises (t385 t388 t346 t389 t412 t306 t413 t441 t346 t349))
(step t443 (cl @p_241) :rule contraction :premises (t442))
(step t444 (cl @p_239) :rule resolution :premises (t356 t443))
(step t445 (cl @p_269) :rule resolution :premises (t343 t346 t349 t444))
(step t446 (cl @p_243) :rule resolution :premises (t313 t445))
(step t447 (cl @p_289) :rule resolution :premises (t64 t306 t446))
(step t448 (cl @p_270) :rule resolution :premises (t38 t447))
(step t449 (cl) :rule resolution :premises (t31 t448 t444 t388 t346))
