unsat
(assume a0 (! (and (! (or (and (= x0 y0) (= y0 x1)) (and (= x0 z0) (= z0 x1))) :named @p_77) (! (or (and (= x1 y1) (= y1 x2)) (and (= x1 z1) (= z1 x2))) :named @p_76) (! (or (and (= x2 y2) (= y2 x3)) (and (= x2 z2) (= z2 x3))) :named @p_75) (! (or (and (= x3 y3) (= y3 x4)) (and (= x3 z3) (= z3 x4))) :named @p_74) (! (or (! (and (= x4 y4) (= y4 x5)) :named @p_72) (! (and (= x4 z4) (= z4 x5)) :named @p_71)) :named @p_73) (! (or (and (= x5 y5) (= y5 x6)) (and (= x5 z5) (= z5 x6))) :named @p_70) (! (or (and (= x6 y6) (= y6 x7)) (and (= x6 z6) (= z6 x7))) :named @p_69) (! (or (and (= x7 y7) (= y7 x8)) (and (= x7 z7) (= z7 x8))) :named @p_68) (! (or (and (= x8 y8) (= y8 x9)) (and (= x8 z8) (= z8 x9))) :named @p_67) (! (or (and (= x9 y9) (= y9 x10)) (and (= x9 z9) (= z9 x10))) :named @p_66) (! (or (and (= x10 y10) (= y10 x11)) (and (= x10 z10) (= z10 x11))) :named @p_65) (! (or (and (= x11 y11) (= y11 x12)) (and (= x11 z11) (= z11 x12))) :named @p_64) (! (or (and (= x12 y12) (= y12 x13)) (and (= x12 z12) (= z12 x13))) :named @p_63) (! (or (and (= x13 y13) (= y13 x14)) (and (= x13 z13) (= z13 x14))) :named @p_62) (! (or (and (= x14 y14) (= y14 x15)) (and (= x14 z14) (= z14 x15))) :named @p_61) (! (or (and (= x15 y15) (= y15 x16)) (and (= x15 z15) (= z15 x16))) :named @p_60) (! (or (and (= x16 y16) (= y16 x17)) (and (= x16 z16) (= z16 x17))) :named @p_59) (! (or (and (= x17 y17) (= y17 x18)) (and (= x17 z17) (= z17 x18))) :named @p_58) (! (or (and (= x18 y18) (= y18 x19)) (and (= x18 z18) (= z18 x19))) :named @p_57) (! (or (and (= x19 y19) (= y19 x20)) (and (= x19 z19) (= z19 x20))) :named @p_56) (! (or (and (= x20 y20) (= y20 x21)) (and (= x20 z20) (= z20 x21))) :named @p_55) (! (or (and (= x21 y21) (= y21 x22)) (and (= x21 z21) (= z21 x22))) :named @p_54) (! (or (and (= x22 y22) (= y22 x23)) (and (= x22 z22) (= z22 x23))) :named @p_53) (! (or (and (= x23 y23) (= y23 x24)) (and (= x23 z23) (= z23 x24))) :named @p_52) (! (or (and (= x24 y24) (= y24 x25)) (and (= x24 z24) (= z24 x25))) :named @p_51) (! (or (and (= x25 y25) (= y25 x26)) (and (= x25 z25) (= z25 x26))) :named @p_50) (! (or (and (= x26 y26) (= y26 x27)) (and (= x26 z26) (= z26 x27))) :named @p_49) (! (or (and (= x27 y27) (= y27 x28)) (and (= x27 z27) (= z27 x28))) :named @p_48) (! (or (and (= x28 y28) (= y28 x29)) (and (= x28 z28) (= z28 x29))) :named @p_47) (! (or (and (= x29 y29) (= y29 x30)) (and (= x29 z29) (= z29 x30))) :named @p_46) (! (or (and (= x30 y30) (= y30 x31)) (and (= x30 z30) (= z30 x31))) :named @p_45) (! (or (and (= x31 y31) (= y31 x32)) (and (= x31 z31) (= z31 x32))) :named @p_44) (! (or (and (= x32 y32) (= y32 x33)) (and (= x32 z32) (= z32 x33))) :named @p_43) (! (or (and (= x33 y33) (= y33 x34)) (and (= x33 z33) (= z33 x34))) :named @p_42) (! (or (and (= x34 y34) (= y34 x35)) (and (= x34 z34) (= z34 x35))) :named @p_41) (! (or (and (= x35 y35) (= y35 x36)) (and (= x35 z35) (= z35 x36))) :named @p_40) (! (or (and (= x36 y36) (= y36 x37)) (and (= x36 z36) (= z36 x37))) :named @p_39) (! (or (and (= x37 y37) (= y37 x38)) (and (= x37 z37) (= z37 x38))) :named @p_38) (! (or (and (= x38 y38) (= y38 x39)) (and (= x38 z38) (= z38 x39))) :named @p_37) (! (or (and (= x39 y39) (= y39 x40)) (and (= x39 z39) (= z39 x40))) :named @p_36) (! (or (and (= x40 y40) (= y40 x41)) (and (= x40 z40) (= z40 x41))) :named @p_35) (! (or (and (= x41 y41) (= y41 x42)) (and (= x41 z41) (= z41 x42))) :named @p_34) (! (or (and (= x42 y42) (= y42 x43)) (and (= x42 z42) (= z42 x43))) :named @p_33) (! (or (and (= x43 y43) (= y43 x44)) (and (= x43 z43) (= z43 x44))) :named @p_32) (! (or (and (= x44 y44) (= y44 x45)) (and (= x44 z44) (= z44 x45))) :named @p_31) (! (or (and (= x45 y45) (= y45 x46)) (and (= x45 z45) (= z45 x46))) :named @p_30) (! (or (and (= x46 y46) (= y46 x47)) (and (= x46 z46) (= z46 x47))) :named @p_29) (! (or (and (= x47 y47) (= y47 x48)) (and (= x47 z47) (= z47 x48))) :named @p_28) (! (or (and (= x48 y48) (= y48 x49)) (and (= x48 z48) (= z48 x49))) :named @p_27) (! (or (and (= x49 y49) (= y49 x50)) (and (= x49 z49) (= z49 x50))) :named @p_26) (! (or (and (= x50 y50) (= y50 x51)) (and (= x50 z50) (= z50 x51))) :named @p_25) (! (or (and (= x51 y51) (= y51 x52)) (and (= x51 z51) (= z51 x52))) :named @p_24) (! (or (and (= x52 y52) (= y52 x53)) (and (= x52 z52) (= z52 x53))) :named @p_23) (! (or (and (= x53 y53) (= y53 x54)) (and (= x53 z53) (= z53 x54))) :named @p_22) (! (or (and (= x54 y54) (= y54 x55)) (and (= x54 z54) (= z54 x55))) :named @p_21) (! (or (and (= x55 y55) (= y55 x56)) (and (= x55 z55) (= z55 x56))) :named @p_20) (! (or (and (= x56 y56) (= y56 x57)) (and (= x56 z56) (= z56 x57))) :named @p_19) (! (or (and (= x57 y57) (= y57 x58)) (and (= x57 z57) (= z57 x58))) :named @p_18) (! (or (and (= x58 y58) (= y58 x59)) (and (= x58 z58) (= z58 x59))) :named @p_17) (! (or (and (= x59 y59) (= y59 x60)) (and (= x59 z59) (= z59 x60))) :named @p_16) (! (or (and (= x60 y60) (= y60 x61)) (and (= x60 z60) (= z60 x61))) :named @p_15) (! (or (and (= x61 y61) (= y61 x62)) (and (= x61 z61) (= z61 x62))) :named @p_14) (! (or (and (= x62 y62) (= y62 x63)) (and (= x62 z62) (= z62 x63))) :named @p_13) (! (or (and (= x63 y63) (= y63 x64)) (and (= x63 z63) (= z63 x64))) :named @p_12) (! (or (and (= x64 y64) (= y64 x65)) (and (= x64 z64) (= z64 x65))) :named @p_11) (! (or (and (= x65 y65) (= y65 x66)) (and (= x65 z65) (= z65 x66))) :named @p_10) (! (or (and (= x66 y66) (= y66 x67)) (and (= x66 z66) (= z66 x67))) :named @p_9) (! (or (and (= x67 y67) (= y67 x68)) (and (= x67 z67) (= z67 x68))) :named @p_8) (! (or (and (= x68 y68) (= y68 x69)) (and (= x68 z68) (= z68 x69))) :named @p_7) (! (or (and (= x69 y69) (= y69 x70)) (and (= x69 z69) (= z69 x70))) :named @p_6) (! (or (and (= x70 y70) (= y70 x71)) (and (= x70 z70) (= z70 x71))) :named @p_5) (! (or (and (= x71 y71) (= y71 x72)) (and (= x71 z71) (= z71 x72))) :named @p_4) (! (or (and (= x72 y72) (= y72 x73)) (and (= x72 z72) (= z72 x73))) :named @p_3) (! (not (! (= x0 x73) :named @p_1)) :named @p_2)) :named @p_78))
(step t0 (cl @p_73) :rule and :premises (a0) :args (4))
(step t1 (cl @p_72 @p_71) :rule or :premises (t0))
(step t2 (cl @p_73 (! (not @p_71) :named @p_522)) :rule or_neg)
(step t3 (cl (not (! (= @p_73 (! (and @p_73 (! (=> @p_73 (! (= x4 x5) :named @p_81)) :named @p_518)) :named @p_519)) :named @p_520)) (! (not @p_73) :named @p_80) @p_519) :rule equiv_pos2)
(step t4 (cl (= (! (= @p_520 true) :named @p_521) @p_520)) :rule equiv_simplify)
(step t5 (cl (not @p_521) @p_520) :rule equiv1 :premises (t4))
(step t6 (cl @p_521) :rule hole)
(step t7 (cl @p_520) :rule resolution :premises (t5 t6))
(step t8 (cl @p_519) :rule resolution :premises (t3 t7 t0))
(step t9 (cl @p_518) :rule and :premises (t8) :args (1))
(step t10 (cl @p_80 @p_81) :rule implies :premises (t9))
(step t11 (cl @p_81 @p_80) :rule reordering :premises (t10))
(step t12 (cl (= (! (= @p_81 false) :named @p_83) (! (not @p_81) :named @p_82))) :rule equiv_simplify)
(step t13 (cl (not @p_83) @p_82) :rule equiv1 :premises (t12))
(step t14 (cl (not (! (= @p_74 (! (and @p_74 (! (=> @p_74 (! (= x3 x4) :named @p_512)) :named @p_514)) :named @p_515)) :named @p_516)) (! (not @p_74) :named @p_513) @p_515) :rule equiv_pos2)
(step t15 (cl (= (! (= @p_516 true) :named @p_517) @p_516)) :rule equiv_simplify)
(step t16 (cl (not @p_517) @p_516) :rule equiv1 :premises (t15))
(step t17 (cl @p_517) :rule hole)
(step t18 (cl @p_516) :rule resolution :premises (t16 t17))
(step t19 (cl @p_74) :rule and :premises (a0) :args (3))
(step t20 (cl @p_515) :rule resolution :premises (t14 t18 t19))
(step t21 (cl @p_514) :rule and :premises (t20) :args (1))
(step t22 (cl @p_513 @p_512) :rule implies :premises (t21))
(step t23 (cl @p_512 @p_513) :rule reordering :premises (t22))
(step t24 (cl @p_512) :rule resolution :premises (t23 t19))
(step t25 (cl (= x4 x3)) :rule symm :premises (t24))
(step t26 (cl (not (! (= @p_75 (! (and @p_75 (! (=> @p_75 (! (= x2 x3) :named @p_506)) :named @p_508)) :named @p_509)) :named @p_510)) (! (not @p_75) :named @p_507) @p_509) :rule equiv_pos2)
(step t27 (cl (= (! (= @p_510 true) :named @p_511) @p_510)) :rule equiv_simplify)
(step t28 (cl (not @p_511) @p_510) :rule equiv1 :premises (t27))
(step t29 (cl @p_511) :rule hole)
(step t30 (cl @p_510) :rule resolution :premises (t28 t29))
(step t31 (cl @p_75) :rule and :premises (a0) :args (2))
(step t32 (cl @p_509) :rule resolution :premises (t26 t30 t31))
(step t33 (cl @p_508) :rule and :premises (t32) :args (1))
(step t34 (cl @p_507 @p_506) :rule implies :premises (t33))
(step t35 (cl @p_506 @p_507) :rule reordering :premises (t34))
(step t36 (cl @p_506) :rule resolution :premises (t35 t31))
(step t37 (cl (= x3 x2)) :rule symm :premises (t36))
(step t38 (cl (not (! (= @p_76 (! (and @p_76 (! (=> @p_76 (! (= x1 x2) :named @p_500)) :named @p_502)) :named @p_503)) :named @p_504)) (! (not @p_76) :named @p_501) @p_503) :rule equiv_pos2)
(step t39 (cl (= (! (= @p_504 true) :named @p_505) @p_504)) :rule equiv_simplify)
(step t40 (cl (not @p_505) @p_504) :rule equiv1 :premises (t39))
(step t41 (cl @p_505) :rule hole)
(step t42 (cl @p_504) :rule resolution :premises (t40 t41))
(step t43 (cl @p_76) :rule and :premises (a0) :args (1))
(step t44 (cl @p_503) :rule resolution :premises (t38 t42 t43))
(step t45 (cl @p_502) :rule and :premises (t44) :args (1))
(step t46 (cl @p_501 @p_500) :rule implies :premises (t45))
(step t47 (cl @p_500 @p_501) :rule reordering :premises (t46))
(step t48 (cl @p_500) :rule resolution :premises (t47 t43))
(step t49 (cl (= x2 x1)) :rule symm :premises (t48))
(step t50 (cl (not (! (= @p_77 (! (and @p_77 (! (=> @p_77 (! (= x0 x1) :named @p_494)) :named @p_496)) :named @p_497)) :named @p_498)) (! (not @p_77) :named @p_495) @p_497) :rule equiv_pos2)
(step t51 (cl (= (! (= @p_498 true) :named @p_499) @p_498)) :rule equiv_simplify)
(step t52 (cl (not @p_499) @p_498) :rule equiv1 :premises (t51))
(step t53 (cl @p_499) :rule hole)
(step t54 (cl @p_498) :rule resolution :premises (t52 t53))
(step t55 (cl @p_77) :rule and :premises (a0) :args (0))
(step t56 (cl @p_497) :rule resolution :premises (t50 t54 t55))
(step t57 (cl @p_496) :rule and :premises (t56) :args (1))
(step t58 (cl @p_495 @p_494) :rule implies :premises (t57))
(step t59 (cl @p_494 @p_495) :rule reordering :premises (t58))
(step t60 (cl @p_494) :rule resolution :premises (t59 t55))
(step t61 (cl (= x1 x0)) :rule symm :premises (t60))
(step t62 (cl (= x4 x0)) :rule trans :premises (t25 t37 t49 t61))
(step t63 (cl (not (! (= @p_70 (! (and @p_70 (! (=> @p_70 (! (= x5 x6) :named @p_488)) :named @p_490)) :named @p_491)) :named @p_492)) (! (not @p_70) :named @p_489) @p_491) :rule equiv_pos2)
(step t64 (cl (= (! (= @p_492 true) :named @p_493) @p_492)) :rule equiv_simplify)
(step t65 (cl (not @p_493) @p_492) :rule equiv1 :premises (t64))
(step t66 (cl @p_493) :rule hole)
(step t67 (cl @p_492) :rule resolution :premises (t65 t66))
(step t68 (cl @p_70) :rule and :premises (a0) :args (5))
(step t69 (cl @p_491) :rule resolution :premises (t63 t67 t68))
(step t70 (cl @p_490) :rule and :premises (t69) :args (1))
(step t71 (cl @p_489 @p_488) :rule implies :premises (t70))
(step t72 (cl @p_488 @p_489) :rule reordering :premises (t71))
(step t73 (cl @p_488) :rule resolution :premises (t72 t68))
(step t74 (cl (not (! (= @p_69 (! (and @p_69 (! (=> @p_69 (! (= x6 x7) :named @p_482)) :named @p_484)) :named @p_485)) :named @p_486)) (! (not @p_69) :named @p_483) @p_485) :rule equiv_pos2)
(step t75 (cl (= (! (= @p_486 true) :named @p_487) @p_486)) :rule equiv_simplify)
(step t76 (cl (not @p_487) @p_486) :rule equiv1 :premises (t75))
(step t77 (cl @p_487) :rule hole)
(step t78 (cl @p_486) :rule resolution :premises (t76 t77))
(step t79 (cl @p_69) :rule and :premises (a0) :args (6))
(step t80 (cl @p_485) :rule resolution :premises (t74 t78 t79))
(step t81 (cl @p_484) :rule and :premises (t80) :args (1))
(step t82 (cl @p_483 @p_482) :rule implies :premises (t81))
(step t83 (cl @p_482 @p_483) :rule reordering :premises (t82))
(step t84 (cl @p_482) :rule resolution :premises (t83 t79))
(step t85 (cl (not (! (= @p_68 (! (and @p_68 (! (=> @p_68 (! (= x7 x8) :named @p_476)) :named @p_478)) :named @p_479)) :named @p_480)) (! (not @p_68) :named @p_477) @p_479) :rule equiv_pos2)
(step t86 (cl (= (! (= @p_480 true) :named @p_481) @p_480)) :rule equiv_simplify)
(step t87 (cl (not @p_481) @p_480) :rule equiv1 :premises (t86))
(step t88 (cl @p_481) :rule hole)
(step t89 (cl @p_480) :rule resolution :premises (t87 t88))
(step t90 (cl @p_68) :rule and :premises (a0) :args (7))
(step t91 (cl @p_479) :rule resolution :premises (t85 t89 t90))
(step t92 (cl @p_478) :rule and :premises (t91) :args (1))
(step t93 (cl @p_477 @p_476) :rule implies :premises (t92))
(step t94 (cl @p_476 @p_477) :rule reordering :premises (t93))
(step t95 (cl @p_476) :rule resolution :premises (t94 t90))
(step t96 (cl (not (! (= @p_67 (! (and @p_67 (! (=> @p_67 (! (= x8 x9) :named @p_470)) :named @p_472)) :named @p_473)) :named @p_474)) (! (not @p_67) :named @p_471) @p_473) :rule equiv_pos2)
(step t97 (cl (= (! (= @p_474 true) :named @p_475) @p_474)) :rule equiv_simplify)
(step t98 (cl (not @p_475) @p_474) :rule equiv1 :premises (t97))
(step t99 (cl @p_475) :rule hole)
(step t100 (cl @p_474) :rule resolution :premises (t98 t99))
(step t101 (cl @p_67) :rule and :premises (a0) :args (8))
(step t102 (cl @p_473) :rule resolution :premises (t96 t100 t101))
(step t103 (cl @p_472) :rule and :premises (t102) :args (1))
(step t104 (cl @p_471 @p_470) :rule implies :premises (t103))
(step t105 (cl @p_470 @p_471) :rule reordering :premises (t104))
(step t106 (cl @p_470) :rule resolution :premises (t105 t101))
(step t107 (cl (not (! (= @p_66 (! (and @p_66 (! (=> @p_66 (! (= x9 x10) :named @p_464)) :named @p_466)) :named @p_467)) :named @p_468)) (! (not @p_66) :named @p_465) @p_467) :rule equiv_pos2)
(step t108 (cl (= (! (= @p_468 true) :named @p_469) @p_468)) :rule equiv_simplify)
(step t109 (cl (not @p_469) @p_468) :rule equiv1 :premises (t108))
(step t110 (cl @p_469) :rule hole)
(step t111 (cl @p_468) :rule resolution :premises (t109 t110))
(step t112 (cl @p_66) :rule and :premises (a0) :args (9))
(step t113 (cl @p_467) :rule resolution :premises (t107 t111 t112))
(step t114 (cl @p_466) :rule and :premises (t113) :args (1))
(step t115 (cl @p_465 @p_464) :rule implies :premises (t114))
(step t116 (cl @p_464 @p_465) :rule reordering :premises (t115))
(step t117 (cl @p_464) :rule resolution :premises (t116 t112))
(step t118 (cl (not (! (= @p_65 (! (and @p_65 (! (=> @p_65 (! (= x10 x11) :named @p_458)) :named @p_460)) :named @p_461)) :named @p_462)) (! (not @p_65) :named @p_459) @p_461) :rule equiv_pos2)
(step t119 (cl (= (! (= @p_462 true) :named @p_463) @p_462)) :rule equiv_simplify)
(step t120 (cl (not @p_463) @p_462) :rule equiv1 :premises (t119))
(step t121 (cl @p_463) :rule hole)
(step t122 (cl @p_462) :rule resolution :premises (t120 t121))
(step t123 (cl @p_65) :rule and :premises (a0) :args (10))
(step t124 (cl @p_461) :rule resolution :premises (t118 t122 t123))
(step t125 (cl @p_460) :rule and :premises (t124) :args (1))
(step t126 (cl @p_459 @p_458) :rule implies :premises (t125))
(step t127 (cl @p_458 @p_459) :rule reordering :premises (t126))
(step t128 (cl @p_458) :rule resolution :premises (t127 t123))
(step t129 (cl (not (! (= @p_64 (! (and @p_64 (! (=> @p_64 (! (= x11 x12) :named @p_452)) :named @p_454)) :named @p_455)) :named @p_456)) (! (not @p_64) :named @p_453) @p_455) :rule equiv_pos2)
(step t130 (cl (= (! (= @p_456 true) :named @p_457) @p_456)) :rule equiv_simplify)
(step t131 (cl (not @p_457) @p_456) :rule equiv1 :premises (t130))
(step t132 (cl @p_457) :rule hole)
(step t133 (cl @p_456) :rule resolution :premises (t131 t132))
(step t134 (cl @p_64) :rule and :premises (a0) :args (11))
(step t135 (cl @p_455) :rule resolution :premises (t129 t133 t134))
(step t136 (cl @p_454) :rule and :premises (t135) :args (1))
(step t137 (cl @p_453 @p_452) :rule implies :premises (t136))
(step t138 (cl @p_452 @p_453) :rule reordering :premises (t137))
(step t139 (cl @p_452) :rule resolution :premises (t138 t134))
(step t140 (cl (not (! (= @p_63 (! (and @p_63 (! (=> @p_63 (! (= x12 x13) :named @p_446)) :named @p_448)) :named @p_449)) :named @p_450)) (! (not @p_63) :named @p_447) @p_449) :rule equiv_pos2)
(step t141 (cl (= (! (= @p_450 true) :named @p_451) @p_450)) :rule equiv_simplify)
(step t142 (cl (not @p_451) @p_450) :rule equiv1 :premises (t141))
(step t143 (cl @p_451) :rule hole)
(step t144 (cl @p_450) :rule resolution :premises (t142 t143))
(step t145 (cl @p_63) :rule and :premises (a0) :args (12))
(step t146 (cl @p_449) :rule resolution :premises (t140 t144 t145))
(step t147 (cl @p_448) :rule and :premises (t146) :args (1))
(step t148 (cl @p_447 @p_446) :rule implies :premises (t147))
(step t149 (cl @p_446 @p_447) :rule reordering :premises (t148))
(step t150 (cl @p_446) :rule resolution :premises (t149 t145))
(step t151 (cl (not (! (= @p_62 (! (and @p_62 (! (=> @p_62 (! (= x13 x14) :named @p_440)) :named @p_442)) :named @p_443)) :named @p_444)) (! (not @p_62) :named @p_441) @p_443) :rule equiv_pos2)
(step t152 (cl (= (! (= @p_444 true) :named @p_445) @p_444)) :rule equiv_simplify)
(step t153 (cl (not @p_445) @p_444) :rule equiv1 :premises (t152))
(step t154 (cl @p_445) :rule hole)
(step t155 (cl @p_444) :rule resolution :premises (t153 t154))
(step t156 (cl @p_62) :rule and :premises (a0) :args (13))
(step t157 (cl @p_443) :rule resolution :premises (t151 t155 t156))
(step t158 (cl @p_442) :rule and :premises (t157) :args (1))
(step t159 (cl @p_441 @p_440) :rule implies :premises (t158))
(step t160 (cl @p_440 @p_441) :rule reordering :premises (t159))
(step t161 (cl @p_440) :rule resolution :premises (t160 t156))
(step t162 (cl (not (! (= @p_61 (! (and @p_61 (! (=> @p_61 (! (= x14 x15) :named @p_434)) :named @p_436)) :named @p_437)) :named @p_438)) (! (not @p_61) :named @p_435) @p_437) :rule equiv_pos2)
(step t163 (cl (= (! (= @p_438 true) :named @p_439) @p_438)) :rule equiv_simplify)
(step t164 (cl (not @p_439) @p_438) :rule equiv1 :premises (t163))
(step t165 (cl @p_439) :rule hole)
(step t166 (cl @p_438) :rule resolution :premises (t164 t165))
(step t167 (cl @p_61) :rule and :premises (a0) :args (14))
(step t168 (cl @p_437) :rule resolution :premises (t162 t166 t167))
(step t169 (cl @p_436) :rule and :premises (t168) :args (1))
(step t170 (cl @p_435 @p_434) :rule implies :premises (t169))
(step t171 (cl @p_434 @p_435) :rule reordering :premises (t170))
(step t172 (cl @p_434) :rule resolution :premises (t171 t167))
(step t173 (cl (not (! (= @p_60 (! (and @p_60 (! (=> @p_60 (! (= x15 x16) :named @p_428)) :named @p_430)) :named @p_431)) :named @p_432)) (! (not @p_60) :named @p_429) @p_431) :rule equiv_pos2)
(step t174 (cl (= (! (= @p_432 true) :named @p_433) @p_432)) :rule equiv_simplify)
(step t175 (cl (not @p_433) @p_432) :rule equiv1 :premises (t174))
(step t176 (cl @p_433) :rule hole)
(step t177 (cl @p_432) :rule resolution :premises (t175 t176))
(step t178 (cl @p_60) :rule and :premises (a0) :args (15))
(step t179 (cl @p_431) :rule resolution :premises (t173 t177 t178))
(step t180 (cl @p_430) :rule and :premises (t179) :args (1))
(step t181 (cl @p_429 @p_428) :rule implies :premises (t180))
(step t182 (cl @p_428 @p_429) :rule reordering :premises (t181))
(step t183 (cl @p_428) :rule resolution :premises (t182 t178))
(step t184 (cl (not (! (= @p_59 (! (and @p_59 (! (=> @p_59 (! (= x16 x17) :named @p_422)) :named @p_424)) :named @p_425)) :named @p_426)) (! (not @p_59) :named @p_423) @p_425) :rule equiv_pos2)
(step t185 (cl (= (! (= @p_426 true) :named @p_427) @p_426)) :rule equiv_simplify)
(step t186 (cl (not @p_427) @p_426) :rule equiv1 :premises (t185))
(step t187 (cl @p_427) :rule hole)
(step t188 (cl @p_426) :rule resolution :premises (t186 t187))
(step t189 (cl @p_59) :rule and :premises (a0) :args (16))
(step t190 (cl @p_425) :rule resolution :premises (t184 t188 t189))
(step t191 (cl @p_424) :rule and :premises (t190) :args (1))
(step t192 (cl @p_423 @p_422) :rule implies :premises (t191))
(step t193 (cl @p_422 @p_423) :rule reordering :premises (t192))
(step t194 (cl @p_422) :rule resolution :premises (t193 t189))
(step t195 (cl (not (! (= @p_58 (! (and @p_58 (! (=> @p_58 (! (= x17 x18) :named @p_416)) :named @p_418)) :named @p_419)) :named @p_420)) (! (not @p_58) :named @p_417) @p_419) :rule equiv_pos2)
(step t196 (cl (= (! (= @p_420 true) :named @p_421) @p_420)) :rule equiv_simplify)
(step t197 (cl (not @p_421) @p_420) :rule equiv1 :premises (t196))
(step t198 (cl @p_421) :rule hole)
(step t199 (cl @p_420) :rule resolution :premises (t197 t198))
(step t200 (cl @p_58) :rule and :premises (a0) :args (17))
(step t201 (cl @p_419) :rule resolution :premises (t195 t199 t200))
(step t202 (cl @p_418) :rule and :premises (t201) :args (1))
(step t203 (cl @p_417 @p_416) :rule implies :premises (t202))
(step t204 (cl @p_416 @p_417) :rule reordering :premises (t203))
(step t205 (cl @p_416) :rule resolution :premises (t204 t200))
(step t206 (cl (not (! (= @p_57 (! (and @p_57 (! (=> @p_57 (! (= x18 x19) :named @p_410)) :named @p_412)) :named @p_413)) :named @p_414)) (! (not @p_57) :named @p_411) @p_413) :rule equiv_pos2)
(step t207 (cl (= (! (= @p_414 true) :named @p_415) @p_414)) :rule equiv_simplify)
(step t208 (cl (not @p_415) @p_414) :rule equiv1 :premises (t207))
(step t209 (cl @p_415) :rule hole)
(step t210 (cl @p_414) :rule resolution :premises (t208 t209))
(step t211 (cl @p_57) :rule and :premises (a0) :args (18))
(step t212 (cl @p_413) :rule resolution :premises (t206 t210 t211))
(step t213 (cl @p_412) :rule and :premises (t212) :args (1))
(step t214 (cl @p_411 @p_410) :rule implies :premises (t213))
(step t215 (cl @p_410 @p_411) :rule reordering :premises (t214))
(step t216 (cl @p_410) :rule resolution :premises (t215 t211))
(step t217 (cl (not (! (= @p_56 (! (and @p_56 (! (=> @p_56 (! (= x19 x20) :named @p_404)) :named @p_406)) :named @p_407)) :named @p_408)) (! (not @p_56) :named @p_405) @p_407) :rule equiv_pos2)
(step t218 (cl (= (! (= @p_408 true) :named @p_409) @p_408)) :rule equiv_simplify)
(step t219 (cl (not @p_409) @p_408) :rule equiv1 :premises (t218))
(step t220 (cl @p_409) :rule hole)
(step t221 (cl @p_408) :rule resolution :premises (t219 t220))
(step t222 (cl @p_56) :rule and :premises (a0) :args (19))
(step t223 (cl @p_407) :rule resolution :premises (t217 t221 t222))
(step t224 (cl @p_406) :rule and :premises (t223) :args (1))
(step t225 (cl @p_405 @p_404) :rule implies :premises (t224))
(step t226 (cl @p_404 @p_405) :rule reordering :premises (t225))
(step t227 (cl @p_404) :rule resolution :premises (t226 t222))
(step t228 (cl (not (! (= @p_55 (! (and @p_55 (! (=> @p_55 (! (= x20 x21) :named @p_398)) :named @p_400)) :named @p_401)) :named @p_402)) (! (not @p_55) :named @p_399) @p_401) :rule equiv_pos2)
(step t229 (cl (= (! (= @p_402 true) :named @p_403) @p_402)) :rule equiv_simplify)
(step t230 (cl (not @p_403) @p_402) :rule equiv1 :premises (t229))
(step t231 (cl @p_403) :rule hole)
(step t232 (cl @p_402) :rule resolution :premises (t230 t231))
(step t233 (cl @p_55) :rule and :premises (a0) :args (20))
(step t234 (cl @p_401) :rule resolution :premises (t228 t232 t233))
(step t235 (cl @p_400) :rule and :premises (t234) :args (1))
(step t236 (cl @p_399 @p_398) :rule implies :premises (t235))
(step t237 (cl @p_398 @p_399) :rule reordering :premises (t236))
(step t238 (cl @p_398) :rule resolution :premises (t237 t233))
(step t239 (cl (not (! (= @p_54 (! (and @p_54 (! (=> @p_54 (! (= x21 x22) :named @p_392)) :named @p_394)) :named @p_395)) :named @p_396)) (! (not @p_54) :named @p_393) @p_395) :rule equiv_pos2)
(step t240 (cl (= (! (= @p_396 true) :named @p_397) @p_396)) :rule equiv_simplify)
(step t241 (cl (not @p_397) @p_396) :rule equiv1 :premises (t240))
(step t242 (cl @p_397) :rule hole)
(step t243 (cl @p_396) :rule resolution :premises (t241 t242))
(step t244 (cl @p_54) :rule and :premises (a0) :args (21))
(step t245 (cl @p_395) :rule resolution :premises (t239 t243 t244))
(step t246 (cl @p_394) :rule and :premises (t245) :args (1))
(step t247 (cl @p_393 @p_392) :rule implies :premises (t246))
(step t248 (cl @p_392 @p_393) :rule reordering :premises (t247))
(step t249 (cl @p_392) :rule resolution :premises (t248 t244))
(step t250 (cl (not (! (= @p_53 (! (and @p_53 (! (=> @p_53 (! (= x22 x23) :named @p_386)) :named @p_388)) :named @p_389)) :named @p_390)) (! (not @p_53) :named @p_387) @p_389) :rule equiv_pos2)
(step t251 (cl (= (! (= @p_390 true) :named @p_391) @p_390)) :rule equiv_simplify)
(step t252 (cl (not @p_391) @p_390) :rule equiv1 :premises (t251))
(step t253 (cl @p_391) :rule hole)
(step t254 (cl @p_390) :rule resolution :premises (t252 t253))
(step t255 (cl @p_53) :rule and :premises (a0) :args (22))
(step t256 (cl @p_389) :rule resolution :premises (t250 t254 t255))
(step t257 (cl @p_388) :rule and :premises (t256) :args (1))
(step t258 (cl @p_387 @p_386) :rule implies :premises (t257))
(step t259 (cl @p_386 @p_387) :rule reordering :premises (t258))
(step t260 (cl @p_386) :rule resolution :premises (t259 t255))
(step t261 (cl (not (! (= @p_52 (! (and @p_52 (! (=> @p_52 (! (= x23 x24) :named @p_380)) :named @p_382)) :named @p_383)) :named @p_384)) (! (not @p_52) :named @p_381) @p_383) :rule equiv_pos2)
(step t262 (cl (= (! (= @p_384 true) :named @p_385) @p_384)) :rule equiv_simplify)
(step t263 (cl (not @p_385) @p_384) :rule equiv1 :premises (t262))
(step t264 (cl @p_385) :rule hole)
(step t265 (cl @p_384) :rule resolution :premises (t263 t264))
(step t266 (cl @p_52) :rule and :premises (a0) :args (23))
(step t267 (cl @p_383) :rule resolution :premises (t261 t265 t266))
(step t268 (cl @p_382) :rule and :premises (t267) :args (1))
(step t269 (cl @p_381 @p_380) :rule implies :premises (t268))
(step t270 (cl @p_380 @p_381) :rule reordering :premises (t269))
(step t271 (cl @p_380) :rule resolution :premises (t270 t266))
(step t272 (cl (not (! (= @p_51 (! (and @p_51 (! (=> @p_51 (! (= x24 x25) :named @p_374)) :named @p_376)) :named @p_377)) :named @p_378)) (! (not @p_51) :named @p_375) @p_377) :rule equiv_pos2)
(step t273 (cl (= (! (= @p_378 true) :named @p_379) @p_378)) :rule equiv_simplify)
(step t274 (cl (not @p_379) @p_378) :rule equiv1 :premises (t273))
(step t275 (cl @p_379) :rule hole)
(step t276 (cl @p_378) :rule resolution :premises (t274 t275))
(step t277 (cl @p_51) :rule and :premises (a0) :args (24))
(step t278 (cl @p_377) :rule resolution :premises (t272 t276 t277))
(step t279 (cl @p_376) :rule and :premises (t278) :args (1))
(step t280 (cl @p_375 @p_374) :rule implies :premises (t279))
(step t281 (cl @p_374 @p_375) :rule reordering :premises (t280))
(step t282 (cl @p_374) :rule resolution :premises (t281 t277))
(step t283 (cl (not (! (= @p_50 (! (and @p_50 (! (=> @p_50 (! (= x25 x26) :named @p_368)) :named @p_370)) :named @p_371)) :named @p_372)) (! (not @p_50) :named @p_369) @p_371) :rule equiv_pos2)
(step t284 (cl (= (! (= @p_372 true) :named @p_373) @p_372)) :rule equiv_simplify)
(step t285 (cl (not @p_373) @p_372) :rule equiv1 :premises (t284))
(step t286 (cl @p_373) :rule hole)
(step t287 (cl @p_372) :rule resolution :premises (t285 t286))
(step t288 (cl @p_50) :rule and :premises (a0) :args (25))
(step t289 (cl @p_371) :rule resolution :premises (t283 t287 t288))
(step t290 (cl @p_370) :rule and :premises (t289) :args (1))
(step t291 (cl @p_369 @p_368) :rule implies :premises (t290))
(step t292 (cl @p_368 @p_369) :rule reordering :premises (t291))
(step t293 (cl @p_368) :rule resolution :premises (t292 t288))
(step t294 (cl (not (! (= @p_49 (! (and @p_49 (! (=> @p_49 (! (= x26 x27) :named @p_362)) :named @p_364)) :named @p_365)) :named @p_366)) (! (not @p_49) :named @p_363) @p_365) :rule equiv_pos2)
(step t295 (cl (= (! (= @p_366 true) :named @p_367) @p_366)) :rule equiv_simplify)
(step t296 (cl (not @p_367) @p_366) :rule equiv1 :premises (t295))
(step t297 (cl @p_367) :rule hole)
(step t298 (cl @p_366) :rule resolution :premises (t296 t297))
(step t299 (cl @p_49) :rule and :premises (a0) :args (26))
(step t300 (cl @p_365) :rule resolution :premises (t294 t298 t299))
(step t301 (cl @p_364) :rule and :premises (t300) :args (1))
(step t302 (cl @p_363 @p_362) :rule implies :premises (t301))
(step t303 (cl @p_362 @p_363) :rule reordering :premises (t302))
(step t304 (cl @p_362) :rule resolution :premises (t303 t299))
(step t305 (cl (not (! (= @p_48 (! (and @p_48 (! (=> @p_48 (! (= x27 x28) :named @p_356)) :named @p_358)) :named @p_359)) :named @p_360)) (! (not @p_48) :named @p_357) @p_359) :rule equiv_pos2)
(step t306 (cl (= (! (= @p_360 true) :named @p_361) @p_360)) :rule equiv_simplify)
(step t307 (cl (not @p_361) @p_360) :rule equiv1 :premises (t306))
(step t308 (cl @p_361) :rule hole)
(step t309 (cl @p_360) :rule resolution :premises (t307 t308))
(step t310 (cl @p_48) :rule and :premises (a0) :args (27))
(step t311 (cl @p_359) :rule resolution :premises (t305 t309 t310))
(step t312 (cl @p_358) :rule and :premises (t311) :args (1))
(step t313 (cl @p_357 @p_356) :rule implies :premises (t312))
(step t314 (cl @p_356 @p_357) :rule reordering :premises (t313))
(step t315 (cl @p_356) :rule resolution :premises (t314 t310))
(step t316 (cl (not (! (= @p_47 (! (and @p_47 (! (=> @p_47 (! (= x28 x29) :named @p_350)) :named @p_352)) :named @p_353)) :named @p_354)) (! (not @p_47) :named @p_351) @p_353) :rule equiv_pos2)
(step t317 (cl (= (! (= @p_354 true) :named @p_355) @p_354)) :rule equiv_simplify)
(step t318 (cl (not @p_355) @p_354) :rule equiv1 :premises (t317))
(step t319 (cl @p_355) :rule hole)
(step t320 (cl @p_354) :rule resolution :premises (t318 t319))
(step t321 (cl @p_47) :rule and :premises (a0) :args (28))
(step t322 (cl @p_353) :rule resolution :premises (t316 t320 t321))
(step t323 (cl @p_352) :rule and :premises (t322) :args (1))
(step t324 (cl @p_351 @p_350) :rule implies :premises (t323))
(step t325 (cl @p_350 @p_351) :rule reordering :premises (t324))
(step t326 (cl @p_350) :rule resolution :premises (t325 t321))
(step t327 (cl (not (! (= @p_46 (! (and @p_46 (! (=> @p_46 (! (= x29 x30) :named @p_344)) :named @p_346)) :named @p_347)) :named @p_348)) (! (not @p_46) :named @p_345) @p_347) :rule equiv_pos2)
(step t328 (cl (= (! (= @p_348 true) :named @p_349) @p_348)) :rule equiv_simplify)
(step t329 (cl (not @p_349) @p_348) :rule equiv1 :premises (t328))
(step t330 (cl @p_349) :rule hole)
(step t331 (cl @p_348) :rule resolution :premises (t329 t330))
(step t332 (cl @p_46) :rule and :premises (a0) :args (29))
(step t333 (cl @p_347) :rule resolution :premises (t327 t331 t332))
(step t334 (cl @p_346) :rule and :premises (t333) :args (1))
(step t335 (cl @p_345 @p_344) :rule implies :premises (t334))
(step t336 (cl @p_344 @p_345) :rule reordering :premises (t335))
(step t337 (cl @p_344) :rule resolution :premises (t336 t332))
(step t338 (cl (not (! (= @p_45 (! (and @p_45 (! (=> @p_45 (! (= x30 x31) :named @p_338)) :named @p_340)) :named @p_341)) :named @p_342)) (! (not @p_45) :named @p_339) @p_341) :rule equiv_pos2)
(step t339 (cl (= (! (= @p_342 true) :named @p_343) @p_342)) :rule equiv_simplify)
(step t340 (cl (not @p_343) @p_342) :rule equiv1 :premises (t339))
(step t341 (cl @p_343) :rule hole)
(step t342 (cl @p_342) :rule resolution :premises (t340 t341))
(step t343 (cl @p_45) :rule and :premises (a0) :args (30))
(step t344 (cl @p_341) :rule resolution :premises (t338 t342 t343))
(step t345 (cl @p_340) :rule and :premises (t344) :args (1))
(step t346 (cl @p_339 @p_338) :rule implies :premises (t345))
(step t347 (cl @p_338 @p_339) :rule reordering :premises (t346))
(step t348 (cl @p_338) :rule resolution :premises (t347 t343))
(step t349 (cl (not (! (= @p_44 (! (and @p_44 (! (=> @p_44 (! (= x31 x32) :named @p_332)) :named @p_334)) :named @p_335)) :named @p_336)) (! (not @p_44) :named @p_333) @p_335) :rule equiv_pos2)
(step t350 (cl (= (! (= @p_336 true) :named @p_337) @p_336)) :rule equiv_simplify)
(step t351 (cl (not @p_337) @p_336) :rule equiv1 :premises (t350))
(step t352 (cl @p_337) :rule hole)
(step t353 (cl @p_336) :rule resolution :premises (t351 t352))
(step t354 (cl @p_44) :rule and :premises (a0) :args (31))
(step t355 (cl @p_335) :rule resolution :premises (t349 t353 t354))
(step t356 (cl @p_334) :rule and :premises (t355) :args (1))
(step t357 (cl @p_333 @p_332) :rule implies :premises (t356))
(step t358 (cl @p_332 @p_333) :rule reordering :premises (t357))
(step t359 (cl @p_332) :rule resolution :premises (t358 t354))
(step t360 (cl (not (! (= @p_43 (! (and @p_43 (! (=> @p_43 (! (= x32 x33) :named @p_326)) :named @p_328)) :named @p_329)) :named @p_330)) (! (not @p_43) :named @p_327) @p_329) :rule equiv_pos2)
(step t361 (cl (= (! (= @p_330 true) :named @p_331) @p_330)) :rule equiv_simplify)
(step t362 (cl (not @p_331) @p_330) :rule equiv1 :premises (t361))
(step t363 (cl @p_331) :rule hole)
(step t364 (cl @p_330) :rule resolution :premises (t362 t363))
(step t365 (cl @p_43) :rule and :premises (a0) :args (32))
(step t366 (cl @p_329) :rule resolution :premises (t360 t364 t365))
(step t367 (cl @p_328) :rule and :premises (t366) :args (1))
(step t368 (cl @p_327 @p_326) :rule implies :premises (t367))
(step t369 (cl @p_326 @p_327) :rule reordering :premises (t368))
(step t370 (cl @p_326) :rule resolution :premises (t369 t365))
(step t371 (cl (not (! (= @p_42 (! (and @p_42 (! (=> @p_42 (! (= x33 x34) :named @p_320)) :named @p_322)) :named @p_323)) :named @p_324)) (! (not @p_42) :named @p_321) @p_323) :rule equiv_pos2)
(step t372 (cl (= (! (= @p_324 true) :named @p_325) @p_324)) :rule equiv_simplify)
(step t373 (cl (not @p_325) @p_324) :rule equiv1 :premises (t372))
(step t374 (cl @p_325) :rule hole)
(step t375 (cl @p_324) :rule resolution :premises (t373 t374))
(step t376 (cl @p_42) :rule and :premises (a0) :args (33))
(step t377 (cl @p_323) :rule resolution :premises (t371 t375 t376))
(step t378 (cl @p_322) :rule and :premises (t377) :args (1))
(step t379 (cl @p_321 @p_320) :rule implies :premises (t378))
(step t380 (cl @p_320 @p_321) :rule reordering :premises (t379))
(step t381 (cl @p_320) :rule resolution :premises (t380 t376))
(step t382 (cl (not (! (= @p_41 (! (and @p_41 (! (=> @p_41 (! (= x34 x35) :named @p_314)) :named @p_316)) :named @p_317)) :named @p_318)) (! (not @p_41) :named @p_315) @p_317) :rule equiv_pos2)
(step t383 (cl (= (! (= @p_318 true) :named @p_319) @p_318)) :rule equiv_simplify)
(step t384 (cl (not @p_319) @p_318) :rule equiv1 :premises (t383))
(step t385 (cl @p_319) :rule hole)
(step t386 (cl @p_318) :rule resolution :premises (t384 t385))
(step t387 (cl @p_41) :rule and :premises (a0) :args (34))
(step t388 (cl @p_317) :rule resolution :premises (t382 t386 t387))
(step t389 (cl @p_316) :rule and :premises (t388) :args (1))
(step t390 (cl @p_315 @p_314) :rule implies :premises (t389))
(step t391 (cl @p_314 @p_315) :rule reordering :premises (t390))
(step t392 (cl @p_314) :rule resolution :premises (t391 t387))
(step t393 (cl (not (! (= @p_40 (! (and @p_40 (! (=> @p_40 (! (= x35 x36) :named @p_308)) :named @p_310)) :named @p_311)) :named @p_312)) (! (not @p_40) :named @p_309) @p_311) :rule equiv_pos2)
(step t394 (cl (= (! (= @p_312 true) :named @p_313) @p_312)) :rule equiv_simplify)
(step t395 (cl (not @p_313) @p_312) :rule equiv1 :premises (t394))
(step t396 (cl @p_313) :rule hole)
(step t397 (cl @p_312) :rule resolution :premises (t395 t396))
(step t398 (cl @p_40) :rule and :premises (a0) :args (35))
(step t399 (cl @p_311) :rule resolution :premises (t393 t397 t398))
(step t400 (cl @p_310) :rule and :premises (t399) :args (1))
(step t401 (cl @p_309 @p_308) :rule implies :premises (t400))
(step t402 (cl @p_308 @p_309) :rule reordering :premises (t401))
(step t403 (cl @p_308) :rule resolution :premises (t402 t398))
(step t404 (cl (not (! (= @p_39 (! (and @p_39 (! (=> @p_39 (! (= x36 x37) :named @p_302)) :named @p_304)) :named @p_305)) :named @p_306)) (! (not @p_39) :named @p_303) @p_305) :rule equiv_pos2)
(step t405 (cl (= (! (= @p_306 true) :named @p_307) @p_306)) :rule equiv_simplify)
(step t406 (cl (not @p_307) @p_306) :rule equiv1 :premises (t405))
(step t407 (cl @p_307) :rule hole)
(step t408 (cl @p_306) :rule resolution :premises (t406 t407))
(step t409 (cl @p_39) :rule and :premises (a0) :args (36))
(step t410 (cl @p_305) :rule resolution :premises (t404 t408 t409))
(step t411 (cl @p_304) :rule and :premises (t410) :args (1))
(step t412 (cl @p_303 @p_302) :rule implies :premises (t411))
(step t413 (cl @p_302 @p_303) :rule reordering :premises (t412))
(step t414 (cl @p_302) :rule resolution :premises (t413 t409))
(step t415 (cl (not (! (= @p_38 (! (and @p_38 (! (=> @p_38 (! (= x37 x38) :named @p_296)) :named @p_298)) :named @p_299)) :named @p_300)) (! (not @p_38) :named @p_297) @p_299) :rule equiv_pos2)
(step t416 (cl (= (! (= @p_300 true) :named @p_301) @p_300)) :rule equiv_simplify)
(step t417 (cl (not @p_301) @p_300) :rule equiv1 :premises (t416))
(step t418 (cl @p_301) :rule hole)
(step t419 (cl @p_300) :rule resolution :premises (t417 t418))
(step t420 (cl @p_38) :rule and :premises (a0) :args (37))
(step t421 (cl @p_299) :rule resolution :premises (t415 t419 t420))
(step t422 (cl @p_298) :rule and :premises (t421) :args (1))
(step t423 (cl @p_297 @p_296) :rule implies :premises (t422))
(step t424 (cl @p_296 @p_297) :rule reordering :premises (t423))
(step t425 (cl @p_296) :rule resolution :premises (t424 t420))
(step t426 (cl (not (! (= @p_37 (! (and @p_37 (! (=> @p_37 (! (= x38 x39) :named @p_290)) :named @p_292)) :named @p_293)) :named @p_294)) (! (not @p_37) :named @p_291) @p_293) :rule equiv_pos2)
(step t427 (cl (= (! (= @p_294 true) :named @p_295) @p_294)) :rule equiv_simplify)
(step t428 (cl (not @p_295) @p_294) :rule equiv1 :premises (t427))
(step t429 (cl @p_295) :rule hole)
(step t430 (cl @p_294) :rule resolution :premises (t428 t429))
(step t431 (cl @p_37) :rule and :premises (a0) :args (38))
(step t432 (cl @p_293) :rule resolution :premises (t426 t430 t431))
(step t433 (cl @p_292) :rule and :premises (t432) :args (1))
(step t434 (cl @p_291 @p_290) :rule implies :premises (t433))
(step t435 (cl @p_290 @p_291) :rule reordering :premises (t434))
(step t436 (cl @p_290) :rule resolution :premises (t435 t431))
(step t437 (cl (not (! (= @p_36 (! (and @p_36 (! (=> @p_36 (! (= x39 x40) :named @p_284)) :named @p_286)) :named @p_287)) :named @p_288)) (! (not @p_36) :named @p_285) @p_287) :rule equiv_pos2)
(step t438 (cl (= (! (= @p_288 true) :named @p_289) @p_288)) :rule equiv_simplify)
(step t439 (cl (not @p_289) @p_288) :rule equiv1 :premises (t438))
(step t440 (cl @p_289) :rule hole)
(step t441 (cl @p_288) :rule resolution :premises (t439 t440))
(step t442 (cl @p_36) :rule and :premises (a0) :args (39))
(step t443 (cl @p_287) :rule resolution :premises (t437 t441 t442))
(step t444 (cl @p_286) :rule and :premises (t443) :args (1))
(step t445 (cl @p_285 @p_284) :rule implies :premises (t444))
(step t446 (cl @p_284 @p_285) :rule reordering :premises (t445))
(step t447 (cl @p_284) :rule resolution :premises (t446 t442))
(step t448 (cl (not (! (= @p_35 (! (and @p_35 (! (=> @p_35 (! (= x40 x41) :named @p_278)) :named @p_280)) :named @p_281)) :named @p_282)) (! (not @p_35) :named @p_279) @p_281) :rule equiv_pos2)
(step t449 (cl (= (! (= @p_282 true) :named @p_283) @p_282)) :rule equiv_simplify)
(step t450 (cl (not @p_283) @p_282) :rule equiv1 :premises (t449))
(step t451 (cl @p_283) :rule hole)
(step t452 (cl @p_282) :rule resolution :premises (t450 t451))
(step t453 (cl @p_35) :rule and :premises (a0) :args (40))
(step t454 (cl @p_281) :rule resolution :premises (t448 t452 t453))
(step t455 (cl @p_280) :rule and :premises (t454) :args (1))
(step t456 (cl @p_279 @p_278) :rule implies :premises (t455))
(step t457 (cl @p_278 @p_279) :rule reordering :premises (t456))
(step t458 (cl @p_278) :rule resolution :premises (t457 t453))
(step t459 (cl (not (! (= @p_34 (! (and @p_34 (! (=> @p_34 (! (= x41 x42) :named @p_272)) :named @p_274)) :named @p_275)) :named @p_276)) (! (not @p_34) :named @p_273) @p_275) :rule equiv_pos2)
(step t460 (cl (= (! (= @p_276 true) :named @p_277) @p_276)) :rule equiv_simplify)
(step t461 (cl (not @p_277) @p_276) :rule equiv1 :premises (t460))
(step t462 (cl @p_277) :rule hole)
(step t463 (cl @p_276) :rule resolution :premises (t461 t462))
(step t464 (cl @p_34) :rule and :premises (a0) :args (41))
(step t465 (cl @p_275) :rule resolution :premises (t459 t463 t464))
(step t466 (cl @p_274) :rule and :premises (t465) :args (1))
(step t467 (cl @p_273 @p_272) :rule implies :premises (t466))
(step t468 (cl @p_272 @p_273) :rule reordering :premises (t467))
(step t469 (cl @p_272) :rule resolution :premises (t468 t464))
(step t470 (cl (not (! (= @p_33 (! (and @p_33 (! (=> @p_33 (! (= x42 x43) :named @p_266)) :named @p_268)) :named @p_269)) :named @p_270)) (! (not @p_33) :named @p_267) @p_269) :rule equiv_pos2)
(step t471 (cl (= (! (= @p_270 true) :named @p_271) @p_270)) :rule equiv_simplify)
(step t472 (cl (not @p_271) @p_270) :rule equiv1 :premises (t471))
(step t473 (cl @p_271) :rule hole)
(step t474 (cl @p_270) :rule resolution :premises (t472 t473))
(step t475 (cl @p_33) :rule and :premises (a0) :args (42))
(step t476 (cl @p_269) :rule resolution :premises (t470 t474 t475))
(step t477 (cl @p_268) :rule and :premises (t476) :args (1))
(step t478 (cl @p_267 @p_266) :rule implies :premises (t477))
(step t479 (cl @p_266 @p_267) :rule reordering :premises (t478))
(step t480 (cl @p_266) :rule resolution :premises (t479 t475))
(step t481 (cl (not (! (= @p_32 (! (and @p_32 (! (=> @p_32 (! (= x43 x44) :named @p_260)) :named @p_262)) :named @p_263)) :named @p_264)) (! (not @p_32) :named @p_261) @p_263) :rule equiv_pos2)
(step t482 (cl (= (! (= @p_264 true) :named @p_265) @p_264)) :rule equiv_simplify)
(step t483 (cl (not @p_265) @p_264) :rule equiv1 :premises (t482))
(step t484 (cl @p_265) :rule hole)
(step t485 (cl @p_264) :rule resolution :premises (t483 t484))
(step t486 (cl @p_32) :rule and :premises (a0) :args (43))
(step t487 (cl @p_263) :rule resolution :premises (t481 t485 t486))
(step t488 (cl @p_262) :rule and :premises (t487) :args (1))
(step t489 (cl @p_261 @p_260) :rule implies :premises (t488))
(step t490 (cl @p_260 @p_261) :rule reordering :premises (t489))
(step t491 (cl @p_260) :rule resolution :premises (t490 t486))
(step t492 (cl (not (! (= @p_31 (! (and @p_31 (! (=> @p_31 (! (= x44 x45) :named @p_254)) :named @p_256)) :named @p_257)) :named @p_258)) (! (not @p_31) :named @p_255) @p_257) :rule equiv_pos2)
(step t493 (cl (= (! (= @p_258 true) :named @p_259) @p_258)) :rule equiv_simplify)
(step t494 (cl (not @p_259) @p_258) :rule equiv1 :premises (t493))
(step t495 (cl @p_259) :rule hole)
(step t496 (cl @p_258) :rule resolution :premises (t494 t495))
(step t497 (cl @p_31) :rule and :premises (a0) :args (44))
(step t498 (cl @p_257) :rule resolution :premises (t492 t496 t497))
(step t499 (cl @p_256) :rule and :premises (t498) :args (1))
(step t500 (cl @p_255 @p_254) :rule implies :premises (t499))
(step t501 (cl @p_254 @p_255) :rule reordering :premises (t500))
(step t502 (cl @p_254) :rule resolution :premises (t501 t497))
(step t503 (cl (not (! (= @p_30 (! (and @p_30 (! (=> @p_30 (! (= x45 x46) :named @p_248)) :named @p_250)) :named @p_251)) :named @p_252)) (! (not @p_30) :named @p_249) @p_251) :rule equiv_pos2)
(step t504 (cl (= (! (= @p_252 true) :named @p_253) @p_252)) :rule equiv_simplify)
(step t505 (cl (not @p_253) @p_252) :rule equiv1 :premises (t504))
(step t506 (cl @p_253) :rule hole)
(step t507 (cl @p_252) :rule resolution :premises (t505 t506))
(step t508 (cl @p_30) :rule and :premises (a0) :args (45))
(step t509 (cl @p_251) :rule resolution :premises (t503 t507 t508))
(step t510 (cl @p_250) :rule and :premises (t509) :args (1))
(step t511 (cl @p_249 @p_248) :rule implies :premises (t510))
(step t512 (cl @p_248 @p_249) :rule reordering :premises (t511))
(step t513 (cl @p_248) :rule resolution :premises (t512 t508))
(step t514 (cl (not (! (= @p_29 (! (and @p_29 (! (=> @p_29 (! (= x46 x47) :named @p_242)) :named @p_244)) :named @p_245)) :named @p_246)) (! (not @p_29) :named @p_243) @p_245) :rule equiv_pos2)
(step t515 (cl (= (! (= @p_246 true) :named @p_247) @p_246)) :rule equiv_simplify)
(step t516 (cl (not @p_247) @p_246) :rule equiv1 :premises (t515))
(step t517 (cl @p_247) :rule hole)
(step t518 (cl @p_246) :rule resolution :premises (t516 t517))
(step t519 (cl @p_29) :rule and :premises (a0) :args (46))
(step t520 (cl @p_245) :rule resolution :premises (t514 t518 t519))
(step t521 (cl @p_244) :rule and :premises (t520) :args (1))
(step t522 (cl @p_243 @p_242) :rule implies :premises (t521))
(step t523 (cl @p_242 @p_243) :rule reordering :premises (t522))
(step t524 (cl @p_242) :rule resolution :premises (t523 t519))
(step t525 (cl (not (! (= @p_28 (! (and @p_28 (! (=> @p_28 (! (= x47 x48) :named @p_236)) :named @p_238)) :named @p_239)) :named @p_240)) (! (not @p_28) :named @p_237) @p_239) :rule equiv_pos2)
(step t526 (cl (= (! (= @p_240 true) :named @p_241) @p_240)) :rule equiv_simplify)
(step t527 (cl (not @p_241) @p_240) :rule equiv1 :premises (t526))
(step t528 (cl @p_241) :rule hole)
(step t529 (cl @p_240) :rule resolution :premises (t527 t528))
(step t530 (cl @p_28) :rule and :premises (a0) :args (47))
(step t531 (cl @p_239) :rule resolution :premises (t525 t529 t530))
(step t532 (cl @p_238) :rule and :premises (t531) :args (1))
(step t533 (cl @p_237 @p_236) :rule implies :premises (t532))
(step t534 (cl @p_236 @p_237) :rule reordering :premises (t533))
(step t535 (cl @p_236) :rule resolution :premises (t534 t530))
(step t536 (cl (not (! (= @p_27 (! (and @p_27 (! (=> @p_27 (! (= x48 x49) :named @p_230)) :named @p_232)) :named @p_233)) :named @p_234)) (! (not @p_27) :named @p_231) @p_233) :rule equiv_pos2)
(step t537 (cl (= (! (= @p_234 true) :named @p_235) @p_234)) :rule equiv_simplify)
(step t538 (cl (not @p_235) @p_234) :rule equiv1 :premises (t537))
(step t539 (cl @p_235) :rule hole)
(step t540 (cl @p_234) :rule resolution :premises (t538 t539))
(step t541 (cl @p_27) :rule and :premises (a0) :args (48))
(step t542 (cl @p_233) :rule resolution :premises (t536 t540 t541))
(step t543 (cl @p_232) :rule and :premises (t542) :args (1))
(step t544 (cl @p_231 @p_230) :rule implies :premises (t543))
(step t545 (cl @p_230 @p_231) :rule reordering :premises (t544))
(step t546 (cl @p_230) :rule resolution :premises (t545 t541))
(step t547 (cl (not (! (= @p_26 (! (and @p_26 (! (=> @p_26 (! (= x49 x50) :named @p_224)) :named @p_226)) :named @p_227)) :named @p_228)) (! (not @p_26) :named @p_225) @p_227) :rule equiv_pos2)
(step t548 (cl (= (! (= @p_228 true) :named @p_229) @p_228)) :rule equiv_simplify)
(step t549 (cl (not @p_229) @p_228) :rule equiv1 :premises (t548))
(step t550 (cl @p_229) :rule hole)
(step t551 (cl @p_228) :rule resolution :premises (t549 t550))
(step t552 (cl @p_26) :rule and :premises (a0) :args (49))
(step t553 (cl @p_227) :rule resolution :premises (t547 t551 t552))
(step t554 (cl @p_226) :rule and :premises (t553) :args (1))
(step t555 (cl @p_225 @p_224) :rule implies :premises (t554))
(step t556 (cl @p_224 @p_225) :rule reordering :premises (t555))
(step t557 (cl @p_224) :rule resolution :premises (t556 t552))
(step t558 (cl (not (! (= @p_25 (! (and @p_25 (! (=> @p_25 (! (= x50 x51) :named @p_218)) :named @p_220)) :named @p_221)) :named @p_222)) (! (not @p_25) :named @p_219) @p_221) :rule equiv_pos2)
(step t559 (cl (= (! (= @p_222 true) :named @p_223) @p_222)) :rule equiv_simplify)
(step t560 (cl (not @p_223) @p_222) :rule equiv1 :premises (t559))
(step t561 (cl @p_223) :rule hole)
(step t562 (cl @p_222) :rule resolution :premises (t560 t561))
(step t563 (cl @p_25) :rule and :premises (a0) :args (50))
(step t564 (cl @p_221) :rule resolution :premises (t558 t562 t563))
(step t565 (cl @p_220) :rule and :premises (t564) :args (1))
(step t566 (cl @p_219 @p_218) :rule implies :premises (t565))
(step t567 (cl @p_218 @p_219) :rule reordering :premises (t566))
(step t568 (cl @p_218) :rule resolution :premises (t567 t563))
(step t569 (cl (not (! (= @p_24 (! (and @p_24 (! (=> @p_24 (! (= x51 x52) :named @p_212)) :named @p_214)) :named @p_215)) :named @p_216)) (! (not @p_24) :named @p_213) @p_215) :rule equiv_pos2)
(step t570 (cl (= (! (= @p_216 true) :named @p_217) @p_216)) :rule equiv_simplify)
(step t571 (cl (not @p_217) @p_216) :rule equiv1 :premises (t570))
(step t572 (cl @p_217) :rule hole)
(step t573 (cl @p_216) :rule resolution :premises (t571 t572))
(step t574 (cl @p_24) :rule and :premises (a0) :args (51))
(step t575 (cl @p_215) :rule resolution :premises (t569 t573 t574))
(step t576 (cl @p_214) :rule and :premises (t575) :args (1))
(step t577 (cl @p_213 @p_212) :rule implies :premises (t576))
(step t578 (cl @p_212 @p_213) :rule reordering :premises (t577))
(step t579 (cl @p_212) :rule resolution :premises (t578 t574))
(step t580 (cl (not (! (= @p_23 (! (and @p_23 (! (=> @p_23 (! (= x52 x53) :named @p_206)) :named @p_208)) :named @p_209)) :named @p_210)) (! (not @p_23) :named @p_207) @p_209) :rule equiv_pos2)
(step t581 (cl (= (! (= @p_210 true) :named @p_211) @p_210)) :rule equiv_simplify)
(step t582 (cl (not @p_211) @p_210) :rule equiv1 :premises (t581))
(step t583 (cl @p_211) :rule hole)
(step t584 (cl @p_210) :rule resolution :premises (t582 t583))
(step t585 (cl @p_23) :rule and :premises (a0) :args (52))
(step t586 (cl @p_209) :rule resolution :premises (t580 t584 t585))
(step t587 (cl @p_208) :rule and :premises (t586) :args (1))
(step t588 (cl @p_207 @p_206) :rule implies :premises (t587))
(step t589 (cl @p_206 @p_207) :rule reordering :premises (t588))
(step t590 (cl @p_206) :rule resolution :premises (t589 t585))
(step t591 (cl (not (! (= @p_22 (! (and @p_22 (! (=> @p_22 (! (= x53 x54) :named @p_200)) :named @p_202)) :named @p_203)) :named @p_204)) (! (not @p_22) :named @p_201) @p_203) :rule equiv_pos2)
(step t592 (cl (= (! (= @p_204 true) :named @p_205) @p_204)) :rule equiv_simplify)
(step t593 (cl (not @p_205) @p_204) :rule equiv1 :premises (t592))
(step t594 (cl @p_205) :rule hole)
(step t595 (cl @p_204) :rule resolution :premises (t593 t594))
(step t596 (cl @p_22) :rule and :premises (a0) :args (53))
(step t597 (cl @p_203) :rule resolution :premises (t591 t595 t596))
(step t598 (cl @p_202) :rule and :premises (t597) :args (1))
(step t599 (cl @p_201 @p_200) :rule implies :premises (t598))
(step t600 (cl @p_200 @p_201) :rule reordering :premises (t599))
(step t601 (cl @p_200) :rule resolution :premises (t600 t596))
(step t602 (cl (not (! (= @p_21 (! (and @p_21 (! (=> @p_21 (! (= x54 x55) :named @p_194)) :named @p_196)) :named @p_197)) :named @p_198)) (! (not @p_21) :named @p_195) @p_197) :rule equiv_pos2)
(step t603 (cl (= (! (= @p_198 true) :named @p_199) @p_198)) :rule equiv_simplify)
(step t604 (cl (not @p_199) @p_198) :rule equiv1 :premises (t603))
(step t605 (cl @p_199) :rule hole)
(step t606 (cl @p_198) :rule resolution :premises (t604 t605))
(step t607 (cl @p_21) :rule and :premises (a0) :args (54))
(step t608 (cl @p_197) :rule resolution :premises (t602 t606 t607))
(step t609 (cl @p_196) :rule and :premises (t608) :args (1))
(step t610 (cl @p_195 @p_194) :rule implies :premises (t609))
(step t611 (cl @p_194 @p_195) :rule reordering :premises (t610))
(step t612 (cl @p_194) :rule resolution :premises (t611 t607))
(step t613 (cl (not (! (= @p_20 (! (and @p_20 (! (=> @p_20 (! (= x55 x56) :named @p_188)) :named @p_190)) :named @p_191)) :named @p_192)) (! (not @p_20) :named @p_189) @p_191) :rule equiv_pos2)
(step t614 (cl (= (! (= @p_192 true) :named @p_193) @p_192)) :rule equiv_simplify)
(step t615 (cl (not @p_193) @p_192) :rule equiv1 :premises (t614))
(step t616 (cl @p_193) :rule hole)
(step t617 (cl @p_192) :rule resolution :premises (t615 t616))
(step t618 (cl @p_20) :rule and :premises (a0) :args (55))
(step t619 (cl @p_191) :rule resolution :premises (t613 t617 t618))
(step t620 (cl @p_190) :rule and :premises (t619) :args (1))
(step t621 (cl @p_189 @p_188) :rule implies :premises (t620))
(step t622 (cl @p_188 @p_189) :rule reordering :premises (t621))
(step t623 (cl @p_188) :rule resolution :premises (t622 t618))
(step t624 (cl (not (! (= @p_19 (! (and @p_19 (! (=> @p_19 (! (= x56 x57) :named @p_182)) :named @p_184)) :named @p_185)) :named @p_186)) (! (not @p_19) :named @p_183) @p_185) :rule equiv_pos2)
(step t625 (cl (= (! (= @p_186 true) :named @p_187) @p_186)) :rule equiv_simplify)
(step t626 (cl (not @p_187) @p_186) :rule equiv1 :premises (t625))
(step t627 (cl @p_187) :rule hole)
(step t628 (cl @p_186) :rule resolution :premises (t626 t627))
(step t629 (cl @p_19) :rule and :premises (a0) :args (56))
(step t630 (cl @p_185) :rule resolution :premises (t624 t628 t629))
(step t631 (cl @p_184) :rule and :premises (t630) :args (1))
(step t632 (cl @p_183 @p_182) :rule implies :premises (t631))
(step t633 (cl @p_182 @p_183) :rule reordering :premises (t632))
(step t634 (cl @p_182) :rule resolution :premises (t633 t629))
(step t635 (cl (not (! (= @p_18 (! (and @p_18 (! (=> @p_18 (! (= x57 x58) :named @p_176)) :named @p_178)) :named @p_179)) :named @p_180)) (! (not @p_18) :named @p_177) @p_179) :rule equiv_pos2)
(step t636 (cl (= (! (= @p_180 true) :named @p_181) @p_180)) :rule equiv_simplify)
(step t637 (cl (not @p_181) @p_180) :rule equiv1 :premises (t636))
(step t638 (cl @p_181) :rule hole)
(step t639 (cl @p_180) :rule resolution :premises (t637 t638))
(step t640 (cl @p_18) :rule and :premises (a0) :args (57))
(step t641 (cl @p_179) :rule resolution :premises (t635 t639 t640))
(step t642 (cl @p_178) :rule and :premises (t641) :args (1))
(step t643 (cl @p_177 @p_176) :rule implies :premises (t642))
(step t644 (cl @p_176 @p_177) :rule reordering :premises (t643))
(step t645 (cl @p_176) :rule resolution :premises (t644 t640))
(step t646 (cl (not (! (= @p_17 (! (and @p_17 (! (=> @p_17 (! (= x58 x59) :named @p_170)) :named @p_172)) :named @p_173)) :named @p_174)) (! (not @p_17) :named @p_171) @p_173) :rule equiv_pos2)
(step t647 (cl (= (! (= @p_174 true) :named @p_175) @p_174)) :rule equiv_simplify)
(step t648 (cl (not @p_175) @p_174) :rule equiv1 :premises (t647))
(step t649 (cl @p_175) :rule hole)
(step t650 (cl @p_174) :rule resolution :premises (t648 t649))
(step t651 (cl @p_17) :rule and :premises (a0) :args (58))
(step t652 (cl @p_173) :rule resolution :premises (t646 t650 t651))
(step t653 (cl @p_172) :rule and :premises (t652) :args (1))
(step t654 (cl @p_171 @p_170) :rule implies :premises (t653))
(step t655 (cl @p_170 @p_171) :rule reordering :premises (t654))
(step t656 (cl @p_170) :rule resolution :premises (t655 t651))
(step t657 (cl (not (! (= @p_16 (! (and @p_16 (! (=> @p_16 (! (= x59 x60) :named @p_164)) :named @p_166)) :named @p_167)) :named @p_168)) (! (not @p_16) :named @p_165) @p_167) :rule equiv_pos2)
(step t658 (cl (= (! (= @p_168 true) :named @p_169) @p_168)) :rule equiv_simplify)
(step t659 (cl (not @p_169) @p_168) :rule equiv1 :premises (t658))
(step t660 (cl @p_169) :rule hole)
(step t661 (cl @p_168) :rule resolution :premises (t659 t660))
(step t662 (cl @p_16) :rule and :premises (a0) :args (59))
(step t663 (cl @p_167) :rule resolution :premises (t657 t661 t662))
(step t664 (cl @p_166) :rule and :premises (t663) :args (1))
(step t665 (cl @p_165 @p_164) :rule implies :premises (t664))
(step t666 (cl @p_164 @p_165) :rule reordering :premises (t665))
(step t667 (cl @p_164) :rule resolution :premises (t666 t662))
(step t668 (cl (not (! (= @p_15 (! (and @p_15 (! (=> @p_15 (! (= x60 x61) :named @p_158)) :named @p_160)) :named @p_161)) :named @p_162)) (! (not @p_15) :named @p_159) @p_161) :rule equiv_pos2)
(step t669 (cl (= (! (= @p_162 true) :named @p_163) @p_162)) :rule equiv_simplify)
(step t670 (cl (not @p_163) @p_162) :rule equiv1 :premises (t669))
(step t671 (cl @p_163) :rule hole)
(step t672 (cl @p_162) :rule resolution :premises (t670 t671))
(step t673 (cl @p_15) :rule and :premises (a0) :args (60))
(step t674 (cl @p_161) :rule resolution :premises (t668 t672 t673))
(step t675 (cl @p_160) :rule and :premises (t674) :args (1))
(step t676 (cl @p_159 @p_158) :rule implies :premises (t675))
(step t677 (cl @p_158 @p_159) :rule reordering :premises (t676))
(step t678 (cl @p_158) :rule resolution :premises (t677 t673))
(step t679 (cl (not (! (= @p_14 (! (and @p_14 (! (=> @p_14 (! (= x61 x62) :named @p_152)) :named @p_154)) :named @p_155)) :named @p_156)) (! (not @p_14) :named @p_153) @p_155) :rule equiv_pos2)
(step t680 (cl (= (! (= @p_156 true) :named @p_157) @p_156)) :rule equiv_simplify)
(step t681 (cl (not @p_157) @p_156) :rule equiv1 :premises (t680))
(step t682 (cl @p_157) :rule hole)
(step t683 (cl @p_156) :rule resolution :premises (t681 t682))
(step t684 (cl @p_14) :rule and :premises (a0) :args (61))
(step t685 (cl @p_155) :rule resolution :premises (t679 t683 t684))
(step t686 (cl @p_154) :rule and :premises (t685) :args (1))
(step t687 (cl @p_153 @p_152) :rule implies :premises (t686))
(step t688 (cl @p_152 @p_153) :rule reordering :premises (t687))
(step t689 (cl @p_152) :rule resolution :premises (t688 t684))
(step t690 (cl (not (! (= @p_13 (! (and @p_13 (! (=> @p_13 (! (= x62 x63) :named @p_146)) :named @p_148)) :named @p_149)) :named @p_150)) (! (not @p_13) :named @p_147) @p_149) :rule equiv_pos2)
(step t691 (cl (= (! (= @p_150 true) :named @p_151) @p_150)) :rule equiv_simplify)
(step t692 (cl (not @p_151) @p_150) :rule equiv1 :premises (t691))
(step t693 (cl @p_151) :rule hole)
(step t694 (cl @p_150) :rule resolution :premises (t692 t693))
(step t695 (cl @p_13) :rule and :premises (a0) :args (62))
(step t696 (cl @p_149) :rule resolution :premises (t690 t694 t695))
(step t697 (cl @p_148) :rule and :premises (t696) :args (1))
(step t698 (cl @p_147 @p_146) :rule implies :premises (t697))
(step t699 (cl @p_146 @p_147) :rule reordering :premises (t698))
(step t700 (cl @p_146) :rule resolution :premises (t699 t695))
(step t701 (cl (not (! (= @p_12 (! (and @p_12 (! (=> @p_12 (! (= x63 x64) :named @p_140)) :named @p_142)) :named @p_143)) :named @p_144)) (! (not @p_12) :named @p_141) @p_143) :rule equiv_pos2)
(step t702 (cl (= (! (= @p_144 true) :named @p_145) @p_144)) :rule equiv_simplify)
(step t703 (cl (not @p_145) @p_144) :rule equiv1 :premises (t702))
(step t704 (cl @p_145) :rule hole)
(step t705 (cl @p_144) :rule resolution :premises (t703 t704))
(step t706 (cl @p_12) :rule and :premises (a0) :args (63))
(step t707 (cl @p_143) :rule resolution :premises (t701 t705 t706))
(step t708 (cl @p_142) :rule and :premises (t707) :args (1))
(step t709 (cl @p_141 @p_140) :rule implies :premises (t708))
(step t710 (cl @p_140 @p_141) :rule reordering :premises (t709))
(step t711 (cl @p_140) :rule resolution :premises (t710 t706))
(step t712 (cl (not (! (= @p_11 (! (and @p_11 (! (=> @p_11 (! (= x64 x65) :named @p_134)) :named @p_136)) :named @p_137)) :named @p_138)) (! (not @p_11) :named @p_135) @p_137) :rule equiv_pos2)
(step t713 (cl (= (! (= @p_138 true) :named @p_139) @p_138)) :rule equiv_simplify)
(step t714 (cl (not @p_139) @p_138) :rule equiv1 :premises (t713))
(step t715 (cl @p_139) :rule hole)
(step t716 (cl @p_138) :rule resolution :premises (t714 t715))
(step t717 (cl @p_11) :rule and :premises (a0) :args (64))
(step t718 (cl @p_137) :rule resolution :premises (t712 t716 t717))
(step t719 (cl @p_136) :rule and :premises (t718) :args (1))
(step t720 (cl @p_135 @p_134) :rule implies :premises (t719))
(step t721 (cl @p_134 @p_135) :rule reordering :premises (t720))
(step t722 (cl @p_134) :rule resolution :premises (t721 t717))
(step t723 (cl (not (! (= @p_10 (! (and @p_10 (! (=> @p_10 (! (= x65 x66) :named @p_128)) :named @p_130)) :named @p_131)) :named @p_132)) (! (not @p_10) :named @p_129) @p_131) :rule equiv_pos2)
(step t724 (cl (= (! (= @p_132 true) :named @p_133) @p_132)) :rule equiv_simplify)
(step t725 (cl (not @p_133) @p_132) :rule equiv1 :premises (t724))
(step t726 (cl @p_133) :rule hole)
(step t727 (cl @p_132) :rule resolution :premises (t725 t726))
(step t728 (cl @p_10) :rule and :premises (a0) :args (65))
(step t729 (cl @p_131) :rule resolution :premises (t723 t727 t728))
(step t730 (cl @p_130) :rule and :premises (t729) :args (1))
(step t731 (cl @p_129 @p_128) :rule implies :premises (t730))
(step t732 (cl @p_128 @p_129) :rule reordering :premises (t731))
(step t733 (cl @p_128) :rule resolution :premises (t732 t728))
(step t734 (cl (not (! (= @p_9 (! (and @p_9 (! (=> @p_9 (! (= x66 x67) :named @p_122)) :named @p_124)) :named @p_125)) :named @p_126)) (! (not @p_9) :named @p_123) @p_125) :rule equiv_pos2)
(step t735 (cl (= (! (= @p_126 true) :named @p_127) @p_126)) :rule equiv_simplify)
(step t736 (cl (not @p_127) @p_126) :rule equiv1 :premises (t735))
(step t737 (cl @p_127) :rule hole)
(step t738 (cl @p_126) :rule resolution :premises (t736 t737))
(step t739 (cl @p_9) :rule and :premises (a0) :args (66))
(step t740 (cl @p_125) :rule resolution :premises (t734 t738 t739))
(step t741 (cl @p_124) :rule and :premises (t740) :args (1))
(step t742 (cl @p_123 @p_122) :rule implies :premises (t741))
(step t743 (cl @p_122 @p_123) :rule reordering :premises (t742))
(step t744 (cl @p_122) :rule resolution :premises (t743 t739))
(step t745 (cl (not (! (= @p_8 (! (and @p_8 (! (=> @p_8 (! (= x67 x68) :named @p_116)) :named @p_118)) :named @p_119)) :named @p_120)) (! (not @p_8) :named @p_117) @p_119) :rule equiv_pos2)
(step t746 (cl (= (! (= @p_120 true) :named @p_121) @p_120)) :rule equiv_simplify)
(step t747 (cl (not @p_121) @p_120) :rule equiv1 :premises (t746))
(step t748 (cl @p_121) :rule hole)
(step t749 (cl @p_120) :rule resolution :premises (t747 t748))
(step t750 (cl @p_8) :rule and :premises (a0) :args (67))
(step t751 (cl @p_119) :rule resolution :premises (t745 t749 t750))
(step t752 (cl @p_118) :rule and :premises (t751) :args (1))
(step t753 (cl @p_117 @p_116) :rule implies :premises (t752))
(step t754 (cl @p_116 @p_117) :rule reordering :premises (t753))
(step t755 (cl @p_116) :rule resolution :premises (t754 t750))
(step t756 (cl (not (! (= @p_7 (! (and @p_7 (! (=> @p_7 (! (= x68 x69) :named @p_110)) :named @p_112)) :named @p_113)) :named @p_114)) (! (not @p_7) :named @p_111) @p_113) :rule equiv_pos2)
(step t757 (cl (= (! (= @p_114 true) :named @p_115) @p_114)) :rule equiv_simplify)
(step t758 (cl (not @p_115) @p_114) :rule equiv1 :premises (t757))
(step t759 (cl @p_115) :rule hole)
(step t760 (cl @p_114) :rule resolution :premises (t758 t759))
(step t761 (cl @p_7) :rule and :premises (a0) :args (68))
(step t762 (cl @p_113) :rule resolution :premises (t756 t760 t761))
(step t763 (cl @p_112) :rule and :premises (t762) :args (1))
(step t764 (cl @p_111 @p_110) :rule implies :premises (t763))
(step t765 (cl @p_110 @p_111) :rule reordering :premises (t764))
(step t766 (cl @p_110) :rule resolution :premises (t765 t761))
(step t767 (cl (not (! (= @p_6 (! (and @p_6 (! (=> @p_6 (! (= x69 x70) :named @p_104)) :named @p_106)) :named @p_107)) :named @p_108)) (! (not @p_6) :named @p_105) @p_107) :rule equiv_pos2)
(step t768 (cl (= (! (= @p_108 true) :named @p_109) @p_108)) :rule equiv_simplify)
(step t769 (cl (not @p_109) @p_108) :rule equiv1 :premises (t768))
(step t770 (cl @p_109) :rule hole)
(step t771 (cl @p_108) :rule resolution :premises (t769 t770))
(step t772 (cl @p_6) :rule and :premises (a0) :args (69))
(step t773 (cl @p_107) :rule resolution :premises (t767 t771 t772))
(step t774 (cl @p_106) :rule and :premises (t773) :args (1))
(step t775 (cl @p_105 @p_104) :rule implies :premises (t774))
(step t776 (cl @p_104 @p_105) :rule reordering :premises (t775))
(step t777 (cl @p_104) :rule resolution :premises (t776 t772))
(step t778 (cl (not (! (= @p_5 (! (and @p_5 (! (=> @p_5 (! (= x70 x71) :named @p_98)) :named @p_100)) :named @p_101)) :named @p_102)) (! (not @p_5) :named @p_99) @p_101) :rule equiv_pos2)
(step t779 (cl (= (! (= @p_102 true) :named @p_103) @p_102)) :rule equiv_simplify)
(step t780 (cl (not @p_103) @p_102) :rule equiv1 :premises (t779))
(step t781 (cl @p_103) :rule hole)
(step t782 (cl @p_102) :rule resolution :premises (t780 t781))
(step t783 (cl @p_5) :rule and :premises (a0) :args (70))
(step t784 (cl @p_101) :rule resolution :premises (t778 t782 t783))
(step t785 (cl @p_100) :rule and :premises (t784) :args (1))
(step t786 (cl @p_99 @p_98) :rule implies :premises (t785))
(step t787 (cl @p_98 @p_99) :rule reordering :premises (t786))
(step t788 (cl @p_98) :rule resolution :premises (t787 t783))
(step t789 (cl (not (! (= @p_4 (! (and @p_4 (! (=> @p_4 (! (= x71 x72) :named @p_92)) :named @p_94)) :named @p_95)) :named @p_96)) (! (not @p_4) :named @p_93) @p_95) :rule equiv_pos2)
(step t790 (cl (= (! (= @p_96 true) :named @p_97) @p_96)) :rule equiv_simplify)
(step t791 (cl (not @p_97) @p_96) :rule equiv1 :premises (t790))
(step t792 (cl @p_97) :rule hole)
(step t793 (cl @p_96) :rule resolution :premises (t791 t792))
(step t794 (cl @p_4) :rule and :premises (a0) :args (71))
(step t795 (cl @p_95) :rule resolution :premises (t789 t793 t794))
(step t796 (cl @p_94) :rule and :premises (t795) :args (1))
(step t797 (cl @p_93 @p_92) :rule implies :premises (t796))
(step t798 (cl @p_92 @p_93) :rule reordering :premises (t797))
(step t799 (cl @p_92) :rule resolution :premises (t798 t794))
(step t800 (cl (not (! (= @p_3 (! (and @p_3 (! (=> @p_3 (! (= x72 x73) :named @p_86)) :named @p_88)) :named @p_89)) :named @p_90)) (! (not @p_3) :named @p_87) @p_89) :rule equiv_pos2)
(step t801 (cl (= (! (= @p_90 true) :named @p_91) @p_90)) :rule equiv_simplify)
(step t802 (cl (not @p_91) @p_90) :rule equiv1 :premises (t801))
(step t803 (cl @p_91) :rule hole)
(step t804 (cl @p_90) :rule resolution :premises (t802 t803))
(step t805 (cl @p_3) :rule and :premises (a0) :args (72))
(step t806 (cl @p_89) :rule resolution :premises (t800 t804 t805))
(step t807 (cl @p_88) :rule and :premises (t806) :args (1))
(step t808 (cl @p_87 @p_86) :rule implies :premises (t807))
(step t809 (cl @p_86 @p_87) :rule reordering :premises (t808))
(step t810 (cl @p_86) :rule resolution :premises (t809 t805))
(step t811 (cl (= x5 x73)) :rule trans :premises (t73 t84 t95 t106 t117 t128 t139 t150 t161 t172 t183 t194 t205 t216 t227 t238 t249 t260 t271 t282 t293 t304 t315 t326 t337 t348 t359 t370 t381 t392 t403 t414 t425 t436 t447 t458 t469 t480 t491 t502 t513 t524 t535 t546 t557 t568 t579 t590 t601 t612 t623 t634 t645 t656 t667 t678 t689 t700 t711 t722 t733 t744 t755 t766 t777 t788 t799 t810))
(step t812 (cl (= @p_81 @p_1)) :rule cong :premises (t62 t811))
(step t813 (cl (= (! (= @p_1 false) :named @p_84) @p_2)) :rule equiv_simplify)
(step t814 (cl @p_84 (! (not @p_2) :named @p_85)) :rule equiv2 :premises (t813))
(step t815 (cl (not @p_85) @p_1) :rule not_not)
(step t816 (cl @p_84 @p_1) :rule resolution :premises (t814 t815))
(step t817 (cl @p_2) :rule and :premises (a0) :args (73))
(step t818 (cl @p_84) :rule resolution :premises (t816 t817))
(step t819 (cl @p_83) :rule trans :premises (t812 t818))
(step t820 (cl @p_82) :rule resolution :premises (t13 t819))
(step t821 (cl @p_80) :rule resolution :premises (t11 t820))
(step t822 (cl @p_522) :rule resolution :premises (t2 t821))
(step t823 (cl @p_73 (! (not @p_72) :named @p_79)) :rule or_neg)
(step t824 (cl @p_79) :rule resolution :premises (t823 t821))
(step t825 (cl) :rule resolution :premises (t1 t822 t824))